---
id: 67eaa957114d373deb3a9149
title: 构建一个货币转换器
challengeType: 25
dashedName: build-a-currency-converter
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你的 `CurrencyConverter` 组件应渲染一个 `input` 元素，用于输入要转换的金额。
2. 你的 `input` 元素应只接受数字。
3. 你的 `CurrencyConverter` 组件应渲染两个 `select` 元素，用于选择要转换的货币（**从**和**到**）。
4. 你的 `select` 元素应至少包含 `USD`、`EUR`、`GBP` 和 `JPY` 选项。
  1. 你可以使用任意汇率，只要不同货币之间不是一对一映射即可。
5. 你的 `CurrencyConverter` 组件应对“从”货币的转换结果进行记忆化处理，使得仅当“从”货币发生变化时才重新计算转换结果，而“到”货币变化时不会重新计算。
6. 你的 `CurrencyConverter` 组件应渲染一个元素，显示格式为 `XX.XX CCC` 的转换金额，其中 `XX.XX` 是转换后的金额，`CCC` 是货币代码。
7. 转换金额应四舍五入保留两位小数。

# --hints--

你应导出一个名为 `CurrencyConverter` 的组件。

```js
async () => {
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);

  assert.property(exports, "CurrencyConverter");
}
```

你应有一个 `input[type="number"]` 元素用于输入要转换的金额。

```js
const inp = document.querySelectorAll('input[type="number"]');
assert.equal(inp.length, 1);
```

你应有两个 `select` 元素。

```js
const selects = document.querySelectorAll('select');
assert.equal(selects.length, 2);
```

`select` 元素应至少包含 `USD`、`EUR`、`GBP` 和 `JPY` 选项。

```js
const selects = [...document.querySelectorAll('select')];

assert.equal(selects.length, 2);

for (const select of selects) {
  const options = [...select.options].map((o) => o.value);
  assert.includeMembers(options, ["USD", "EUR", "GBP", "JPY"]);
}
```

更改第一个 `select` 元素的值应导致转换结果被重新计算。

```js
async () => {
  function spyOn(
   obj,
   method
  ) {
   const original = obj[method];
   const calls = [];

   const fn = (cb, deps) => {
    const result = original(() => {calls.push(1); return cb();}, deps);
    return result;
   };

   obj[method] = fn;

   fn.calls = calls;
   return fn;
  }
  const abuseMemo = spyOn(React, "useMemo");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);
  const s = await __helpers.prepTestComponent(exports.CurrencyConverter);
  const [first, _second] = s.querySelectorAll('select');
  assert.exists(first);
  await React.act(async () => {
   // 找到第一个未被选中的选项
   const notSelected = [...first.options].find((o) => !o.selected);
   first.value = notSelected.value;
   const ev = new Event("change", { bubbles: true, cancelable: false });
   first[Object.keys(first).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: first});
  });
  assert.equal(abuseMemo.calls.length, 2);
}
```

更改第二个 `select` 元素的值不应导致转换结果被重新计算。

```js
async () => {
  function spyOn(
   obj,
   method
  ) {
   const original = obj[method];
   const calls = [];

   const fn = (cb, deps) => {
    const result = original(() => {calls.push(1); return cb();}, deps);
    return result;
   };

   obj[method] = fn;

   fn.calls = calls;
   return fn;
  }
  const abuseMemo = spyOn(React, "useMemo");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);
  const s = await __helpers.prepTestComponent(exports.CurrencyConverter);
  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await React.act(async () => {
   // 找到第一个未被选中的选项
   const notSelected = [...second.options].find((o) => !o.selected);
   second.value = notSelected.value;
   const ev = new Event("change", { bubbles: true, cancelable: false });
   second[Object.keys(second).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: second});
  });
  assert.equal(abuseMemo.calls.length, 1);
}
```

更改第一个 `select` 元素的值应导致页面上的文本发生变化。

```js
async () => {
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const nonFormContentBefore = getInnerTextExcept(s, "input,select");

  const [first, _second] = s.querySelectorAll('select');
  assert.exists(first);
  await React.act(async () => {
   // 找到第一个未被选中的选项
   const notSelected = [...first.options].find((o) => !o.selected);
   first.value = notSelected.value;
   const ev = new Event("change", { bubbles: true, cancelable: false });
   first[Object.keys(first).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: first});
  });

  const nonFormContentAfter = getInnerTextExcept(s, "input,select");

  try {
   assert.notEqual(nonFormContentBefore, nonFormContentAfter);
  } catch (e) {
   console.error(e);
   throw e;
  }
}
```

更改第二个 `select` 元素的值应导致页面上的文本发生变化。

```js
async () => {
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const nonFormContentBefore = getInnerTextExcept(s, "input,select");

  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await React.act(async () => {
   // 找到第一个未被选中的选项
   const notSelected = [...second.options].find((o) => !o.selected);
   second.value = notSelected.value;
   const ev = new Event("change", { bubbles: true, cancelable: false });
   second[Object.keys(second).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: second});
  });

  const nonFormContentAfter = getInnerTextExcept(s, "input,select");

  try {
   assert.notEqual(nonFormContentBefore, nonFormContentAfter);
  } catch (e) {
   console.error(e);
   throw e;
  }
}
```

转换金额应以 `XX.XX CCC` 的格式显示，其中 `XX.XX` 是保留两位小数的金额，`CCC` 是货币代码。

```js
async () => {
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const inp = s.querySelector('input[type="number"]');
  assert.exists(inp);
  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await React.act(async () => {
   // 找到第一个未被选中的选项
   const notSelected = [...second.options].find((o) => !o.selected);
   second.value = notSelected.value;
   const ev = new Event("change", { bubbles: true, cancelable: false });
   second[Object.keys(second).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: second});
   inp.value = 10;
   const ev2 = new Event("change", { bubbles: true, cancelable: false });
   inp[Object.keys(inp).find((k) => k.startsWith("__reactProps"))].onChange({...ev2, target: inp});
  });

  const nonFormContent = getInnerTextExcept(s, "input,select");

  try {
   const currencyCode = second.value;
   const reg = new RegExp(`\\d+\\.\\d{2} ${currencyCode}`);
   assert.match(nonFormContent, reg);
  } catch (e) {
   console.error(e);
   throw e;
  }
}
```

转换后的金额应与输入金额不同。

```js
async () => {
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const inp = s.querySelector('input[type="number"]');
  assert.exists(inp);
  const [first, second] = s.querySelectorAll('select');
  assert.exists(first);
  assert.exists(second);
  for (let i = 0; i < first.options.length; i++) {
   for (let j = 0; j < second.options.length; j++) {
    if (first.options[i].value === second.options[j].value) {
      continue;
    }

    await React.act(async () => {
      first.value = first.options[i].value;
      const ev = new Event("change", { bubbles: true, cancelable: false });
      first[Object.keys(first).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: first});
      second.value = second.options[j].value;
      const ev2 = new Event("change", { bubbles: true, cancelable: false });
      second[Object.keys(second).find((k) => k.startsWith("__reactProps"))].onChange({...ev2, target: second});

      inp.value = 10;
      const ev3 = new Event("change", { bubbles: true, cancelable: false });
      inp[Object.keys(inp).find((k) => k.startsWith("__reactProps"))].onChange({...ev3, target: inp});
    });

    const nonFormContent = getInnerTextExcept(s, "input,select");
    const { amount } = nonFormContent.match(/(?<amount>\d+\.\d{2}) [A-Z]{3}/).groups;

    try {
      assert.notEqual(Number(amount), Number(inp.value));
    } catch (e) {
      console.error(e);
      throw e;
    }
   }
  }
}
```

# --before-all--

```js
async function delay(time) {
  return new Promise((resolve) => setTimeout(resolve, time));
}

function getInnerTextExcept(doc, removingSelector) {
  const body = doc.cloneNode(true);

  const squareElements = body.querySelectorAll(removingSelector);
  squareElements.forEach(element => {
   element.parentNode.removeChild(element)
  });

  return body.innerText;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
   <meta charset="UTF-8" />
   <title>Currency Converter</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
   <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
   ></script>
   <link rel="stylesheet" href="styles.css" />
</head>

<body>
   <div id="root"></div>
   <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
   >
    import { CurrencyConverter } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
   </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useMemo } = React;

export function CurrencyConverter() {

}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
   <meta charset="UTF-8" />
   <title>Currency Converter</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
   <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
   ></script>
   <link rel="stylesheet" href="styles.css" />
</head>

<body>
   <div id="root"></div>
   <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
   >
    import { CurrencyConverter } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
   </script>
</body>

</html>
```

```css
body {
   background-color: #0a0a23;
   font-family: Lato, sans-serif;
}

main {
   background-color: #1b1b32;
   padding: 20px;
   border-radius: 8px;
   box-shadow: 0px 4px 6px rgba(0, 0, 0, 0.1);
   text-align: center;
   width: 400px;
   margin: 50px auto;
}

h1 {
   margin-bottom: 16px;
   color: #f5f6f7;
}

label {
   color: #f5f6f7;
}

input,
select {
   width: 100%;
   display: block;
   box-sizing: border-box;
   padding: 8px;
   margin-bottom: 12px;
   margin-top: 4px;
   border: 1px solid #99c9ff;
   border-radius: 4px;
}

.conversion-display {
   font-size: 16px;
   color: #f5f6f7;
   margin-bottom: 10px;
}

p {
   font-size: 18px;
   font-weight: bold;
   color: #acd157;
}
```

```jsx
export function CurrencyConverter() {
  const [amount, setAmount] = React.useState(1);
  const [startCurrency, setStartCurrency] = React.useState("USD");
  const [endCurrency, setEndCurrency] = React.useState("EUR");

  const exchangeRates = {
   USD: 1,
   EUR: 0.85,
   GBP: 0.75,
   JPY: 110
  }

  const convertedAmounts = React.useMemo(() => {
   const converted = {};
   Object.keys(exchangeRates).forEach((curr) => {
    converted[curr] = ((amount / exchangeRates[startCurrency]) * exchangeRates[curr]).toFixed(2);
   });
   return converted;
  }, [amount, startCurrency]);

  return (
   <main>
    <h1>Currency Converter</h1>
    <p className="conversion-display">{startCurrency} to {endCurrency} Conversion</p>
    <input
      type="number"
      value={amount}
      onChange={(e) => setAmount(Number(e.target.value))}
    />
    <label>
      Start Currency:
      <select value={startCurrency} onChange={(e) => setStartCurrency(e.target.value)}>
       {Object.keys(exchangeRates).map((curr) => (
        <option key={curr} value={curr}>
          {curr}
        </option>
       ))}
      </select>
    </label>
    <label>
      Target Currency:
      <select value={endCurrency} onChange={(e) => setEndCurrency(e.target.value)}>
       {Object.keys(exchangeRates).map((curr) => (
        <option key={curr} value={curr}>
          {curr}
        </option>
       ))}
      </select>
    </label>
    <p>Converted Amount: <span>{convertedAmounts[endCurrency]}</span> {endCurrency}</p>
   </main>
  );
}
```

