---
id: 66f2836c459cfb16ae76f24f
title: 构建金字塔生成器
challengeType: 26
dashedName: lab-pyramid-generator
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个名为 `pyramid` 的函数，它接收三个参数。
1. 第一个参数应为一个字符串，表示在金字塔中重复的模式字符。
1. 第二个参数应为一个整数，表示金字塔的行数。
1. 第三个参数应为一个布尔值。
1. `pyramid` 函数应返回一个字符串，其中模式字符被重复并排列成一个顶点朝上的金字塔（当第三个参数为 `false` 时）。
1. 当第三个参数为 `true` 时，金字塔的顶点应朝下。
1. 顶点行应只有一个模式字符，每增加一行，模式字符数量比上一行多两个。
1. 每一行应以足够的空格开头，以使每行的中心字符在同一列对齐。
1. 金字塔应以换行符开始和结束。

例如，调用 `pyramid("o", 4, false)` 应返回如下输出：

```js

    o
  ooo
 ooooo
ooooooo

```

# --hints--

你应该有一个名为 `pyramid` 的函数。

```js
assert.isFunction(pyramid);
```

你的 `pyramid` 函数应有三个参数。

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)` 应返回 `"\n   o\n  ooo\n ooooo\nooooooo\n"`。

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)` 应返回 `"\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n"`。

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
     const rows = []
     for (let i = 1; i <= count; i++) {
          if (isInverted) {
                rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
          } else {
                rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
          }        
     }
     return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```

