---
id: 66f686b8ebdb982fa8e14330
title: 构建日期转换程序
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

在本实验中，你将编写一个日期转换程序，将给定日期转换为不同的格式。例如，当前日期 `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` 可以被格式化为以下两种方式：

- `(MM/DD/YYYY)：9/27/2024`
- `(Month Day, Year)：September 27, 2024`

**目标：** 完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要创建一个名为 `currentDate` 的变量，并使用 `Date` 对象为其赋值为当前日期和时间。
2. 你需要创建一个名为 `currentDateFormat` 的变量，并赋值为字符串 `Current Date and Time: [current date]`。将 `[current date]` 替换为 `currentDate` 变量的结果。
3. 你需要将 `currentDateFormat` 的值输出到控制台。
4. 你需要创建一个名为 `formatDateMMDDYYYY` 的函数，该函数接收一个日期对象作为参数。你可以自定义参数名。
5. 你的 `formatDateMMDDYYYY` 函数应将当前日期转换为 `月/日/年` 格式，并返回字符串 `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`。
6. 你需要创建一个名为 `formatDateLong` 的函数，该函数接收一个日期对象作为参数。
7. 你的 `formatDateLong` 函数应将当前日期转换为 `Month Day, Year` 格式，并返回字符串 `Formatted Date (Month Day, Year): [formatted date]`。

**注意：** 为了通过测试，格式化日期时请确保使用 `en-US` 作为本地化设置。

# --hints--

你应该有一个名为 `currentDate` 的变量，使用 `Date` 对象保存当前日期和时间。

```js
// 排除时间戳
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

你应该有一个名为 `currentDateFormat` 的变量，保存格式为 `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>` 的当前日期。

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

你应该将 `currentDateFormat` 的值输出到控制台。

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

你应该有一个名为 `formatDateMMDDYYYY` 的函数

```js
assert.isFunction(formatDateMMDDYYYY);
```

`formatDateMMDDYYYY` 函数应接收一个参数。

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

当日期对象为 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 时，`formatDateMMDDYYYY` 应返回 `Formatted Date (MM/DD/YYYY): 9/27/2024`。

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// 防止硬编码

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

你应该有一个名为 `formatDateLong` 的函数

```js
assert.isFunction(formatDateLong);
```

`formatDateLong` 函数应接收一个参数。

```js
assert.lengthOf(formatDateLong, 1);
```

当日期对象为 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 时，`formatDateLong` 应返回 `Formatted Date (Month Day, Year): September 27, 2024`。

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// 防止硬编码

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

