---
id: 67e3a6b7f60b4085588189e6
title: 构建井字棋游戏
challengeType: 25
dashedName: build-a-tic-tac-toe-game
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应创建一个 `Board` 组件，在 3x3 网格中渲染九个带有 `square` 类的 `button` 元素。
2. 点击 `button.square` 元素应在元素内依次显示 `X` 和 `O`。
3. 一旦有玩家获胜，点击任何 `button.square` 都不应再产生变化。
4. 你应创建一个 `button#reset` 元素，点击时重置游戏。
5. 应显示一条消息，指示 `X` 或 `O` 获胜，或如果结果为平局则显示平局。

# --hints--

你应导出一个 `Board` 组件。

```js
async () => {
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);

  assert.property(exports, "Board");
}
```

你应有九个 `button.square` 元素。

```js
const els = document.querySelectorAll("button.square");
assert.equal(els.length, 9);
```

`button.square` 元素应排列为 3x3 网格。

```js
// TODO: 也许要强制按钮大小一致？
const els = document.querySelectorAll("button.square");
const squares = Array.from(els);

const coords = squares.map((square) => {
  const rect = square.getBoundingClientRect();
  return rect;
});

const xTolerance = coords[0].width / 10;
const yTolerance = coords[0].height / 10;

try {
  assert.isBelow(coords[0].x, coords[1].x, "第一个格子应在第二个格子的左侧");
  assert.isBelow(coords[1].x, coords[2].x, "第二个格子应在第三个格子的左侧");
  assert.approximately(coords[0].y, coords[1].y, yTolerance, "第一个格子应与第二个格子在同一高度");
  assert.approximately(coords[1].y, coords[2].y, yTolerance, "第二个格子应与第三个格子在同一高度");

  assert.isBelow(coords[3].x, coords[4].x, "第四个格子应在第五个格子的左侧");
  assert.isBelow(coords[4].x, coords[5].x, "第五个格子应在第六个格子的左侧");
  assert.approximately(coords[3].y, coords[4].y, yTolerance, "第四个格子应与第五个格子在同一高度");
  assert.approximately(coords[4].y, coords[5].y, yTolerance, "第五个格子应与第六个格子在同一高度");

  assert.isBelow(coords[6].x, coords[7].x, "第七个格子应在第八个格子的左侧");
  assert.isBelow(coords[7].x, coords[8].x, "第八个格子应在第九个格子的左侧");
  assert.approximately(coords[6].y, coords[7].y, yTolerance, "第七个格子应与第八个格子在同一高度");
  assert.approximately(coords[7].y, coords[8].y, yTolerance, "第八个格子应与第九个格子在同一高度");

  assert.isBelow(coords[0].y, coords[3].y, "第一个格子应在第四个格子上方");
  assert.isBelow(coords[3].y, coords[6].y, "第四个格子应在第七个格子上方");
  assert.approximately(coords[0].x, coords[3].x, xTolerance, "第一个格子应与第四个格子在同一宽度");
  assert.approximately(coords[3].x, coords[6].x, xTolerance, "第四个格子应与第七个格子在同一宽度");

  assert.isBelow(coords[1].y, coords[4].y, "第二个格子应在第五个格子上方");
  assert.isBelow(coords[4].y, coords[7].y, "第五个格子应在第八个格子上方");
  assert.approximately(coords[1].x, coords[4].x, xTolerance, "第二个格子应与第五个格子在同一宽度");
  assert.approximately(coords[4].x, coords[7].x, xTolerance, "第五个格子应与第八个格子在同一宽度");

  assert.isBelow(coords[2].y, coords[5].y, "第三个格子应在第六个格子上方");
  assert.isBelow(coords[5].y, coords[8].y, "第六个格子应在第九个格子上方");
  assert.approximately(coords[2].x, coords[5].x, xTolerance, "第三个格子应与第六个格子在同一宽度");
  assert.approximately(coords[5].x, coords[8].x, xTolerance, "第六个格子应与第九个格子在同一宽度");
} catch (e) {
  console.error(e)
  throw e;
}
```

第一次点击 `button.square` 元素应在元素内显示 `X`。

```js
async () => {
  const el = document.querySelector("button.square");
  el.click();

  await delay(50);

  try {
    assert.include(el.textContent, "X");
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

点击 `button#reset` 元素应重置游戏。

```js
async () => {
  // 注意：此测试故意放在前面，因为后续测试依赖此功能。
  const el = document.querySelector("button.square");
  el.click();
  await delay(50);

  try {
    await reset(assert);
    assert.notInclude(el.textContent, "X");
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

第二次点击 `button.square` 元素应在元素内显示 `O`。

```js
async () => {
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  els[0].click();
  await delay(50);
  els[1].click();

  await delay(50);

  try {
    assert.include(els[1].textContent, "O");
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

后续点击 `button.square` 元素应在元素内依次显示 `X` 和 `O`。

```js
async () => {
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  try {
    for (let i = 3; i < els.length + 3; i++) {
      const wrappedI = i % els.length;
      els[wrappedI].click();
      await delay(50);
      assert.include(els[wrappedI].textContent, (i - 3) % 2 === 0 ? "X" : "O");
    }
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

点击已被使用的 `button.square` 元素应无变化。

```js
async () => {
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  try {
    for (let i = 3; i < els.length + 3; i++) {
      const wrappedI = i % els.length;
      // 连续点击同一个按钮，确保不会变化
      els[wrappedI].click();
      await delay(50);
      els[wrappedI].click();
      await delay(50);
      assert.include(els[wrappedI].textContent, (i - 3) % 2 === 0 ? "X" : "O");
    }
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

游戏结束后点击 `button.square` 元素应无变化。

```js
async () => {
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  try {
    // 获胜后，点击空格子，确保无变化
    for (let i = 0; i < 3; i++) {
      const x = els[i];
      x.click();
      await delay(50);
      assert.include(x.textContent, "X");

      const o = els[i + 3];
      o.click();
      await delay(50);
      if (i === 2) {
        assert.notInclude(o.textContent, "O");
      } else {
        assert.include(o.textContent, "O");
      }
    }
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

游戏应显示 `X` 或 `O` 获胜的消息。

```js
async () => {
  // 进入接近获胜状态
  // 检查 dom
  // 点击获胜格子
  // 检查 dom 是否有变化
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  try {
    for (let i = 0; i < 2; i++) {
      const x = els[i];
      x.click();
      await delay(50);
      const o = els[i + 3];
      o.click();
      await delay(50);
    }

    const preXWin = getInnerTextExcept("button.square");
    // 点击 X 获胜的格子
    els[2].click();
    await delay(50);

    const postXWin = getInnerTextExcept("button.square");
    assert.notEqual(preXWin, postXWin);

    await reset(assert);

    for (let i = 0; i < 2; i++) {
      const x = els[i];
      x.click();
      await delay(50);
      const o = els[i + 3];
      o.click();
      await delay(50);
    }
    els[6].click();
    await delay(50);

    const preOWin = getInnerTextExcept("button.square");
    // 点击 O 获胜的格子
    els[5].click();
    await delay(50);

    const postOWin = getInnerTextExcept("button.square");
    assert.notEqual(preOWin, postOWin);
    // `O` 和 `X` 获胜时应有不同的显示。
    assert.notEqual(postXWin, postOWin);
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

游戏应显示平局的消息。

```js
async () => {
  // 进入接近平局状态
  // 检查 dom
  // 点击最后一个格子
  // 检查 dom 是否有变化
  await reset(assert);
  const els = document.querySelectorAll("button.square");
  try {
    for (let i = 3; i < els.length + 2; i++) {
      const wrappedI = i % els.length;
      els[wrappedI].click();
      await delay(50);
    }

    const pre = getInnerTextExcept("button.square");
    els[2].click();
    await delay(50);

    const post = getInnerTextExcept("button.square");
    assert.notEqual(pre, post);
  } catch(e) {
    console.error(e);
    throw e;
  }
}
```

# --before-all--

```js
async function delay(time) {
  return new Promise((resolve) => setTimeout(resolve, time));
}

async function reset(assert) {
  const reset = document.querySelector("#reset");
  assert.exists(reset, "button#reset 应存在");
  reset.click();
  await delay(50);
}

// 获取文档中除匹配选择器外的文本内容。
function getInnerTextExcept(removingSelector) {
  const body = document.body.cloneNode(true);

  const squareElements = body.querySelectorAll(`${removingSelector},script`);
  squareElements.forEach(element => {
    element.parentNode.removeChild(element)
  });

  return body.innerText;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>井字棋</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Board } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Board />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState } = React;

export function Board() {

}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>井字棋</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Board } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Board />);
    </script>
</body>

</html>
```

```css
* {
    font-family: "Secular One", sans-serif;
    font-weight: 400;
    font-style: normal;
}

.board {
    display: flex;
    flex-direction: column;
    align-items: center;
}

.status {
    margin: 10px;
    font-size: 1.2em;
}

.board-row {
    display: flex;
}

.square {
    width: 60px;
    height: 60px;
    font-size: 1.5em;
    margin: 5px;
    background: #fff;
    border: 1px solid #999;
    cursor: pointer;
    border-radius: 5px;
}

.square:focus {
    outline: none;
}

#reset {
    margin-top: 20px;
    padding: 10px 20px;
    font-size: 1em;
    cursor: pointer;
}
```

```jsx
export function Board() {
  const [squares, setSquares] = React.useState(Array(9).fill(null));
  const [isXNext, setIsXNext] = React.useState(true);
  const winner = calculateWinner(squares);
  const isDraw = squares.every(square => square !== null) && !winner;

  function handleClick(index) {
    if (squares[index] || winner || isDraw) return;
    const nextSquares = squares.slice();
    nextSquares[index] = isXNext ? 'X' : 'O';
    setSquares(nextSquares);
    setIsXNext(!isXNext);
  };

  function resetGame() {
    setSquares(Array(9).fill(null));
    setIsXNext(true);
  };

  function renderSquare(index) {
    return <Square value={squares[index]} onClick={() => handleClick(index)} />;
  };

  return (
    <div className="board">
      <h1>井字棋</h1>
      <div className="status">
        {winner ? `获胜者: ${winner}` : isDraw ? "平局！" : `下一位玩家: ${isXNext ? 'X' : 'O'}`}
      </div>
      <div className="board-row">
        {renderSquare(0)}
        {renderSquare(1)}
        {renderSquare(2)}
      </div>
      <div className="board-row">
        {renderSquare(3)}
        {renderSquare(4)}
        {renderSquare(5)}
      </div>
      <div className="board-row">
        {renderSquare(6)}
        {renderSquare(7)}
        {renderSquare(8)}
      </div>
      <button id="reset" onClick={resetGame}>重置游戏</button>
    </div>
  );
};

function Square({ value, onClick }) {
  return (
    <button className="square" onClick={onClick}>
      {value}
    </button>
  );
};

function calculateWinner(squares) {
  const lines = [
    [0, 1, 2], [3, 4, 5], [6, 7, 8],
    [0, 3, 6], [1, 4, 7], [2, 5, 8],
    [0, 4, 8], [2, 4, 6]
  ];
  for (let line of lines) {
    const [a, b, c] = line;
    if (squares[a] && squares[a] === squares[b] && squares[a] === squares[c]) {
      return squares[a];
    }
  }
  return null;
};
```

