---
id: 66db529d37ad966480ebb633
title: 构建一个午餐选择程序
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

在本实验中，你将构建一个帮助管理午餐选项的程序。你将使用一个午餐数组，向数组中添加和移除项目，并随机选择一个午餐选项。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你应创建一个名为 `lunches` 的变量，并将其赋值为空数组，用于存储午餐项目。

2. 你应创建一个函数 `addLunchToEnd`，该函数以数组为第一个参数，字符串为第二个参数。该函数应：
  - 将字符串添加到数组末尾。
  - 向控制台输出 `"[Lunch Item] added to the end of the lunch menu."`，其中 `[Lunch Item]` 是传入的字符串。
  - 返回更新后的数组。

3. 你应创建一个函数 `addLunchToStart`，该函数以数组为第一个参数，字符串为第二个参数。该函数应：
  - 将字符串添加到数组开头。
  - 向控制台输出 `"[Lunch Item] added to the start of the lunch menu."`，其中 `[Lunch Item]` 是传入的字符串。
  - 返回更新后的数组。

4. 你应创建一个函数 `removeLastLunch`，该函数以数组为参数。该函数应：
   - 移除数组的最后一个元素。
   - 如果移除成功，向控制台输出 `"[Lunch Item] removed from the end of the lunch menu."`，其中 `[Lunch Item]` 是被移除的元素。
   - 如果数组为空，向控制台输出 `"No lunches to remove."`。
   - 返回更新后的数组。

5. 你应创建一个函数 `removeFirstLunch`，该函数以数组为参数。该函数应：
   - 移除数组的第一个元素。
   - 如果移除成功，向控制台输出 `"[Lunch Item] removed from the start of the lunch menu."`，其中 `[Lunch Item]` 是被移除的元素。
   - 如果数组为空，向控制台输出 `"No lunches to remove."`。
   - 返回更新后的数组。

6. 你应创建一个函数 `getRandomLunch`，该函数以数组为参数。该函数应：
   - 从数组中随机选择一个元素。
   - 如果选择成功，向控制台输出 `"Randomly selected lunch: [Lunch Item]"`，其中 `[Lunch Item]` 是数组中的一个随机元素。
   - 如果数组为空，向控制台输出 `"No lunches available."`。

7. 你应创建一个函数 `showLunchMenu`，该函数以数组为参数，并且：
   - 如果数组中有元素，向控制台输出 `"Menu items: [Lunch Item], [Lunch Item]..."`，其中每个 `[Lunch item]` 是数组中的一个元素，顺序与数组一致。
   - 如果数组为空，向控制台输出 `"The menu is empty."`。

# --hints--

你应声明一个变量 `lunches` 并将其赋值为空数组，用于存储午餐项目。

```js
assert.isArray(lunches);
```

你应定义一个函数 `addLunchToEnd`。

```js
assert.isFunction(addLunchToEnd);
```

函数 `addLunchToEnd` 应有两个参数。

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` 应向控制台输出 `"Tacos added to the end of the lunch menu."`。

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` 应返回 `["Pizza", "Tacos", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

你应定义一个函数 `addLunchToStart`。

```js
assert.isFunction(addLunchToStart);
```

函数 `addLunchToStart` 应有两个参数。

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` 应向控制台输出 `"Sushi added to the start of the lunch menu."`。

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` 应返回 `["Pizza", "Burger", "Sushi"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

你应定义一个函数 `removeLastLunch`。

```js
assert.isFunction(removeLastLunch)
```

函数 `removeLastLunch` 应有一个参数。

```js
assert.lengthOf(removeLastLunch, 1);
```

当输入数组为空时，`removeLastLunch` 应向控制台输出 `"No lunches to remove."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` 应向控制台输出 `"Toast removed from the end of the lunch menu."`。

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` 应返回 `["Sushi", "Pizza"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

你应定义一个函数 `removeFirstLunch`。

```js
assert.isFunction(removeFirstLunch);
```

函数 `removeFirstLunch` 应有一个参数。

```js
assert.lengthOf(removeFirstLunch, 1);
```

当输入数组为空时，`removeFirstLunch` 应向控制台输出 `"No lunches to remove."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` 应向控制台输出 `"Salad removed from the start of the lunch menu."`。

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` 应返回 `["Pizza", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`、`addLunchToStart`、`removeLastLunch` 和 `removeFirstLunch` 应返回传入的同一个数组（经过更新）。

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

你应定义一个函数 `getRandomLunch`。

```js
assert.isFunction(getRandomLunch);
```

函数 `getRandomLunch` 应有一个参数。

```js
assert.lengthOf(getRandomLunch, 1);
```

当输入数组为空时，`getRandomLunch` 应向控制台输出 `"No lunches available."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

当输入数组不为空时，`getRandomLunch` 应以 `"Randomly selected lunch: [Lunch Item]"` 的格式向控制台输出。

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);

  // 检查是否正确输出第一个元素
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // 第二个元素
  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // 第三个元素
  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  // 恢复 Math.random
  Math.random = tempRandom;
  console.log = temp;
}
```

你应定义一个函数 `showLunchMenu`。

```js
assert.isFunction(showLunchMenu);
```

函数 `showLunchMenu` 应有一个参数。

```js
assert.lengthOf(showLunchMenu, 1);
```

当输入数组为空时，`showLunchMenu` 应向控制台输出 `"The menu is empty."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` 应向控制台输出 `"Menu items: Greens, Corns, Beans"`。

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` 应向控制台输出 `"Menu items: Pizza, Burger, Fries, Salad"`。

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// 向列表末尾添加新的午餐
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// 向列表开头添加新的午餐
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// 从列表移除最后一个午餐
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
  console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
  console.log("No lunches to remove.");
  }
  return lunches;
}

// 从列表移除第一个午餐
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
  console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
  console.log("No lunches to remove.");
  }
  return lunches;
}

// 从列表中随机获取一个午餐
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
  console.log("No lunches available.");
  }
  else {
  const randomIndex = Math.floor(Math.random() * lunches.length);
  const randomLunch = lunches[randomIndex];
  console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// 显示列表中的所有午餐
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
  console.log("The menu is empty.");
  }
  else {
  console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```

