---
id: 67d83df6f82eda3868dd2a84
title: 构建一个RPG角色
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

在本实验中，你将通过构建一个用于RPG冒险的角色生成小应用，练习Python的基础知识。

完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你需要有一个名为 `create_character` 的函数。
1. 该函数应依次接收角色名和三个属性：力量（strength）、智力（intelligence）、魅力（charisma）。
1. 角色名需要进行如下校验：
    - 如果角色名不是字符串，函数应返回 `The character name should be a string`。
    - 如果角色名长度超过10个字符，函数应返回 `The character name is too long`。
    - 如果角色名包含空格，函数应返回 `The character name should not contain spaces`。
1. 属性也需要校验：
    - 如果有一个或多个属性不是整数，函数应返回 `All stats should be integers`。
    - 如果有一个或多个属性小于1，函数应返回 `All stats should be no less than 1`。
    - 如果有一个或多个属性大于4，函数应返回 `All stats should be no more than 4`。
    - 如果所有属性之和不等于7，函数应返回 `The character should start with 7 points`。
1. 如果所有值都通过校验，函数应返回一个包含四行的字符串：
    - 第一行是角色名
    - 第2-4行分别以属性缩写 `STR`、`INT`、`CHA`（按此顺序）开头，后跟一个空格，然后是与属性值相等数量的实心圆点（`●`），再补足到10个空心圆点（`○`）。例如：如果力量为3，则应有3个实心圆点和7个空心圆点。圆点已在编辑器中给出。

`create_character("ren", 4, 2, 1)` 应返回如下字符串：

    ```md
    ren
    STR ●●●●○○○○○○
    INT ●●○○○○○○○○
    CHA ●○○○○○○○○○
    ```

注意：虽然 `str` 和 `int` 是属性的常用缩写，但它们是Python的保留关键字，不应作为变量名使用。

# --hints--

你需要有一个名为 `create_character` 的函数。

```js
({
     test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
     )
})
```

当 `create_character` 的第一个参数不是字符串时，应返回 `The character name should be a string`。

```js
({
     test: () => runPython(
          `
wrong_types = [3, 5, {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
     assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'

# 即使数值也不对，类型错误优先

assert create_character(1, 1, 1, 1) == 'The character name should be a string'
`
     )
})
```

当 `create_character` 的第一个参数长度超过10时，应返回 `The character name is too long`。

```js
({
     test: () => runPython(
          `
assert create_character("aaabbbcccdd", 4, 2, 1) == 'The character name is too long'
`
     )
})
```

当 `create_character` 的第一个参数包含空格时，应返回 `The character name should not contain spaces`。

```js
({
     test: () => runPython(
          `
assert create_character("cha cha", 4, 1, 2) == 'The character name should not contain spaces'
`
     )
})
```

当 `create_character` 的第二、三或四个参数不是整数时，应返回 `All stats should be integers`。

```js
({
     test: () => runPython(
          `
wrong_types = ["3", "5", {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}, "str", "friend"]
for wrong_type in wrong_types:
     assert create_character("friend", wrong_type, 2, 1) == 'All stats should be integers'
     assert create_character("friend", 2, wrong_type, 1) == 'All stats should be integers'
     assert create_character("friend", 2, 1, wrong_type) == 'All stats should be integers'
`
     )
})
```

当 `create_character` 的第二、三或四个参数小于1时，应返回 `All stats should be no less than 1`。

```js
({
     test: () => runPython(
          `
expected = "All stats should be no less than 1"

assert create_character("ren", 0, 4, 3) == expected
assert create_character("ren", 4, 0, 3) == expected
assert create_character("ren", 4, 3, 0) == expected
assert create_character("ren", -1, 4, 4) == expected
assert create_character("ren", 4, -1, 4) == expected
assert create_character("ren", 4, 4, -1) == expected
`
     )
})
```

当 `create_character` 的第二、三或四个参数大于4时，应返回 `All stats should be no more than 4`。

```js
({
     test: () => runPython(
          `
assert create_character("ren", 1, 5, 1) == 'All stats should be no more than 4'
assert create_character("ren", 5, 1, 1) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 5) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 7) == 'All stats should be no more than 4'
assert create_character("ren", 1, 7, 1) == 'All stats should be no more than 4'
assert create_character("ren", 7, 1, 1) == 'All stats should be no more than 4'
`
     )
})
```

当 `create_character` 的第二、三、四个参数之和不等于7时，应返回 `The character should start with 7 points`。

```js
({
     test: () => runPython(
          `
assert create_character("ren", 4, 4, 4) == 'The character should start with 7 points'
assert create_character("ren", 1, 1, 1) == 'The character should start with 7 points'
assert create_character("ren", 1, 2, 3) == 'The character should start with 7 points'
`
     )
})
```

`create_character("ren", 4, 2, 1)` 应返回 `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`。

```js
({
     test: () => runPython(
          `
assert create_character("ren", 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
     )
})
```

当 `create_character` 使用有效参数时，应按要求输出角色属性。

```js
({
     test: () => runPython(
          `
assert create_character("Bob", 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character("aaabbbcccd", 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
     )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
     if type(character_name) != str:
          return "The character name should be a string"
     if len(character_name) > 10:
          return "The character name is too long"
     if ' ' in character_name:
          return "The character name should not contain spaces"

     if type(strength) != int or type(intelligence) != int or type(charisma) != int:
         return "All stats should be integers"
     if strength < 1 or intelligence < 1 or charisma < 1:
          return "All stats should be no less than 1"
     if strength > 4 or intelligence > 4 or charisma > 4:
          return "All stats should be no more than 4"
     if strength + intelligence + charisma != 7:
          return "The character should start with 7 points"

     character_sheet = f'{character_name}\n'
     character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
     character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
     character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
     return character_sheet
```

