---
id: 6718d2d59337c822ecb697ff
title: 构建一个银行账户管理程序
challengeType: 26
dashedName: build-a-bank-account-management-program
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该定义一个名为 `BankAccount` 的类，并实现构造函数。构造函数应将默认 `balance`（余额）设置为 `0`，并初始化一个空数组 `transactions` 用于存储交易记录（对象）。

2. 每个存储在 `transactions` 数组中的交易应为一个包含两个属性的对象：`type` 和 `amount`。`type` 属性应为 `deposit`（存款）或 `withdraw`（取款），`amount` 属性为存入或取出的金额。

3. 你应该定义一个名为 `deposit` 的方法，接收存款金额作为参数。当存款金额大于 `0` 时，应：

    - 向 `transactions` 数组添加一个新对象，`type` 为 `deposit`，`amount` 为存入金额。
    - 更新余额。
    - 返回 `"Successfully deposited $[amount]. New balance: $[balance]"`。

4. 如果金额小于或等于 `0`，`deposit` 方法应返回 `"Deposit amount must be greater than zero."`。

5. 你应该定义一个名为 `withdraw` 的方法，接收取款金额作为参数。该方法应根据取款更新当前余额。当取款金额大于 `0` 且小于等于当前余额时，应：

    - 向 `transactions` 数组添加一个新对象，`type` 为 `withdraw`，`amount` 为取出金额。
    - 更新余额。
    - 返回 `"Successfully withdrew $[amount]. New balance: $[balance]"`。

6. 如果取款金额小于或等于 `0` 或大于当前余额，`withdraw` 方法应返回 `"Insufficient balance or invalid amount."`。

7. 你应该定义一个名为 `checkBalance` 的方法，以 `"Current balance: $[balance]"` 的格式返回当前余额。

8. 你应该定义一个名为 `listAllDeposits` 的方法，遍历 `transactions` 数组并以 `"Deposits: amount,amount,..."` 的格式返回所有存款。

9. 你应该定义一个名为 `listAllWithdrawals` 的方法，遍历 `transactions` 数组并以 `"Withdrawals: amount,amount,..."` 的格式返回所有取款。

10. 你应该创建一个名为 `myAccount` 的 `BankAccount` 实例。

11. 你的 `myAccount` 银行账户应至少有五笔交易。

12. 你的 `myAccount` 银行账户应至少有两笔存款。

13. 你的 `myAccount` 银行账户应至少有两笔取款。

14. 你的 `myAccount` 银行账户余额应大于 `$100`。

# --hints--

你应该定义一个名为 `BankAccount` 的类。

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

`BankAccount` 对象初始应有 `balance` 为 `0`，并有一个空数组 `transactions` 用于存储交易记录。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

你应该有一个接收存款金额参数的 `deposit` 方法。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

你应该有一个接收取款金额参数的 `withdraw` 方法。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

你应该有一个检查当前余额的 `checkBalance` 方法。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

你应该有一个列出所有存款的 `listAllDeposits` 方法。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

你应该有一个列出所有取款的 `listAllWithdrawals` 方法。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` 应返回 `"Successfully deposited $100. New balance: $100"`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// 防止硬编码
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` 应返回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// 防止硬编码
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` 应返回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

当账户余额为 `100` 时，`BankAccount.withdraw(150)` 应返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// 防止硬编码
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` 应返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// 防止硬编码
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` 应返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

当账户余额为 `200` 时，`BankAccount.withdraw(150)` 应返回 `"Successfully withdrew $150. New balance: $50"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

当账户余额为 `200` 时，`BankAccount.checkBalance()` 应返回 `"Current balance: $200"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// 防止硬编码
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

当你存入 `10`、`35`、`90` 时，`listAllDeposits` 方法应返回 `"Deposits: 10,35,90"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

当你取出 `20`、`50`、`100` 时，`listAllWithdrawals` 方法应返回 `"Withdrawals: 20,50,100"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

你应该有一个名为 `myAccount` 的 `BankAccount` 实例。

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

你的 `myAccount` 银行账户应至少有五笔交易。

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

你的 `myAccount` 银行账户应至少有两笔存款。

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

你的 `myAccount` 银行账户应至少有两笔取款。

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

你的 `myAccount` 银行账户余额应大于 `$100`。

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
     constructor() {
          this.balance = 0;
          this.transactions = [];
     }

     deposit(amount) {
          if (amount > 0) {
                this.transactions.push({ type: 'deposit', amount });
                this.balance += amount;
                return `Successfully deposited $${amount}. New balance: $${this.balance}`;
          } else {
                return 'Deposit amount must be greater than zero.';
          }
     }

     withdraw(amount) {
          if (amount > 0 && amount <= this.balance) {
                this.transactions.push({ type: 'withdraw', amount });
                this.balance -= amount;
                return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
          } else {
                return 'Insufficient balance or invalid amount.';
          }
     }

     checkBalance() {
          return `Current balance: $${this.balance}`;
     }

     listAllDeposits() {
          const deposits = this.transactions
                .filter(transaction => transaction.type === 'deposit')
                .map(transaction => transaction.amount);
          return `Deposits: ${deposits}`;
     }

     listAllWithdrawals() {
          const withdrawals = this.transactions
                .filter(transaction => transaction.type === 'withdraw')
                .map(transaction => transaction.amount);
          return `Withdrawals: ${withdrawals}`;
     }
}

const myAccount = new BankAccount();

console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```

