---
id: 5e44413e903586ffb414c94e
title: 创建一个预算 App
challengeType: 27
forumTopicId: 462361
dashedName: build-a-budget-app
---

# --description--

完成 `Category` 类。它应该能够根据不同的预算类别（如 *food*、*clothing* 和 *entertainment*）实例化对象。创建对象时，会传入类别的名称。该类应该有一个名为 `ledger` 的实例变量，它是一个列表。该类还应包含以下方法：

- `deposit` 方法，接受一个金额和描述。如果未提供描述，则默认为空字符串。该方法应将一个形如 `{'amount': amount, 'description': description}` 的对象附加到 ledger 列表中。
- `withdraw` 方法，与 `deposit` 方法类似，但传入的金额应作为负数存储在 ledger 中。如果资金不足，则不应向 ledger 添加任何内容。如果取款成功，此方法应返回 `True`，否则返回 `False`。
- `get_balance` 方法，根据已发生的存款和取款返回预算类别的当前余额。
- `transfer` 方法，接受一个金额和另一个预算类别作为参数。该方法应添加一笔取款，金额为指定金额，描述为“Transfer to [目标预算类别]”。然后，该方法应向另一个预算类别添加一笔存款，金额为指定金额，描述为“Transfer from [源预算类别]”。如果资金不足，则不应向任一 ledger 添加任何内容。如果转账成功，此方法应返回 `True`，否则返回 `False`。
- `check_funds` 方法，接受一个金额作为参数。如果金额大于预算类别的余额，则返回 `False`，否则返回 `True`。`withdraw` 方法和 `transfer` 方法都应使用此方法。

打印预算对象时，应显示：

- 一行 30 个字符的标题行，其中类别名称在由 `*` 字符组成的行中居中。
- ledger 中的项目列表。每行应显示描述和金额。应显示描述的前 23 个字符，然后是金额。金额应右对齐，包含两位小数，最多显示 7 个字符。
- 一行显示类别总计。

这是一个使用示例：

```py
food = Category('Food')
food.deposit(1000, 'deposit')
food.withdraw(10.15, 'groceries')
food.withdraw(15.89, 'restaurant and more food for dessert')
clothing = Category('Clothing')
food.transfer(50, clothing)
print(food)
```

这是输出示例：

```bash
*************Food*************
initial deposit        1000.00
groceries               -10.15
restaurant and more foo -15.89
Transfer to Clothing    -50.00
Total: 923.96
```

除了 `Category` 类之外，创建一个名为 `create_spend_chart` 的函数（在类外部），该函数接受一个类别列表作为参数。它应该返回一个条形图字符串。

该图表应显示传递给函数的每个类别中所花费的百分比。花费的百分比应仅根据取款计算，而不包括存款，并且它应该是每个类别花费金额占所有类别总花费金额的百分比。图表左侧应有 0 - 100 的标签。条形图中的“条形”应由“o”字符组成。每个条形的高度应向下取整到最接近的 10。条形下方的水平线应超出最后一个条形两个空格。每个类别名称应垂直写在条形下方。顶部应有一个标题，内容为“Percentage spent by category”。

此函数将使用最多四个类别进行测试。

请仔细查看下面的输出示例，并确保输出的间距与示例完全匹配。

```bash
Percentage spent by category
100|          
 90|          
 80|          
 70|          
 60| o        
 50| o        
 40| o        
 30| o        
 20| o  o     
 10| o  o  o  
    0| o  o  o  
        ----------
         F  C  A  
         o  l  u  
         o  o  t  
         d  t  o  
                h     
                i     
                n     
                g     
```

注意：按 F12 打开浏览器控制台以查看更详细的测试输出。

# --hints--

`deposit` 方法应在 `ledger` 实例变量中创建一个特定的对象。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_deposit(self):
                self.food.deposit(900, "deposit")
                actual = self.food.ledger[0]
                expected = {"amount": 900, "description": "deposit"}
                self.assertEqual(actual, expected, 'Expected "deposit" method to create a specific object in the ledger instance variable.')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

调用不带描述的 `deposit` 方法应创建一个空描述。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_deposit_no_description(self):
                self.food.deposit(45.56)
                actual = self.food.ledger[0]
                expected = {"amount": 45.56, "description": ""}
                self.assertEqual(actual, expected, 'Expected calling "deposit" method with no description to create a blank description.')        
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`withdraw` 方法应在 `ledger` 实例变量中创建一个特定的对象。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_withdraw(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                actual = self.food.ledger[1]
                expected = {"amount": -45.67, "description": "milk, cereal, eggs, bacon, bread"}
                self.assertEqual(actual, expected, 'Expected "withdraw" method to create a specific object in the ledger instance variable.')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

调用不带描述的 `withdraw` 方法应创建一个空描述。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_withdraw_no_description(self):
                self.food.deposit(900, "deposit")
                good_withdraw = self.food.withdraw(45.67)
                actual = self.food.ledger[1]
                expected = {"amount": -45.67, "description": ""}
                self.assertEqual(actual, expected, 'Expected "withdraw" method with no description to create a blank description.')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果取款成功，`withdraw` 方法应返回 `True`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_withdraw_no_description(self):
                self.food.deposit(900, "deposit")
                good_withdraw = self.food.withdraw(45.67)
                self.assertEqual(good_withdraw, True, 'Expected "withdraw" method to return "True".')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

调用 `food.deposit(900, 'deposit')` 和 `food.withdraw(45.67, 'milk, cereal, eggs, bacon, bread')` 应返回余额 `854.33`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_get_balance(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                actual = self.food.get_balance()
                expected = 854.33
                self.assertEqual(actual, expected, 'Expected balance to be 854.33')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

在类别对象上调用 `transfer` 方法应在该类别对象中创建一个特定的 ledger 项目。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                transfer_amount = 20
                good_transfer = self.food.transfer(transfer_amount, self.entertainment)
                actual = self.food.ledger[2]
                expected = {"amount": -transfer_amount, "description": "Transfer to Entertainment"}
                self.assertEqual(actual, expected, 'Expected "transfer" method to create a specific ledger item in food object.')
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果转账成功，`transfer` 方法应返回 `True`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                transfer_amount = 20 
                good_transfer = self.food.transfer(transfer_amount, self.entertainment)        
                self.assertEqual(good_transfer, True, 'Expected "transfer" method to return "True".')        
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

在类别对象上调用 `transfer` 应减少该类别对象的余额。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                transfer_amount = 20
                food_balance_before = self.food.get_balance()        
                good_transfer = self.food.transfer(transfer_amount, self.entertainment)
                food_balance_after = self.food.get_balance()
                self.assertEqual(food_balance_before - food_balance_after, transfer_amount, 'Expected "transfer" method to reduce balance in food object.')
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`transfer` 方法应增加作为其参数传递的类别对象的余额。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                transfer_amount = 20        
                entertainment_balance_before = self.entertainment.get_balance()
                good_transfer = self.food.transfer(transfer_amount, self.entertainment)        
                entertainment_balance_after = self.entertainment.get_balance()        
                self.assertEqual(entertainment_balance_after - entertainment_balance_before, transfer_amount, 'Expected "transfer" method to increase balance in entertainment object.')  
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`transfer` 方法应在作为其参数传递的类别对象中创建一个特定的 ledger 项目。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                transfer_amount = 20        
                good_transfer = self.food.transfer(transfer_amount, self.entertainment)
                actual = self.entertainment.ledger[0]
                expected = {"amount": transfer_amount, "description": "Transfer from Food"}
                self.assertEqual(actual, expected, 'Expected "transfer" method to create a specific ledger item in entertainment object.')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果传递给 `check_funds` 方法的金额大于类别余额，则该方法应返回 `False`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_check_funds(self):
                self.food.deposit(10, "deposit")
                actual = self.food.check_funds(20)
                expected = False
                self.assertEqual(actual, expected, 'Expected "check_funds" method to be False')   
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果传递给 `check_funds` 方法的金额不大于类别余额，则该方法应返回 `True`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_check_funds(self):
                self.food.deposit(10, "deposit")
                actual = self.food.check_funds(10)
                expected = True
                self.assertEqual(actual, expected, 'Expected "check_funds" method to be True')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果取款未成功，`withdraw` 方法应返回 `False`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_withdraw_no_funds(self):
                self.food.deposit(100, "deposit")
                good_withdraw = self.food.withdraw(100.10)
                self.assertEqual(good_withdraw, False, 'Expected "withdraw" method to return "False".')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

如果转账未成功，`transfer` 方法应返回 `False`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_transfer_no_funds(self):
                self.food.deposit(100, "deposit")
                good_transfer = self.food.transfer(100.10, self.entertainment)
                self.assertEqual(good_transfer, False, 'Expected "transfer" method to return "False".')    
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

打印 `Category` 实例应给出对象的不同字符串表示形式。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")

        def test_to_string(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(45.67, "milk, cereal, eggs, bacon, bread")
                self.food.transfer(20, self.entertainment)
                actual = str(self.food)
                expected = "*************Food*************\\ndeposit                 900.00\\nmilk, cereal, eggs, bac -45.67\\nTransfer to Entertainme -20.00\\nTotal: 834.33"
                self.assertEqual(actual, expected, 'Expected different string representation of object.')
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表顶部的标题应为 `Percentage spent by category`。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_create_spend_chart(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(105.55)
                chart = budget.create_spend_chart([self.food])
                expected = "Percentage spent by category"
                self.assertEqual(chart.split("\\n")[0], expected, "Chart should have correct title.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表左侧应有正确的百分比。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")

        def test_create_spend_chart(self):
                self.food.deposit(900, "deposit")
                self.food.withdraw(105.55)
                chart = budget.create_spend_chart([self.food])
                percentages = ["100|", " 90|", " 80|", " 70|", " 60|", " 50|", " 40|", " 30|", " 20|", " 10|", "  0|"]
                for line, percent in zip(chart.split("\\n")[1:], percentages):
                     self.assertTrue(line.startswith(percent), "Chart correct percentages in the vertical axis.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表上每个条形的高度应向下取整到最接近的 10。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(78)
                self.entertainment.withdraw(22)
                self.business.withdraw(8)

        def test_create_spend_chart_rounding_close_to_upper_and_lower_ten(self):
                chart_lines = budget.create_spend_chart([self.food, self.entertainment]).split("\\n")[1:12]
                result_lines = '''100|
 90|
 80|
 70| o
 60| o
 50| o
 40| o
 30| o
 20| o  o
 10| o  o
    0| o  o'''.split("\\n")

                self.assertEqual(len(chart_lines), len(result_lines), "Lines missing in chart.")
                for actual, expected in zip(chart_lines, result_lines):
                        self.assertTrue(actual.startswith(expected), "Expected different rounding of bars.")


        def test_create_spend_chart_rounding_single_digit(self):
                chart_lines = budget.create_spend_chart([self.business, self.food, self.entertainment]).split("\\n")[1:12]
                result_lines = '''100|
 90|
 80|
 70|    o
 60|    o
 50|    o
 40|    o
 30|    o
 20|    o  o
 10|    o  o
    0| o  o  o'''.split("\\n")

                self.assertEqual(len(chart_lines), len(result_lines), "Lines missing in chart.")
                for actual, expected in zip(chart_lines, result_lines):
                        self.assertTrue(actual.startswith(expected), "Expected different rounding of bars.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表中的每一行应具有相同的长度。不同类别的条形应由两个空格隔开，最后一个条形后应有两个额外的空格。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(78)
                self.entertainment.withdraw(22)
                self.business.withdraw(8)


        def test_create_spend_chart_chart_lines_have_expected_length(self):
                chart_categories = [[self.food, self.entertainment], [self.business, self.food, self.entertainment]]

                expected_lengths = [len(line) for line in ["  0| o  o  ", "  0| o  o  o  "]]
                expected_chart_lines = 11

                for categories, expected_length in zip(chart_categories, expected_lengths):
                        chart_lines = budget.create_spend_chart(categories).split("\\n")[1:12]

                        self.assertEqual(len(chart_lines), expected_chart_lines, "Lines missing in chart.")
                        for actual in chart_lines:
                                self.assertEqual(len(actual), expected_length, "Expected different length of the chart line. Check that all spacing is exact.")
`);
        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 应正确显示条形下方的水平线。每个类别使用三个 `-` 字符，并且总共超出最后一个条形两个字符。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(105.55)
                self.entertainment.withdraw(33.40)
                self.business.withdraw(10.99)

        def test_create_spend_chart_horizontal_bar(self):
                chart_categories = [[self.business], [self.business, self.food], [self.business, self.food, self.entertainment]]
                horizontal_lines = ["    ----", "    -------", "    ----------"]
                for categories, expected in zip(chart_categories, horizontal_lines):
                        actual = budget.create_spend_chart(categories).split("\\n")[12]
                        self.assertEqual(actual, expected, "Expected different horizontal bar. Check that all spacing is exact.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表的末尾不应有换行符。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")

        def test_create_spend_chart_no_ending_new_line(self):
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(105.55)
                self.entertainment.withdraw(33.40)
                self.business.withdraw(10.99)
                actual = budget.create_spend_chart([self.business, self.food, self.entertainment])
                self.assertFalse(actual.endswith("\\n"), "Expected chart to not have new line at the end.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 图表应将每个类别名称垂直写在条形下方。每行应具有相同的长度，每个类别应由两个空格隔开，最后一个类别后应有两个额外的空格。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(105.55)
                self.entertainment.withdraw(33.40)
                self.business.withdraw(10.99)

        def test_create_spend_chart_names_two_categories(self):
                chart = budget.create_spend_chart([self.food, self.entertainment])
                actual = "\\n".join(chart.split("\\n")[13:]).rstrip("\\n")
                expected = "     F  E  \\n     o  n  \\n     o  t  \\n     d  e  \\n        r  \\n        t  \\n        a  \\n        i  \\n        n  \\n        m  \\n        e  \\n        n  \\n        t  "
                self.assertEqual(actual, expected, "Expected different category names written vertically below the bar. Check that all spacing is exact.")

        def test_create_spend_chart_names_three_categories(self):
                chart = budget.create_spend_chart([self.business, self.food, self.entertainment])
                actual = "\\n".join(chart.split("\\n")[13:]).rstrip("\\n")
                expected = "     B  F  E  \\n     u  o  n  \\n     s  o  t  \\n     i  d  e  \\n     n     r  \\n     e     t  \\n     s     a  \\n     s     i  \\n           n  \\n           m  \\n           e  \\n           n  \\n           t  "
                self.assertEqual(actual, expected, "Expected different category names written vertically below the bar. Check that all spacing is exact.")
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

`create_spend_chart` 应打印不同的图表表示形式。检查所有间距是否准确。按 F12 打开浏览器控制台以获取更多详细信息。

```js
({
    test: () => {
        pyodide.FS.writeFile('/home/pyodide/budget.py', code);
        pyodide.FS.writeFile('/home/pyodide/test_module.py',`
import unittest
import budget
from importlib import reload

reload(budget)
class UnitTests(unittest.TestCase):
        maxDiff = None
        def setUp(self):
                self.food = budget.Category("Food")
                self.entertainment = budget.Category("Entertainment")
                self.business = budget.Category("Business")

        def test_create_spend_chart(self):
                self.food.deposit(900, "deposit")
                self.entertainment.deposit(900, "deposit")
                self.business.deposit(900, "deposit")
                self.food.withdraw(105.55)
                self.entertainment.withdraw(33.40)
                self.business.withdraw(10.99)
                actual = budget.create_spend_chart([self.business, self.food, self.entertainment])
                expected = "Percentage spent by category\\n100|          \\n 90|          \\n 80|          \\n 70|    o     \\n 60|    o     \\n 50|    o     \\n 40|    o     \\n 30|    o     \\n 20|    o  o  \\n 10|    o  o  \\n  0| o  o  o  \\n    ----------\\n     B  F  E  \\n     u  o  n  \\n     s  o  t  \\n     i  d  e  \\n     n     r  \\n     e     t  \\n     s     a  \\n     s     i  \\n           n  \\n           m  \\n           e  \\n           n  \\n           t  "
                self.assertEqual(actual, expected, 'Expected different chart representation. Check that all spacing is exact.')
`);

        const testCode = `
from unittest import main
from importlib import reload
import test_module
reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
        const out = runPython(testCode);
        assert(out);
    }
})
```

# --seed--

## --seed-contents--

```py
class Category:
        pass

def create_spend_chart(categories):
        pass
```

# --solutions--

```py
class Category:

        def __init__(self, name):
                self.name = name
                self.ledger = []
                self.balance = 0
                self.spent = 0

        def __str__(self):
                first_line = f'{self.name.center(30, "*")}\n'
                lines = ''
                total = f'Total: {format(self.balance, ".2f")}'

                for n in range(len(self.ledger)):
                        descr = self.ledger[n]["description"][:23]
                        am = format(float(self.ledger[n]["amount"]), ".2f")[:7]
                        lines = lines + f'{descr:<23}{am:>7}\n'

                return f'{first_line}{lines}{total}'

        def deposit(self, amount, description=''):
                self.ledger.append({
                        'amount': float(amount),
                        'description': description
                })
                self.balance = self.balance + float(amount)

        def withdraw(self, amount, description=''):
                if self.check_funds(amount):
                        self.ledger.append({
                                'amount': -float(amount),
                                'description': description
                        })
                        self.balance = self.balance - float(amount)
                        self.spent = self.spent + float(amount)
                        return True
                else:
                        return False

        def get_balance(self):
                return self.balance

        def transfer(self, amount, category):
                if self.check_funds(amount):
                        # withdraw
                        self.ledger.append({
                                'amount': -float(amount),
                                'description': f'Transfer to {category.name}'
                        })
                        self.balance = self.balance - float(amount)
                        # deposit
                        category.deposit(amount, f'Transfer from {self.name}')

                        return True
                else:
                        return False

        def check_funds(self, amount):
                if float(amount) > self.balance:
                        return False
                else:
                        return True


def create_spend_chart(categories):
        total_expenses = 0
        obj = {}
        col1 = []
        str = []
        final_str = 'Percentage spent by category\n'
        label_max_length = 0
        label_strings = []

        for category in categories:
                total_expenses = total_expenses + category.spent
                obj[category.name] = {'expenses': category.spent}
                obj[category.name]['label'] = list(category.name)
                if len(obj[category.name]['label']) > label_max_length:
                        label_max_length = len(obj[category.name]['label'])

        for category in categories:
                obj[category.name]['percent'] = (
                        (category.spent / total_expenses * 100) // 10) * 10
                obj[category.name]['column'] = []
                for i in range(0, 110, 10):
                        if obj[category.name]['percent'] >= i:
                                obj[category.name]['column'].insert(0, 'o')
                        else:
                                obj[category.name]['column'].insert(0, ' ')

        for i in range(0, 110, 10):
                col1.insert(0, i)

        for i in range(11):
                str.append("")
                for key in obj:
                        str[i] += (f'{obj[key]["column"][i]}  ')
                final_str += f'{col1[i]:>3}| {str[i]}\n'
        final_str += f'    {"-"*(1+3*len(obj))}\n   '

        for i in range(label_max_length):
                label_strings.append('  ')
                for k in obj:
                        if len(obj[k]['label']) < label_max_length:
                                obj[k]['label'].extend(
                                        f'{" "*(label_max_length-len(obj[k]["label"]))}')

                        label_strings[i] += f'{obj[k]["label"][i]}  '
                if i < label_max_length - 1:
                        label_strings[i] += '\n   '
                final_str += label_strings[i]

        print(final_str)
        return (final_str)

```
