---
id: 66f55eac933ff64ce654ca74
title: 构建一个 Markdown 转 HTML 转换器
challengeType: 25
dashedName: build-a-markdown-to-html-converter
demoType: onClick
---

# --description--

Markdown 是一种用于为纯文本文档添加格式化元素的标记语言。在本实验中，所有 HTML 和 CSS 都已为你提供。你需要使用 JavaScript 完成 Markdown 转 HTML 转换器应用，使其能够将基本的 Markdown 语法转换为 HTML 元素。

**注意：** 最终结果不会是一个全面的 Markdown 转 HTML 转换器，但你可以根据需要为其添加额外功能。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个名为 `convertMarkdown` 的函数，该函数不接受参数。
1. `convertMarkdown` 函数应使用正则表达式将 `#markdown-input` 中的 markdown 输入转换为 HTML，并返回包含 HTML 代码的字符串。
1. `convertMarkdown` 函数应将一级、二级和三级标题分别转换为对应的 `h1`、`h2` 和 `h3` 元素。Markdown 中的标题由与级别相同数量的 `#` 字符加一个空格和标题文本组成。`#` 字符应位于行首或前面只有空格。
1. `convertMarkdown` 函数应将加粗文本转换为 `strong` 元素。Markdown 中的加粗文本由一对双星号或一对双下划线包裹。
1. `convertMarkdown` 函数应将斜体文本转换为 `em` 元素。Markdown 中的斜体文本由一对星号或一对下划线包裹。
1. `convertMarkdown` 函数应将图片语法转换为 `img` 元素。Markdown 中的图片语法为 `![alt-text](image-source)`，其中 `alt-text` 是 `alt` 属性值，`image-source` 是 `src` 属性值。
1. `convertMarkdown` 函数应将链接语法转换为锚点元素。Markdown 中的链接语法为 `[link text](URL)`，其中 `link text` 是锚点文本，`URL` 是 `href` 属性值。
1. `convertMarkdown` 函数应将引用语法转换为 `blockquote` 元素。Markdown 中的引用由 `>` 加一个空格和引用文本组成。`>` 字符应位于行首或前面只有空格。
1. 当你在 `#markdown-input` 输入文本时，`convertMarkdown` 返回的原始 HTML 代码应显示在 `#html-output` 中。
1. 当你在 `#markdown-input` 输入文本时，`convertMarkdown` 返回的 HTML 代码应在 `#preview` 中渲染。

以下表格展示了 `convertMarkdown` 需要支持的所有 Markdown 语法及其对应的 HTML：

<table>
    <thead>
        <tr>
            <th>Markdown</th>
            <th>HTML</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td><code># heading 1</code></td>
            <td><code>&lth1&gtheading 1&lt/h1&gt</code></td>
        </tr>
        <tr>
            <td><code>## heading 2</code></td>
            <td><code>&lth2&gtheading 2&lt/h2&gt</code></td>
        </tr>
        <tr>
            <td><code>### heading 3</code></td>
            <td><code>&lth3&gtheading 3&lt/h3&gt</code></td>
        </tr>
        <tr>
            <td><code>**bold text**</code> 或 <code>__bold text__</code></td>
            <td><code>&ltstrong&gtbold text&lt/strong&gt</code></td>
        </tr>
        <tr>
            <td><code>*italic text*</code> 或 <code>_italic text_</code></td>
            <td><code>&ltem&gtitalic text&lt/em&gt</code></td>
        </tr>
        <tr>
            <td><code>![alt-text](image-source)</code></td>
            <td><code>&ltimg alt="alt-text" src="image-source"&gt</code></td>
        </tr>
        <tr>
            <td><code>[link text](URL)</code></td>
            <td><code>&lta href="URL"&gtlink text&lt/a&gt</code></td>
        </tr>
        <tr>
            <td><code>> quote</code></td>
            <td><code>&ltblockquote&gtquote&lt/blockquote&gt</code></td>
        </tr>
    </tbody>
</table>

**注意：** 请确保在 HTML 中引入你的 JavaScript 文件。

# --hints--

你应该有一个名为 `convertMarkdown` 的函数。

```js
assert.isFunction(convertMarkdown);
```

当 `#markdown-input` 的值为 `# title 1` 时，`convertMarkdown()` 应返回 `<h1>title 1</h1>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

当 `#markdown-input` 的值为 `# title 1` 时，`<h1>title 1</h1>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

当 `#markdown-input` 的值为 `# title 1` 时，应将文本为 `title 1` 的 `h1` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

当 `#markdown-input` 的值为 `some text # title 1` 时，`convertMarkdown()` 不应将 `# title 1` 转换为 `h1` 元素。

```js
const input = document.querySelector("#markdown-input");
input.value = "some text # title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH1 = testDiv.querySelector("h1")
assert.notExists(testH1);
```

当 `#markdown-input` 的值为 `# title 1`，新一行是 `# alternate title` 时，`convertMarkdown()` 应返回 `<h1>title 1</h1><h1>alternate title</h1>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1\n# alternate title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 1");
assert.equal(headings[1].innerText, "alternate title");
```

当 `#markdown-input` 的值为 `## title 2` 时，`convertMarkdown()` 应返回 `<h2>title 2</h2>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 2");
```

当 `#markdown-input` 的值为 `## title 2` 时，`<h2>title 2</h2>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

当 `#markdown-input` 的值为 `## title 2` 时，应将文本为 `title 2` 的 `h2` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

当 `#markdown-input` 的值为 `some text ## title 2` 时，`convertMarkdown()` 不应将 `## title 2` 转换为 `h2` 元素。

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH2 = testDiv.querySelector("h2")
assert.notExists(testH2);
```

当 `#markdown-input` 的值为 `## title 2`，新一行是 `## title 2 alt` 时，`convertMarkdown()` 应返回 `<h2>title 2</h2><h2>title 2 alt</h2>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2\n## title 2 alt";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 2");
assert.equal(headings[1].innerText, "title 2 alt");
```

当 `#markdown-input` 的值为 `### title 3` 时，`convertMarkdown()` 应返回 `<h3>title 3</h3>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

当 `#markdown-input` 的值为 `### title 3` 时，`<h3>title 3</h3>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

当 `#markdown-input` 的值为 `### title 3` 时，应将文本为 `title 3` 的 `h3` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

当 `#markdown-input` 的值为 `some text ### title 3` 时，`convertMarkdown()` 不应将 `### title 3` 转换为 `h3` 元素。

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH3 = testDiv.querySelector("h3")
assert.notExists(testH3);
```

当 `#markdown-input` 的值为 `### title 3`，新一行是 `### third title` 时，`convertMarkdown()` 应返回 `<h3>title 3</h3><h3>third title</h3>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3\n### third title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 3");
assert.equal(headings[1].innerText, "third title");
```

当 `#markdown-input` 的值为 `**this is bold**` 时，`convertMarkdown()` 应返回 `<strong>this is bold</strong>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(strongs[0].innerText, "this is bold");
```

当 `#markdown-input` 的值为 `**this is bold**` 时，`<strong>this is bold</strong>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

当 `#markdown-input` 的值为 `**this is bold**` 时，应将文本为 `this is bold` 的 `strong` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

当 `#markdown-input` 的值为 `**this is bold**`，新一行是 `**this is also bold**` 时，`convertMarkdown()` 应返回 `<strong>this is bold</strong><strong>this is also bold</strong>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**\n**this is also bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

当 `#markdown-input` 的值为 `__this is bold__` 时，`<strong>this is bold</strong>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

当 `#markdown-input` 的值为 `__this is bold__` 时，应将文本为 `this is bold` 的 `strong` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

当 `#markdown-input` 的值为 `__this is bold__`，新一行是 `__this is also bold__` 时，`convertMarkdown()` 应返回 `<strong>this is bold</strong><strong>this is also bold</strong>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "__this is bold__\n__this is also bold__";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

当 `#markdown-input` 的值为 `*this is italic*` 时，`convertMarkdown()` 应返回 `<em>this is italic</em>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `*this is italic*` 时，`<em>this is italic</em>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `*this is italic*` 时，应将文本为 `this is italic` 的 `em` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `*this is italic*`，新一行是 `*this is also italic*` 时，`convertMarkdown()` 应返回 `<em>this is italic</em><em>this is also italic</em>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*\n*this is also italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

当 `#markdown-input` 的值为 `_this is italic_` 时，`convertMarkdown()` 应返回 `<em>this is italic</em>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `_this is italic_` 时，`<em>this is italic</em>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `_this is italic_` 时，应将文本为 `this is italic` 的 `em` 元素作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

当 `#markdown-input` 的值为 `_this is italic_`，新一行是 `_this is also italic_` 时，`convertMarkdown()` 应返回 `<em>this is italic</em><em>this is also italic</em>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_\n_this is also italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

当 `#markdown-input` 的值为 `# **title 1**` 或 `# __title 1__` 时，`convertMarkdown()` 应返回 `<h1><strong>title 1</strong></h1>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "# **title 1**";
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = convertMarkdown();
let headings = testDiv1.querySelectorAll("h1");
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

input.value = "# __title 1__";
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = convertMarkdown();
headings = testDiv2.querySelectorAll("h1");
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

当 `#markdown-input` 的值为 `# **title 1**` 或 `# __title 1__` 时，`<h1><strong>title 1</strong></h1>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = output.innerText;
let headings = testDiv1.querySelectorAll("h1")
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

output.innerText = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = output.innerText;
headings = testDiv2.querySelectorAll("h1")
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

当 `#markdown-input` 的值为 `# **title 1**` 或 `# __title 1__` 时，应将 `<h1><strong>title 1</strong></h1>` 作为 `#preview` 的 innerHTML。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
let headings = preview.querySelectorAll("h1")
let testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

preview.innerHTML = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
headings = preview.querySelectorAll("h1")
testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

当 `#markdown-input` 的值为 `![alt-text](image-source)` 时，`convertMarkdown()` 应返回 `<img alt="alt-text" src="image-source">`。

```js
const input = document.querySelector("#markdown-input");
input.value = "![alt-text](image-source)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

当 `#markdown-input` 的值为 `![alt-text](image-source)` 时，`<img alt="alt-text" src="image-source">` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

当 `#markdown-input` 的值为 `![alt-text](image-source)` 时，应将 `<img alt="alt-text" src="image-source">` 作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const imgs = preview.querySelectorAll("img");
assert.lengthOf(preview.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

当 `#markdown-input` 的值为 `![alt-text](image-source)`，新一行是 `![alt-text-2](image-source-2)` 时，`convertMarkdown()` 应返回 `<img alt="alt-text" src="image-source"><img alt="alt-text-2" src="image-source-2">`。

```js
const input = document.querySelector("#markdown-input");
input.value = "!![alt-text](image-source)\n![alt-text-2](image-source-2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(imgs, 2);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
assert.equal(imgs[1].alt, "alt-text-2");
assert.isTrue(imgs[1].src.endsWith("image-source-2"));
```

当 `#markdown-input` 的值为 `[link text](URL)` 时，`convertMarkdown()` 应返回 `<a href="URL">link text</a>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

当 `#markdown-input` 的值为 `[link text](URL)` 时，`<a href="URL">link text</a>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

当 `#markdown-input` 的值为 `[link text](URL)` 时，应将 `<a href="URL">link text</a>` 作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const anchors = preview.querySelectorAll("a");
assert.lengthOf(preview.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

当 `#markdown-input` 的值为 `[link text](URL)`，新一行是 `[link text 2](URL2)` 时，`convertMarkdown()` 应返回 `<a href="URL">link text</a><a href="URL2">link text 2</a>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)\n[link text 2](URL2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(anchors, 2);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
assert.isTrue(anchors[1].href.endsWith("URL2"));
assert.equal(anchors[1].innerText, "link text 2");
```

当 `#markdown-input` 的值为 `> this is a quote` 时，`convertMarkdown()` 应返回 `<blockquote>this is a quote</blockquote>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

当 `#markdown-input` 的值为 `> this is a quote` 时，`<blockquote>this is a quote</blockquote>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

当 `#markdown-input` 的值为 `> this is a quote` 时，应将 `<blockquote>this is a quote</blockquote>` 作为 `#preview` 的子元素添加。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

当 `#markdown-input` 的值为 `> this is a quote`，新一行是 `> this is another quote` 时，`convertMarkdown()` 应返回 `<blockquote>this is a quote</blockquote><blockquote>this is another quote</blockquote>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote\n> this is another quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(quotes, 2);
assert.equal(quotes[0].innerText, "this is a quote");
assert.equal(quotes[1].innerText, "this is another quote");
```

当 `#markdown-input` 的值为 `some text > not a quote anymore` 时，`convertMarkdown()` 不应将 `> not a quote anymore` 转换为 `blockquote` 元素。

```js
const input = document.querySelector("#markdown-input");
input.value = "some text > not a quote anymore";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote")
assert.notExists(testQuote);
```

当 `#markdown-input` 的值为 `> **this is a *quote***` 时，`convertMarkdown()` 应返回 `<blockquote><strong>this is a <em>quote</em></strong></blockquote>`。

```js
const input = document.querySelector("#markdown-input");
input.value = "> **this is a *quote***";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

当 `#markdown-input` 的值为 `> **this is a *quote***` 时，`<blockquote><strong>this is a <em>quote</em></strong></blockquote>` 应显示在 `#html-output` 中。

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

当 `#markdown-input` 的值为 `> **this is a *quote***` 时，应将 `<blockquote><strong>this is a <em>quote</em></strong></blockquote>` 作为 `#preview` 的 innerHTML。

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
const strongs = preview.querySelectorAll("blockquote>strong");
const italics = preview.querySelectorAll("blockquote>strong>em");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

你的 HTML 中应只有一个 `script` 元素。

```js
assert.lengthOf(document.querySelectorAll("script"), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>Markdown 转 HTML 转换器</title>
        <link rel="stylesheet" href="styles.css">
</head>

<body>
        <h1>Markdown 转 HTML 转换器</h1>
        <div id="container">
                <div class="container">
                        <h2>Markdown 输入：</h2>
                        <textarea id="markdown-input" placeholder="在此输入你的 markdown..."></textarea>
                </div>
                <div class="container">
                        <h2>原始 HTML 输出：</h2>
                        <div id="html-output"></div>
                </div>
                <div class="container">
                        <h2>HTML 预览：</h2>
                        <div id="preview"></div>
                </div>
        </div>
</body>

</html>
```

```css
* {
         box-sizing: border-box;
}
 body {
         font-family: Arial, sans-serif;
         padding: 20px;
}
 #markdown-input {
         width: 100%;
         height: 100px;
}
 #html-output, #preview {
         height: 100px;
         display: inline-block;
         width: 100%;
         border: 1px solid #ccc;
         padding: 10px;
         margin: auto;
         white-space: pre-wrap;
         background-color: #f9f9f9;
}
 @media (min-width: 600px) {
         #markdown-input, #html-output, #preview {
                 height: 200px;
                 margin: 0;
        }
         #container {
                 display: flex;
                 justify-content: space-evenly;
                 gap: 10px;
        }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>Markdown 转 HTML 转换器</title>
        <link rel="stylesheet" href="styles.css">
</head>

<body>
        <h1>Markdown 转 HTML 转换器</h1>
        <div id="container">
                <div class="container">
                        <h2>Markdown 输入：</h2>
                        <textarea id="markdown-input" placeholder="在此输入你的 markdown..."></textarea>
                </div>
                <div class="container">
                        <h2>原始 HTML 输出：</h2>
                        <div id="html-output"></div>
                </div>
                <div class="container">
                        <h2>HTML 预览：</h2>
                        <div id="preview"></div>
                </div>
        </div>
        <script src="script.js"></script>
</body>

</html>
```

```css
* {
         box-sizing: border-box;
}
 body {
         font-family: Arial, sans-serif;
         padding: 20px;
}
 #markdown-input {
         width: 100%;
         height: 100px;
}
 #html-output, #preview {
         height: 100px;
         display: inline-block;
         width: 100%;
         border: 1px solid #ccc;
         padding: 10px;
         margin: auto;
         white-space: pre-wrap;
         background-color: #f9f9f9;
}
 @media (min-width: 600px) {
         #markdown-input, #html-output, #preview {
                 height: 200px;
                 margin: 0;
        }
         #container {
                 display: flex;
                 justify-content: space-evenly;
                 gap: 10px;
        }
}
```

```js
const patternReplacementArray = [
        { h6: [/^[ \t]*######[ \t](.+$)/gm, "<h6>$1</h6>"] },
        { h5: [/^[ \t]*#####[ \t](.+$)/gm, "<h5>$1</h5>"] },
        { h4: [/^[ \t]*####[ \t](.+$)/gm, "<h4>$1</h4>"] },
        { h3: [/^[ \t]*###[ \t](.+$)/gm, "<h3>$1</h3>"] },
        { h2: [/^[ \t]*##[ \t](.+$)/gm, "<h2>$1</h2>"] },
        { h1: [/^[ \t]*#[ \t](.+$)/gm, "<h1>$1</h1>"] },
        { hr: [/(?<=\n|^)---(?=\n|$)/, "<hr>"] },
        { strong: [/((?:\*|_){2})(.*?)\1/g, "<strong>$2</strong>"] },
        { em: [/(\*|_)(.*?)\1/g, "<em>$2</em>"] },
        { img: [/!\[\s*(.*?)\s*\]\(\s*(.*?)\s*\)/g, '<img alt="$1" src="$2">'] },
        { a: [/\[(.*?)\]\(\s*(.*?)\s*\)/g, '<a href="$2">$1</a>'] },
        { blockquote1: [/^[ \t]*> ([^\n]+?)$/gm, "<blockquote>$1</blockquote>\n"] },
        { blockquote2: [/(?<=<blockquote>)(.+?)<\/blockquote>\n\n^([^<\n]+?)$/gm, "$1 $2</blockquote>\n"] },
        { ul: [/(?:^[-*+]\s+.*?$\n)+\s*?/gm, "<ul>\n$&</ul>"] },
        { ol: [/(?:^\d*\.\s+?\w.*?$\n)+\s*?/gm, "<ol>\n$&</ol>"] },
        { uli: [/^[-*+]\s+(.*?$)\n/gm, "<li>$1</li>\n"] },
        { oli: [/^\d*\.\s+?(.*?$)\n/gm, "<li>$1</li>\n"] },
        { p1: [/^([^<>\s].*?)(?=\n\s*$)/gsm, "<p>$1</p>"] },
        { p2: [/^\w*?(?:<strong>|<em>).+(?:<\/strong>|<\/em>)\s*\w*?[ \t]*$/gm, "<p>$&</p>"] }
]

const markdownInput = document.getElementById('markdown-input');
const convertMarkdown = () => {
        let markdown = markdownInput.value;
        let matchUList = false;
        let matchOList = false;
        patternReplacementArray.forEach((item) => {
                const [key] = Object.keys(item);
                const [pattern, replacement] = Object.values(item)[0];
                if (key === "blockquote2") {
                        while (markdown.match(pattern)) {
                                markdown = markdown.replace(pattern, replacement);
                        }
                } else if (key === "ul") {
                        if (markdown.match(pattern)) {
                                matchUList = true;
                                markdown = markdown.replace(pattern, replacement);
                        }

                } else if (key === "ol") {
                        if (markdown.match(pattern)) {
                                matchOList = true;
                                markdown = markdown.replace(pattern, replacement);
                        }

                } else if (key === "uli" && !matchUList) {
                        return;
                } else if (key === "oli" && !matchOList) {
                        return;
                } else {
                        markdown = markdown.replace(pattern, replacement);
                }
        })
        return markdown;
}

markdownInput.addEventListener('input', () => {
        const html = convertMarkdown();
        const out = document.getElementById('html-output');
        const preview = document.getElementById('preview');
        out.textContent = html;
        preview.innerHTML = html;
});
```

