---
id: 66ebd4ae2812430bb883c787
title: 构建一个活动中心
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

在本实验中，你将利用语义化 HTML 元素来创建网页结构。你需要添加内容和图片，让它看起来像一个真实的活动中心。

完成以下用户故事并通过所有测试即可完成本实验。

**用户需求：**

1. 你应该有一个 `header` 元素。
1. 在 `header` 元素内，应该有一个包含文本 `Event Hub` 的 `h1` 元素，以及一个 `nav` 元素。
1. 在 `nav` 元素内，应该有一个包含两个项目的无序列表，每个项目都包含指向页面不同部分的链接。第一个项目文本为 `Upcoming Events`，第二个项目文本为 `Past Events`。
1. 每个链接都应该用 `a` 元素表示，并带有 `href` 属性，分别链接到页面对应部分：`#upcoming-events` 和 `#past-events`。
1. 你应该有一个包含页面不同部分的 `main` 元素。
1. 在 `main` 元素内，应该有两个 `section` 元素。
1. 第一个 `section` 元素应有 `id` 属性，值为 `upcoming-events`。
1. 在 `#upcoming-events` 区域内，你应该有：

    - 一个文本为 `Upcoming Events` 的 `h2` 元素。
    - 两个 `article` 元素。每个 article 代表一个活动，且应包含：
        - 一个活动标题的 `h3` 元素。
        - 一个活动描述的 `p` 元素。你可以在底部添加日期。

1. 第二个 `section` 元素应有 `id` 属性，值为 `past-events`。
1. 在 `#past-events` 区域内，你应该有：

    - 一个文本为 `Past Events` 的 `h2` 元素。
    - 两个 `article` 元素。每个 article 代表一个过去的活动，且应包含：
      - 一个活动标题的 `h3` 元素，
      - 一个活动描述的 `p` 元素。你可以在底部添加日期。
      - 一个图片元素，`src` 属性指向图片文件，`alt` 属性为图片描述。

**注意：** 活动描述和日期可以自拟。你可以使用以下图片链接：

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`

# --hints--

你应该有一个 `header` 元素。

```js
assert.isNotNull(document.querySelector("header"));
```

你的 `header` 元素应在 `body` 标签的开头。

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

在 `header` 元素内，应该有一个包含文本 `Event Hub` 的 `h1` 元素。

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText, "Event Hub");
```

在 `header` 元素内，`h1` 元素后应有一个 `nav` 元素。

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

你的 `nav` 元素应包含一个有两个项目的无序列表。

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

无序列表的第一个项目应为一个链接。

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

无序列表的第二个项目应为一个链接。

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

无序列表第一个项目的文本应为 `Upcoming Events`。

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText, "Upcoming Events");
```

无序列表第一个项目的 `href` 应为 `#upcoming-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

无序列表第二个项目的文本应为 `Past Events`。

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText, "Past Events");
```

无序列表第二个项目的 `href` 应为 `#past-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

在 `header` 元素闭合标签后应有一个 `main` 元素。

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

在 `main` 元素内应有两个 `section` 元素。

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

第一个 `section` 元素应有 `id` 属性，值为 `upcoming-events`。

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

第二个 `section` 元素应有 `id` 属性，值为 `past-events`。

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

在 `#upcoming-events` 区域内应有一个文本为 `Upcoming Events` 的 `h2` 元素。

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText, "Upcoming Events");
```

在 `#upcoming-events` 区域内，`h2` 元素下应有两个 `article` 元素。

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#upcoming-events` 区域内的每个 `article` 元素都应有一个活动标题的 `h3` 元素。

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#upcoming-events` 区域内的每个 `article` 元素都应有一个活动描述的段落元素。

```js
const pElement = document.querySelectorAll('#upcoming-events article h3+p').length;
assert.strictEqual(pElement, 2);
```

在 `#past-events` 区域内应有一个文本为 `Past Events` 的 `h2` 元素。

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText, "Past Events");
```

在 `#past-events` 区域内，`h2` 元素下应有两个 `article` 元素。

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#past-events` 区域内的每个 `article` 元素都应有一个活动标题的 `h3` 元素。

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#past-events` 区域内的每个 `article` 元素都应有一个活动描述的段落元素。

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
     assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` 区域内的每个 `article` 元素都应有一个图片元素。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` 区域内的每个图片元素都应有 `src` 属性，指向图片文件。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
     assert.exists(img.getAttribute("src"));
}
```

`#past-events` 区域内的每个图片元素都应有 `alt` 属性，描述图片内容。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
     assert.exists(img.getAttribute("alt"));
}
```

每个 `h3` 元素都应有活动标题。

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText)));
```

每个 `p` 元素都应有活动描述。

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText)));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
     <meta charset="utf-8">
     <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Event Hub</title>
</head>
<body>
     <header>
          <h1>Event Hub</h1>
          <nav>
                <ul>
                     <li><a href="#upcoming-events">Upcoming Events</a></li>
                     <li><a href="#past-events">Past Events</a></li>
                </ul>
          </nav>
     </header>
     <main>
          <section id="upcoming-events">
                <h2>Upcoming Events</h2>
                <article>
                     <h3>AI & Machine Learning Conference 2024</h3>
                     <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                     <p>Date: August 10, 2024</p>
                </article>
                <article>
                     <h3>Web Development Bootcamp</h3>
                     <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                     <p>Date: September 5, 2024</p>
                </article>
          </section>
          <section id="past-events">
                <h2>Past Events</h2>
                <article>
                     <h3>Cybersecurity Summit 2024</h3>
                     <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                     <p>Date: June 15, 2024</p>
                     <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
                </article>
                <article>
                     <h3>Blockchain Expo 2024</h3>
                     <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                     <p>Date: July 20, 2024</p>
                     <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
                </article>
          </section>
     </main>
</body>
</html>
```

