---
id: 66d75dd0aa65a71600dc669b
title: 构建一个库存管理程序
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

在本实验中，你将构建一个库存管理程序，允许你添加、更新、查找和移除库存中的产品。你将使用一个对象数组来表示库存，每个对象包含 `name`（产品名称）和 `quantity`（数量）两个键。

完成以下用户需求并通过所有测试即可完成本实验。

**用户需求：**

1. 你应该声明一个名为 `inventory` 的空数组，用于存储产品对象。每个对象包含一个键 `name`，其值为小写字符串，以及一个键 `quantity`，其值为整数。
1. 你应该创建一个名为 `findProductIndex` 的函数，该函数以产品名称为参数，返回对应产品对象在 `inventory` 数组中的索引。该函数应始终使用产品名称的小写形式进行查找。如果未找到产品，则返回 `-1`。
1. 你应该创建一个名为 `addProduct` 的函数，该函数以产品对象为参数。
1. 如果产品已存在于库存中，`addProduct` 函数应更新其 `quantity` 值，并在控制台输出产品名称加空格和 `quantity updated`。
1. 如果产品不存在于库存中，`addProduct` 函数应将产品添加到 `inventory` 数组，并在控制台输出产品名称加空格和 `added to inventory`。
1. 你应该创建一个名为 `removeProduct` 的函数，该函数以产品名称和数量为参数。
1. `removeProduct` 函数应从库存中对应产品对象的数量中减去传入的数量，并在控制台输出 `Remaining <product-name> pieces: <product-quantity>`，其中 `<product-name>` 替换为产品名称，`<product-quantity>` 替换为产品数量。
1. 如果减去后的数量为零，`removeProduct` 应将该产品对象从库存中移除。如果库存中的数量不足以进行减法，`removeProduct` 应在控制台输出 `Not enough <product-name> available, remaining pieces: <product-quantity>`。
1. 如果要移除的产品在库存中不存在，`removeProduct` 应在控制台输出 `<product-name> not found`。

**注意：** 为避免冲突，运行测试时仅保留用户需求中提到的日志输出。

# --hints--

你应该声明一个名为 `inventory` 的空数组。

```js
assert.match(__helpers.removeJSComments(code), /(let|const)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

你应该有一个名为 `findProductIndex` 的函数。

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` 应返回 `inventory` 数组中 `name` 等于 `"flour"` 的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` 应返回 `inventory` 数组中 `name` 等于 `"flour"` 的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` 在 `inventory` 数组中未找到 `name` 等于 `"flour"` 的对象时应返回 `-1`。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

你应该有一个名为 `addProduct` 的函数。

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` 应将 `5` 加到 `inventory` 数组中 `name` 等于 `"flour"` 的对象的 `quantity` 值上。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` 在 `inventory` 数组中找到 `name` 等于 `"flour"` 的对象时应输出 `flour quantity updated`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` 在 `inventory` 数组中未找到 `name` 等于 `"flour"` 的对象时应将 `{name: "flour", quantity: 5}` 添加到 `inventory` 数组。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` 在 `inventory` 数组中未找到 `name` 等于 `"flour"` 的对象时应输出 `flour added to inventory`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

你应该创建一个名为 `removeProduct` 的函数。

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` 在 `inventory` 数组中未找到 `name` 等于 `"flour"` 的对象时应输出 `flour not found`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` 应从 `inventory` 数组中 `name` 等于 `"flour"` 的对象的 `quantity` 值中减去 `5`。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` 在 `inventory` 等于 `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]` 时应输出 `Remaining flour pieces: 15`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

如果减去后的数量为零，`removeProduct` 应将该产品对象从库存中移除。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` 在 `inventory` 等于 `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]` 时应输出 `Not enough flour available, remaining pieces: 5`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// 示例用法

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```

