---
id: 669e81368e52b3a5c35a2dc5
title: 构建一个视频合辑页面
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个 `main` 元素作为 `body` 元素的唯一子元素。
1. 你应该有一个 `h1` 元素，作为页面主题。
1. 你应该在 `h1` 元素下方有一个介绍页面主题的段落。
1. 你应该在第一个段落下方有三个 `section` 元素。
1. 每个 section 应包含一个 `h2` 元素、一个段落和一个 `iframe` 元素，顺序依次为上述三者。
1. 三个 `iframe` 元素的 `src` 属性应设置为有效的视频链接。
1. 每个 `iframe` 元素还应有 `title` 属性描述嵌入内容，并有 `height` 和 `width` 属性设置合适的尺寸。

# --hints--

你应该在 `body` 元素内有一个 `main` 元素。

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

你的 `main` 元素应是 `body` 元素的唯一子元素。

```js
assert.equal(document.querySelector('body').children.length, 2);
```

你应该在 `main` 元素内有一个 `h1` 元素，作为页面主题。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你应该在 `h1` 元素下方有一个介绍页面主题的段落。

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

你应该在第一个 `p` 元素下方有三个 `section` 元素。

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

每个 `section` 元素应以 `h2` 元素作为该部分的标题开头。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

每个 `section` 元素应包含一个 `p` 元素作为第二个子元素，用于介绍视频内容。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

每个 `section` 元素应包含一个 `iframe` 元素作为第三个子元素。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

你的第一个 `iframe` 元素应有一个有效的视频 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

你的第一个 `iframe` 元素应有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

你的第一个 `iframe` 元素应有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

你的第一个 `iframe` 元素应有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

你的第二个 `iframe` 元素应有一个有效的视频 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

你的第二个 `iframe` 元素应有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

你的第二个 `iframe` 元素应有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

你的第二个 `iframe` 元素应有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

你的第三个 `iframe` 元素应有一个有效的视频 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

你的第三个 `iframe` 元素应有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

你的第三个 `iframe` 元素应有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

你的第三个 `iframe` 元素应有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
    <meta charset="utf-8">
    <title>视频合辑页面</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
    <meta charset="utf-8">
    <title>视频合辑页面</title>
</head>

<body>
    <main>
        <h1>前端网页开发</h1>
        <p>前端网页开发涉及使用 HTML、CSS 和 JavaScript 设计和构建网站的用户界面。它专注于创建视觉美观、交互性强且用户友好的网站。前端开发者确保网站的可访问性、性能优化以及遵循 Web 标准，将创意与技术专长相结合。</p>
        <section>
            <h2>HTML</h2>
            <p>HTML（超文本标记语言）构成网页的基础结构和布局。它通过一系列元素来展示文本、图片、列表、表格等静态内容。</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html 与编程入门视频" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS（层叠样式表）用于控制网页上 HTML 元素的视觉呈现。它支持自定义样式，包括颜色、字体、布局和间距，使网站更具吸引力。CSS 还支持响应式设计，确保网页在各种设备上都能良好显示。</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css 教程视频" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript 是一种为静态网页添加交互性的编程语言，实现动态内容更新、表单验证、动画效果以及基于用户操作和事件的响应行为。</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript 编程视频" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```

