---
id: 66c06d618d075c7f7f1b890a
title: 构建一个算命机
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

在本实验中，你将定义五个运势，并随机选择其中一个向用户展示。

请完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要初始化五个变量 `fortune1`、`fortune2`、`fortune3`、`fortune4` 和 `fortune5`，它们的值为你选择的字符串。你可以使用以下选项：

  - `"你的猫今天看起来非常可爱。"`
  - `"明天天气会很好。"`
  - `"要小心你的新邻居。"`
  - `"你很快会发现一个新爱好。"`
  - `"今天最好避免红色。"`

2. 你需要生成一个 1 到 5（包含 1 和 5）之间的随机数，并将其赋值给变量 `randomNumber`。

3. 你需要创建一个 `selectedFortune` 变量，并根据以下规则为其赋值：

  - 如果 `randomNumber` 是 1，将 `fortune1` 的值赋给 `selectedFortune`。
  - 如果 `randomNumber` 是 2，将 `fortune2` 的值赋给 `selectedFortune`。
  - 如果 `randomNumber` 是 3，将 `fortune3` 的值赋给 `selectedFortune`。
  - 如果 `randomNumber` 是 4，将 `fortune4` 的值赋给 `selectedFortune`。
  - 如果 `randomNumber` 是 5，将 `fortune5` 的值赋给 `selectedFortune`。

4. 你需要将 `selectedFortune` 输出到控制台。

# --hints--

你需要用字符串值初始化 `fortune1`。

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

你需要用字符串值初始化 `fortune2`。

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

你需要用字符串值初始化 `fortune3`。

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

你需要用字符串值初始化 `fortune4`。

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

你需要用字符串值初始化 `fortune5`。

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

你需要使用 `Math.random()` 方法生成一个随机数。

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

你需要生成一个 1 到 5 之间的随机数，并将其赋值给变量 `randomNumber`。

```js
assert.isNotNull(randomNumber);
assert.include([1,2,3,4,5], randomNumber);
```

你需要有一个 `selectedFortune` 变量，并根据 `randomNumber` 的值为其赋值。

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` 应该与其对应的运势相匹配。例如，如果 `randomNumber` 是 1，`selectedFortune` 应等于 `fortune1`，以此类推。

```js
const condition1 = randomNumber === 1 && selectedFortune === fortune1;
const condition2 = randomNumber === 2 && selectedFortune === fortune2;
const condition3 = randomNumber === 3 && selectedFortune === fortune3;
const condition4 = randomNumber === 4 && selectedFortune === fortune4;
const condition5 = randomNumber === 5 && selectedFortune === fortune5;

assert.isTrue(condition1 || condition2 || condition3 || condition4 || condition5);
```

你需要将 `selectedFortune` 输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "你的猫今天看起来非常可爱。";
const fortune2 = "明天天气会很好。";
const fortune3 = "要小心你的新邻居。";
const fortune4 = "你很快会发现一个新爱好。";
const fortune5 = "今天最好避免红色。";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```

