---
id: 66da326c02141df538f29ba5
title: 构建结账页面
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个文本为 `Checkout` 的 `h1` 元素。
1. 你应该在 `h1` 元素后紧跟两个 `section` 元素。
1. 你应该在第一个 section 内有一个文本为 `Your Cart` 的 `h2` 元素。
1. 你应该在第一个 section 内有一个商品图片，并包含合适的替代文本。你可以使用此图片：`https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 你应该在第二个 section 内有一个文本为 `Payment Information` 的 `h2` 元素。
1. 你应该在第二个 section 内有一个 `form` 元素。
1. 你应该在表单内有一个 `id` 和 `name` 都为 `card-name` 的输入框，并有与之关联的 `label`。
1. 你应该在表单内有一个 `id` 和 `name` 都为 `card-number` 的输入框，并有与之关联的 `label`。
1. 你的输入框中至少有两个是必填的，并且有 `aria-required` 属性且值为 `true`。

# --hints--

你应该有一个文本为 `Checkout` 的 `h1` 元素。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

你的页面上只能有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该在 `h1` 元素后至少有两个 `section` 元素。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

你应该在第一个 section 内有一个文本为 `Your Cart` 的 `h2` 元素。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

你应该在第一个 section 内至少有一张带有替代文本的图片。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

你应该在第二个 section 内有一个文本为 `Payment Information` 的 `h2` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

你应该在第二个 section 元素内有一个 `form` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

你应该在表单内有一个 `id` 和 `name` 都为 `card-name` 的输入框。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

你应该在表单内有一个 `id` 和 `name` 都为 `card-number` 的输入框。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

所有不是 `type="submit"` 的输入框都应该有与之关联的 `label` 元素。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

你的输入框中至少有两个是必填的，并且有 `aria-required` 属性且值为 `true`。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="number" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="number" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```

