---
id: 587d78b0367417b2b2512b05
title: 制作一个技术文档页面
challengeType: 25
demoType: onClick
dashedName: build-a-technical-documentation-page
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求:**

1. 你可以看到一个 `main` 元素，并且它有对应的 `id="main-doc"`，其中包含页面的主要内容（技术文档）。
1. 在 `#main-doc` 元素内，我能看见至少 5 个 `section` 元素，每个元素都有一个 class 为 `main-section`。至少要有五个。
1. 每个 `.main-section` 中的第一个元素应该是 `header` 元素，其中包含描述该部分主题的文本。
1. class 为 `main-section` 的每个 `section` 元素应该有一个与包含在其中的每个 `header` 的文本相对应的 `id`。任何空格都应替换为下划线（例如，包含标题 "JavaScript and Java" 的部分应有对应的 `id="JavaScript_and_Java"`）。
1. `.main-section` 元素总共应包含至少十个 `p` 元素（不是每个都要十个）。
1. `.main-section` 元素总共应包含至少五个 `code` 元素（不是每个都要五个）。
1. `.main-section` 元素总共应包含至少五个 `li` 项（不是每个都要五个）。
1. 你可以看到一个 `nav` 元素，并且它有对应的 `id="navbar"`。
1. 导航栏元素应包含一个 `header` 元素，其中包含描述技术文档主题的文本。
1. 此外，导航栏应包含 class 为 `nav-link` 的链接元素（`a`）。每个 class 为 `main-section` 的元素都应有一个对应的链接。
1. `#navbar` 中的 `header` 元素必须位于导航栏中所有链接（`a`）元素之前。
1. 每个 class 为 `nav-link` 的元素应包含与每个 `section` 内 `header` 文本相对应的文本（例如，如果你有一个 "Hello world" 部分/标题，你的导航栏应有一个包含文本 "Hello world" 的元素）。
1. 当你点击导航栏中的元素时，页面应跳转到 `#main-doc` 元素中对应的部分（例如，如果你点击包含 "Hello world" 文本的 `.nav-link` 元素，页面应跳转到具有该 id 的 `section` 元素，并包含相应的 header）。
1. 在常规尺寸设备（笔记本、台式机）上，`id="navbar"` 的元素应显示在屏幕左侧，并始终可见。
1. 你的技术文档应至少使用一个媒体查询。

**注意：** 请确保在 HTML 中链接你的样式表并应用 CSS。

# --hints--

你应该有一个 `id` 为 `main-doc` 的 `main` 元素。

```js
const el = document.getElementById('main-doc');
assert.isNotNull(el);
```

你至少应该有 5 个 class 为 `main-section` 的 `section` 元素。

```js
const els = document.querySelectorAll('#main-doc section');
assert.isAtLeast(els.length, 5);
```

你所有的 `.main-section` 元素都应该是 `section` 元素。

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (el.tagName !== 'SECTION') {
    assert.fail();
  }
});
assert.isAbove(els.length, 0);
```

你至少应该有 5 个 `.main-section` 元素，它们是 `#main-doc` 的子元素。

```js
const els = document.querySelectorAll('#main-doc .main-section');
assert.isAtLeast(els.length, 5);
```

每个 `.main-section` 的第一个子元素都应该是一个 `header` 元素。

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (el.firstElementChild?.tagName !== 'HEADER') assert.fail();
});
assert.isNotEmpty(els);
```

你的 `header` 元素不应为空。

```js
const els = document.querySelectorAll('header');
els.forEach(el => {
  if (el.innerText?.length <= 0) assert.fail();
});
assert.isNotEmpty(els);
```

你所有的 `.main-section` 元素都应该有 `id`。

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (!el.id || el.id === '') assert.fail();
});
assert.isNotEmpty(els);
```

每个 `.main-section` 都应该有一个与其第一个子元素的文本匹配的 `id`，把子元素的文本中的空格都替换为下划线（`_`）用于 id。

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  const text = el.firstElementChild?.innerText?.replaceAll(' ', '_');
  if (el.id?.toUpperCase() !== text?.toUpperCase()) assert.fail();
});
assert.isNotEmpty(els);
```

在你的 `.main-section` 元素中总计应有至少 10 个 `p` 元素

```js
const els = document.querySelectorAll('.main-section p');
assert.isAtLeast(els.length, 10);
```

所有 `.main-section` 元素内总计应有至少 5 个 `code` 元素。

```js
const els = document.querySelectorAll('.main-section code');
assert.isAtLeast(els.length, 5);
```

所有 `.main-section` 元素内总计应有至少 5 个 `li` 元素。

```js
const els = document.querySelectorAll('.main-section li');
assert.isAtLeast(els.length, 5);
```

你应该有一个 `id` 为 `navbar` 的 `nav` 元素。

```js
const el = document.getElementById('navbar');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'NAV');
```

你的 `#navbar` 应该只有一个 `header` 元素。

```js
const els = document.querySelectorAll('#navbar header');
assert.lengthOf(els, 1);
```

你应该至少有一个 class 为 `nav-link` 的 `a` 元素。

```js
const els = document.querySelectorAll('a.nav-link');
assert.isAtLeast(els.length, 1);
```

你所有的 `.nav-link` 元素都应该是锚点（`a`）元素。

```js
const els = document.querySelectorAll('.nav-link');
els.forEach(el => {
  if (el.tagName !== 'A') assert.fail();
});
assert.isNotEmpty(els);
```

你所有的 `.nav-link` 元素都应该在 `#navbar` 中。

```js
const els1 = document.querySelectorAll('.nav-link');
const els2 = document.querySelectorAll('#navbar .nav-link');
assert.isNotEmpty(els2);
assert.strictEqual(els1.length, els2.length);
```

你应该有相同数量的 `.nav-link` 和 `.main-section` 元素。

```js
const els1 = document.querySelectorAll('.main-section');
const els2 = document.querySelectorAll('.nav-link');
assert.isNotEmpty(els1);
assert.isNotEmpty(els2);
assert.strictEqual(els1.length, els2.length);
```

`#navbar` 中的 `header` 元素必须在 `#navbar` 中的任何链接（`a`）之前。

```js
const navLinks = document.querySelectorAll('#navbar a.nav-link');
const header = document.querySelector('#navbar header');
navLinks.forEach(navLink => {
  if (
    header.compareDocumentPosition(navLink) & Node.DOCUMENT_POSITION_PRECEDING
  )
    assert.fail();
});
assert.isNotNull(header);
```

每个 `.nav-link` 应该有与其相关 `section` 的 `header` 文本相对应的文本（例如，如果你有一个 “Hello world” 部分/标题，你的 `#navbar` 应该有一个 `.nav-link` 包含文本 “Hello world”）。

```js
const headerText = Array.from(document.querySelectorAll('.main-section')).map(
  el => el.firstElementChild?.innerText?.trim().toUpperCase()
);
const linkText = Array.from(document.querySelectorAll('.nav-link')).map(el =>
  el.innerText?.trim().toUpperCase()
);
const remainder = headerText.filter(str => linkText.indexOf(str) === -1);

assert.isNotEmpty(headerText);
assert.isNotEmpty(linkText);
assert.isEmpty(remainder);
```

每个 `.nav-link` 都应该有一个 `href` 属性，该属性链接到其对应的 `.main-section`（例如，如果你单击包含文本 “Hello world” 的 `.nav-link` 元素，页面导航到具有该 id 的 `section` 元素）。

```js
const hrefValues = Array.from(document.querySelectorAll('.nav-link')).map(el =>
  el.getAttribute('href')
);
const mainSectionIDs = Array.from(
  document.querySelectorAll('.main-section')
).map(el => el.id);
const missingHrefValues = mainSectionIDs.filter(
  str => hrefValues.indexOf('#' + str) === -1
);
assert.isNotEmpty(hrefValues);
assert.isNotEmpty(mainSectionIDs);
assert.isEmpty(missingHrefValues, 0);
```

你的 `#navbar` 元素应该始终位于视口的顶部。

```js
const el = document.getElementById('navbar');
const left1 = el?.offsetLeft;
assert.isNotNull(el);
assert.isAtLeast(left1, -15);
assert.isAtMost(left1, 15);
```

你的技术文档项目应该使用至少一个媒体查询。

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el =>
  el.getAttribute('media')
);
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh">
  <head>
    <meta charset="UTF-8" />
    <title>技术文档页面</title>
  </head>

  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <nav id="navbar">
      <header>JS 文档</header>
      <ul>
        <li><a class="nav-link" href="#Introduction">Introduction</a></li>
        <li>
          <a class="nav-link" href="#What_you_should_already_know"
            >What you should already know</a
          >
        </li>
        <li>
          <a class="nav-link" href="#JavaScript_and_Java"
            >JavaScript and Java</a
          >
        </li>
        <li><a class="nav-link" href="#Hello_world">Hello world</a></li>
        <li><a class="nav-link" href="#Variables">Variables</a></li>
        <li>
          <a class="nav-link" href="#Declaring_variables"
            >Declaring variables</a
          >
        </li>
        <li><a class="nav-link" href="#Variable_scope">Variable scope</a></li>
        <li>
          <a class="nav-link" href="#Global_variables">Global variables</a>
        </li>
        <li><a class="nav-link" href="#Constants">Constants</a></li>
        <li><a class="nav-link" href="#Data_types">Data types</a></li>
        <li>
          <a class="nav-link" href="#if...else_statement"
            >if...else statement</a
          >
        </li>
        <li><a class="nav-link" href="#while_statement">while statement</a></li>
        <li>
          <a class="nav-link" href="#Function_declarations"
            >Function declarations</a
          >
        </li>
        <li><a class="nav-link" href="#Reference">Reference</a></li>
      </ul>
    </nav>
    <main id="main-doc">
      <section class="main-section" id="Introduction">
        <header>Introduction</header>
        <article>
          <p>
            JavaScript 是一种跨平台、面向对象的脚本语言。它是一种小巧且轻量级的语言。在主机环境（例如 Web 浏览器）中，JavaScript 可以连接到其环境的对象，以实现对它们的编程控制。
          </p>

          <p>
            JavaScript 包含标准对象库，如 Array、Date 和 Math，以及一组核心语言元素，如运算符、控制结构和语句。核心 JavaScript 可以通过补充额外对象来扩展以实现多种用途，例如：
          </p>
          <ul>
            <li>
              客户端 JavaScript 通过提供控制浏览器及其文档对象模型（DOM）的对象来扩展核心语言。例如，客户端扩展允许应用在 HTML 表单上放置元素并响应用户事件，如鼠标点击、表单输入和页面导航。
            </li>
            <li>
              服务端 JavaScript 通过提供与在服务器上运行 JavaScript 相关的对象来扩展核心语言。例如，服务端扩展允许应用与数据库通信、在应用多次调用间保持信息连续性，或在服务器上执行文件操作。
            </li>
          </ul>
        </article>
      </section>
      <section class="main-section" id="What_you_should_already_know">
        <header>What you should already know</header>
        <article>
          <p>本指南假设你具备以下基础知识：</p>

          <ul>
            <li>
              对互联网和万维网（WWW）有基本了解。
            </li>
            <li>熟练掌握超文本标记语言（HTML）。</li>
            <li>
              有一定的编程经验。如果你是编程新手，请尝试主页面上关于 JavaScript 的教程链接。
            </li>
          </ul>
        </article>
      </section>
      <section class="main-section" id="JavaScript_and_Java">
        <header>JavaScript and Java</header>
        <article>
          <p>
            JavaScript 和 Java 在某些方面相似，但在其他方面有根本区别。JavaScript 语言类似于 Java，但没有 Java 的静态类型和强类型检查。JavaScript 遵循大多数 Java 表达式语法、命名约定和基本控制流结构，这也是它从 LiveScript 更名为 JavaScript 的原因。
          </p>

          <p>
            与 Java 的基于声明的编译时类系统不同，JavaScript 支持基于少量数据类型的运行时系统，这些数据类型表示数字、布尔值和字符串。JavaScript 采用基于原型的对象模型，而不是更常见的基于类的对象模型。原型模型提供了动态继承，即继承内容可因对象而异。JavaScript 还支持无需特殊声明的函数。函数可以作为对象的属性，作为松散类型的方法执行。
          </p>
          <p>
            与 Java 相比，JavaScript 是一种非常自由的语言。你不必声明所有变量、类和方法。不必关心方法是 public、private 还是 protected，也不必实现接口。变量、参数和函数返回类型都不是显式类型。
          </p>
        </article>
      </section>
      <section class="main-section" id="Hello_world">
        <header>Hello world</header>
        <article>
          要开始编写 JavaScript，请打开 Scratchpad 并编写你的第一个 "Hello world" JavaScript 代码：
          <code
            >function greetMe(yourName) { alert("Hello " + yourName); }
            greetMe("World");
          </code>

          选中代码并按 Ctrl+R，即可在浏览器中看到效果！
        </article>
      </section>
      <section class="main-section" id="Variables">
        <header>Variables</header>
        <p>
          你可以使用变量作为应用中值的符号名称。变量名（即标识符）需遵循一定规则。
        </p>
        <p>
          JavaScript 标识符必须以字母、下划线（_）或美元符号（$）开头；后续字符也可以是数字（0-9）。由于 JavaScript 区分大小写，字母包括 "A" 到 "Z"（大写）和 "a" 到 "z"（小写）。
        </p>
        <p>
          你可以在标识符中使用 ISO 8859-1 或 Unicode 字母，如 å 和 ü，也可以使用 Unicode 转义序列作为标识符字符。合法名称示例：Number_hits、temp99、_name。
        </p>
      </section>
      <section class="main-section" id="Declaring_variables">
        <header>Declaring variables</header>
        <article>
          你可以通过三种方式声明变量：
          <p>
            使用关键字 var。例如，<code>var x = 42.</code> 这种语法可用于声明局部和全局变量。
          </p>
          <p>
            直接赋值。例如，<code>x = 42.</code> 这总是声明一个全局变量，会产生严格的 JavaScript 警告。你不应使用这种方式。
          </p>
          <p>
            使用关键字 let。例如，<code>let y = 13.</code> 这种语法可用于声明块作用域的局部变量。详见下文的变量作用域。
          </p>
        </article>
      </section>
      <section class="main-section" id="Variable_scope">
        <header>Variable scope</header>
        <article>
          <p>
            当你在任何函数外声明变量时，它被称为全局变量，因为它对当前文档中的任何其他代码都可用。当你在函数内声明变量时，它被称为局部变量，仅在该函数内可用。
          </p>

          <p>
            ECMAScript 2015 之前的 JavaScript 没有块级作用域；在块内声明的变量属于所在函数（或全局作用域）。例如，以下代码会输出 5，因为 x 的作用域是声明它的函数（或全局上下文），而不是 if 语句块。
          </p>
          <code>if (true) { var x = 5; } console.log(x); // 5</code>
          <p>
            使用 ECMAScript 2015 引入的 let 声明时，这一行为会改变。
          </p>

          <code
            >if (true) { let y = 5; } console.log(y); // ReferenceError: y is
            not defined</code
          >
        </article>
      </section>
      <section class="main-section" id="Global_variables">
        <header>Global variables</header>
        <article>
          <p>
            全局变量实际上是全局对象的属性。在网页中，全局对象是 window，因此你可以通过 window.variable 语法设置和访问全局变量。
          </p>

          <p>
            因此，你可以通过指定窗口或框架名称，从另一个窗口或框架访问在一个窗口或框架中声明的全局变量。例如，如果在文档中声明了变量 phoneNumber，你可以在 iframe 中通过 parent.phoneNumber 访问该变量。
          </p>
        </article>
      </section>
      <section class="main-section" id="Constants">
        <header>Constants</header>
        <article>
          <p>
            你可以使用 const 关键字创建只读的命名常量。常量标识符的语法与变量标识符相同：必须以字母、下划线或美元符号开头，可以包含字母、数字或下划线字符。
          </p>

          <code>const PI = 3.14;</code>
          <p>
            常量在脚本运行期间不能通过赋值改变值，也不能被重新声明。必须初始化为一个值。
          </p>

          <p>
            常量的作用域规则与 let 块级作用域变量相同。如果省略 const 关键字，则该标识符被视为变量。
          </p>

          <p>
            你不能在同一作用域内用与函数或变量相同的名称声明常量。例如：
          </p>

          <code
            >// 这会导致错误 function f() {}; const f = 5; // 这也会导致错误 function f() { const g = 5; var g;
            //statements }</code
          >
          但对象属性不受保护，因此以下语句可以正常执行。
          <code
            >const MY_OBJECT = {"key": "value"}; MY_OBJECT.key =
            "otherValue";</code
          >
        </article>
      </section>
      <section class="main-section" id="Data_types">
        <header>Data types</header>
        <article>
          <p>最新的 ECMAScript 标准定义了七种数据类型：</p>
          <ul>
            <li>
              <p>六种原始数据类型：</p>
              <ul>
                <li>Boolean。true 和 false。</li>
                <li>
                  null。表示空值的特殊关键字。由于 JavaScript 区分大小写，null 不等于 Null、NULL 或其他变体。
                </li>
                <li>
                  undefined。顶级属性，其值为 undefined。
                </li>
                <li>Number。42 或 3.14159。</li>
                <li>String。"Howdy"</li>
                <li>
                  Symbol（ECMAScript 2015 新增）。其实例是唯一且不可变的数据类型。
                </li>
              </ul>
            </li>

            <li>以及 Object</li>
          </ul>
          虽然这些数据类型数量较少，但它们使你的应用能够执行有用的功能。对象和函数是该语言的其他基本元素。你可以将对象视为值的命名容器，将函数视为应用可以执行的过程。
        </article>
      </section>
      <section class="main-section" id="if...else_statement">
        <header>if...else statement</header>
        <article>
          使用 if 语句在逻辑条件为 true 时执行语句。使用可选的 else 子句在条件为 false 时执行语句。if 语句如下所示：

          <code>if (condition) { statement_1; } else { statement_2; }</code>
          condition 可以是任何求值为 true 或 false 的表达式。详见 Boolean 了解哪些值为 true 或 false。如果 condition 为 true，则执行 statement_1；否则执行 statement_2。statement_1 和 statement_2 可以是任何语句，包括嵌套的 if 语句。
          <p>
            你还可以使用 else if 复合语句，按顺序测试多个条件，如下所示：
          </p>
          <code
            >if (condition_1) { statement_1; } else if (condition_2) {
            statement_2; } else if (condition_n) { statement_n; } else {
            statement_last; }
          </code>
          多条件时，只有第一个求值为 true 的条件会被执行。要执行多条语句，请用块语句（{ ... }）包裹。通常，尤其是嵌套 if 语句时，建议总是使用块语句：

          <code
            >if (condition) { statement_1_runs_if_condition_is_true;
            statement_2_runs_if_condition_is_true; } else {
            statement_3_runs_if_condition_is_false;
            statement_4_runs_if_condition_is_false; }</code
          >
          建议不要在条件表达式中使用简单赋值，因为赋值可能与等号混淆。例如，不要使用如下代码：
          <code>if (x = y) { /* statements here */ }</code> 如果需要在条件表达式中使用赋值，常见做法是为赋值加额外括号。例如：

          <code>if ((x = y)) { /* statements here */ }</code>
        </article>
      </section>
      <section class="main-section" id="while_statement">
        <header>while statement</header>
        <article>
          while 语句在指定条件为 true 时执行其语句。while 语句如下所示：

          <code>while (condition) statement</code> 如果条件变为 false，循环内的语句停止执行，控制权转移到循环后的语句。

          <p>
            条件测试发生在循环语句执行之前。如果条件为 true，则执行语句并再次测试条件。如果条件为 false，则停止执行，控制权转移到 while 后的语句。
          </p>

          <p>
            要执行多条语句，请使用块语句（{ ... }）将这些语句分组。
          </p>

          示例：

          <p>
            以下 while 循环在 n 小于 3 时迭代：
          </p>

          <code>var n = 0; var x = 0; while (n &lt; 3) { n++; x += n; }</code>
          <p>
            每次迭代，循环递增 n 并将该值加到 x。因此，x 和 n 的值如下：
          </p>

          <ul>
            <li>第一次循环后：n = 1，x = 1</li>
            <li>第二次循环后：n = 2，x = 3</li>
            <li>第三次循环后：n = 3，x = 6</li>
          </ul>
          <p>
            完成第三次循环后，n &lt; 3 条件不再成立，循环终止。
          </p>
        </article>
      </section>
      <section class="main-section" id="Function_declarations">
        <header>Function declarations</header>
        <article>
          函数定义（也称为函数声明或函数语句）由 function 关键字后跟以下内容组成：

          <ul>
            <li>函数名。</li>
            <li>
              括号内以逗号分隔的参数列表。
            </li>
            <li>
              用大括号 { } 包裹的 JavaScript 语句，定义函数体。
            </li>
          </ul>
          <p>
            例如，以下代码定义了一个名为 square 的简单函数：
          </p>

          <code>function square(number) { return number * number; }</code>
          <p>
            函数 square 接收一个参数 number。函数体只有一条语句，返回参数自身的平方。return 语句指定函数返回的值。
          </p>
          <code>return number * number;</code>
          <p>
            原始类型参数（如数字）以值传递给函数；值传递给函数后，如果函数改变参数值，这一更改不会在全局或调用函数中反映。
          </p>
        </article>
      </section>
      <section class="main-section" id="Reference">
        <header>Reference</header>
        <article>
          <ul>
            <li>
              本页面所有文档内容均摘自
              <a
                href="https://developer.mozilla.org/en-US/docs/Web/JavaScript/Guide"
                target="_blank"
                >MDN</a
              >
            </li>
          </ul>
        </article>
      </section>
    </main>
  </body>
</html>
```

```css
html,
body {
  min-width: 290px;
  color: #4d4e53;
  background-color: #ffffff;
  font-family: 'Open Sans', Arial, sans-serif;
  line-height: 1.5;
}

#navbar {
  position: fixed;
  min-width: 290px;
  top: 0px;
  left: 0px;
  width: 300px;
  height: 100%;
  border-right: solid;
  border-color: rgba(0, 22, 22, 0.4);
}

header {
  color: black;
  margin: 10px;
  text-align: center;
  font-size: 1.8em;
  font-weight: thin;
}

#main-doc header {
  text-align: left;
  margin: 0px;
}

#navbar ul {
  height: 88%;
  padding: 0;
  overflow-y: auto;
  overflow-x: hidden;
}

#navbar li {
  color: #4d4e53;
  border-top: 1px solid;
  list-style: none;
  position: relative;
  width: 100%;
}

#navbar a {
  display: block;
  padding: 10px 30px;
  color: #4d4e53;
  text-decoration: none;
  cursor: pointer;
}

#main-doc {
  position: absolute;
  margin-left: 310px;
  padding: 20px;
  margin-bottom: 110px;
}

section article {
  color: #4d4e53;
  margin: 15px;
  font-size: 0.96em;
}

section li {
  margin: 15px 0px 0px 20px;
}

code {
  display: block;
  text-align: left;
  white-space: pre-line;
  position: relative;
  word-break: normal;
  word-wrap: normal;
  line-height: 2;
  background-color: #f7f7f7;
  padding: 15px;
  margin: 10px;
  border-radius: 5px;
}

@media only screen and (max-width: 815px) {
  /* 针对手机： */
  #navbar ul {
    border: 1px solid;
    height: 207px;
  }

  #navbar {
    background-color: white;
    position: absolute;
    top: 0;
    padding: 0;
    margin: 0;
    width: 100%;
    max-height: 275px;
    border: none;
    z-index: 1;
    border-bottom: 2px solid;
  }

  #main-doc {
    position: relative;
    margin-left: 0px;
    margin-top: 270px;
  }
}

@media only screen and (max-width: 400px) {
  #main-doc {
    margin-left: -10px;
  }

  code {
    margin-left: -20px;
    width: 100%;
    padding: 15px;
    padding-left: 10px;
    padding-right: 45px;
    min-width: 233px;
  }
}
```

