---
id: 66ba762af611169359d9d369
title: 构建一篇报纸文章
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你需要将 HTML 文档的根 `font-size` 设置为 `24px`。
1. 你需要有一个类名为 `newspaper` 的元素，包含所有其他元素。
1. 你的 `.newspaper` 元素应设置 `font-size` 为 `16px`，字体为 `Open Sans`，后备字体为 `sans-serif`。
1. 在 `.newspaper` 元素内，你至少需要有七个元素：一个类名为 `name` 的报纸名称元素、一个类名为 `date` 的文章日期元素、一个类名为 `headline` 的标题元素、一个类名为 `sub-headline` 的副标题元素、一个类名为 `author` 的作者元素，以及两个类名为 `text` 的段落元素。所有这些元素都应填写你的文章信息。
1. 你的 `.name` 元素的 `font-size` 应为根元素的两倍，并使用 `Times New Roman` 字体，后备字体为 `serif`。
1. 你的 `.name` 和 `.author` 元素应通过 CSS 设置为全部大写字母。
1. 你的 `.headline` 元素的 `font-size` 应为其父元素的两倍，并加粗显示。
1. 你的 `.sub-headline` 元素应设置 `font-weight` 为 `100`，`font-size` 为其父元素的 1.5 倍，并为斜体。
1. 你的 `.author` 元素应通过 CSS 设置为加粗。
1. 你的 `.text` 元素应设置 `text-indent` 为 `20px`。
1. 你的 `.text` 元素应设置 `line-height` 为其父元素字体大小的两倍。
1. 你的 `.text` 元素的首字母应加粗且为其父元素字体大小的两倍。请使用 `::first-letter` 选择器实现。

# --hints--

你应有一个 `html` 选择器，将其 `font-size` 设置为 `24px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

你应有一个类名为 `newspaper` 的元素。

```js
assert.exists(document.querySelector('.newspaper'));
```

你应有一个 `.newspaper` 选择器，将其元素的 `font-size` 设置为 `16px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

你应有一个 `.newspaper` 选择器，将其元素的 `font-family` 设置为 `'Open Sans', sans-serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

你应在 `.newspaper` 元素内有一个类名为 `name` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

你的 `.name` 元素应包含报纸名称。

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

你应在 `.newspaper` 元素内有一个类名为 `date` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

你的 `.date` 元素应包含日期。

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

你应在 `.newspaper` 元素内有一个类名为 `headline` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

你的 `.headline` 元素应包含文章标题。

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

你应在 `.newspaper` 元素内有一个类名为 `sub-headline` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

你的 `.sub-headline` 元素应包含文章副标题。

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

你应在 `.newspaper` 元素内有一个类名为 `author` 的元素。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

你的 `.author` 元素应包含文章作者。

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

你应在 `.newspaper` 元素内至少有两个类名为 `text` 的段落元素。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

你的 `.text` 元素应包含文章正文。

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

你应有一个 `.name` 选择器，将其元素的 `font-size` 设置为 `2rem`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

你应有一个 `.name` 选择器，将其元素的 `font-family` 设置为 `'Times New Roman', serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

`.name` 元素应设置 `text-transform` 属性为 `uppercase`。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.author` 元素应设置 `text-transform` 属性为 `uppercase`。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

你应有一个 `.headline` 选择器，将其元素的 `font-size` 设置为 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

你应有一个 `.headline` 选择器，将其元素的 `font-weight` 设置为 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

你应有一个 `.sub-headline` 选择器，将其元素的 `font-weight` 设置为 `100`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

你应有一个 `.sub-headline` 选择器，将其元素的 `font-size` 设置为 `1.5em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

你应有一个 `.sub-headline` 选择器，将其元素的 `font-style` 设置为 `italic`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

你应有一个 `.author` 选择器，将其元素的 `font-weight` 设置为 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

你应有一个 `.text` 选择器，将其元素的 `text-indent` 设置为 `20px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

你应有一个 `.text` 选择器，将其元素的 `line-height` 设置为 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

你应有一个 `.text::first-letter` 选择器，将其元素的 `font-weight` 设置为 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

你应有一个 `.text::first-letter` 选择器，将其元素的 `font-size` 设置为 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>每日趣闻</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">每日趣闻</h1>
    <p class="date">2024年7月5日</p>
    <h2 class="headline">突发：外婆艾德娜拯救地球</h2>
    <h3 class="sub-headline">外星人入侵被技术达人外婆的 Wi-Fi 密码击退</h3>
    <p class="author">作者：简·多伊</p>
    <p class="text">在命运的奇妙转折下，外婆艾德娜因其巧妙的 Wi-Fi 安全措施，意外挫败了一场外星人入侵。当外星飞船降临小镇，恐慌蔓延时，艾德娜外婆镇定自若，重置了只有她自己知道的神秘路由器密码。</p>
    <p class="text">外星人面对意想不到的技术壁垒，始终无法攻破外婆的网络防线。最终，他们沮丧而困惑地撤退，准备重新评估入侵策略，镇上的居民对外婆的机智既感到宽慰又惊叹不已。</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```

