---
id: 673b567e3ba535dda140d278
title: 构建一个投票系统
challengeType: 26
dashedName: build-a-voting-system
---

# --description--

在本实验中，你将构建一个投票系统，使用 `Map` 创建投票选项，并用 `Set` 防止重复投票。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你应该将 `poll` 变量初始化为一个新的 `Map` 对象。

2. 你应该有一个名为 `addOption` 的函数，接收一个参数 `option`。

3. 在 `addOption` 函数中：

  - 如果 `option` 在投票中不存在，应将其添加到投票中，并以一个空的 `Set` 作为其值用于追踪投票者。你还应返回消息 `Option "<option>" added to the poll.`

  - 如果 `option` 已存在，应返回消息 `Option "<option>" already exists.`

  - 如果尝试添加空选项，函数应返回消息 `Option cannot be empty.`

4. 你应该有一个名为 `vote` 的函数，接收两个参数，`option`（要投票的选项）和 `voterId`（投票者的唯一ID）。

5. 在 `vote` 函数中：

  - 如果 `option` 不存在于投票中，函数应返回消息 `Option "<option>" does not exist.`

  - 如果 `option` 存在，函数应检查该 `voterId` 是否已经为此选项投过票。

  - 如果该投票者已经投过票，应显示消息 `Voter <voterId> has already voted for "<option>".`

  - 如果该投票者尚未投票，应将 `voterId` 添加到该选项的 `Set` 中，并返回消息 `Voter <voterId> voted for "<option>".`

6. 你的 `poll` 中应至少有三个选项。

7. 你的 `poll` 中应至少有三张选票。

8. 你应该有一个名为 `displayResults` 的函数，以如下格式返回投票结果：

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
示例输出

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

你应该有一个初始化为新 `Map` 对象的 `poll` 变量。

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

你应该定义一个接收参数 `option` 的 `addOption` 函数。

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

你应该定义一个接收两个参数 `option` 和 `voterId` 的 `vote` 函数。

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

你应该定义一个 `displayResults` 函数用于显示投票结果。

```js
assert.isFunction(displayResults);
```

你的 `poll` 中应至少有三个选项。

```js
assert.lengthOf(poll, 3);
```

你的 `poll` 中应至少有三张选票。

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

你应确保每个投票选项都映射到一个 `Set` 对象。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

`addOption("Egypt")` 应返回 `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

添加空选项应返回 `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

当 `Turkey` 已添加时，`addOption("Turkey")` 应返回 `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

当投票选项中存在 `Malysia` 时，`vote("Malysia", "traveler1")` 应返回 `Voter traveler1 voted for "Malysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malysia');
  assert.equal(
   vote('Malysia', 'traveler1'),
   'Voter traveler1 voted for "Malysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

`vote` 应更新某个选项的投票者 `Set`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

当 `traveler1` 再次为 `Algeria` 投票时，`vote("Algeria", "traveler1")` 应返回 `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
   vote('Algeria', 'traveler1'),
   'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

重复投票不应增加 `Set` 的大小。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

当 `Nigeria` 不在投票选项中时，`vote("Nigeria", "traveler2")` 应返回 `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
   vote('Nigeria', 'traveler2'),
   'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

一个唯一选项应能收到多张选票。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

`displayResults()` 应以正确格式返回结果。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
   displayResults(),
   'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
   poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
   return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
   poll.set(option, new Set());
   return `Option "${option}" added to the poll.`;
  } else {
   return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
   return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
   return `Voter ${voterId} has already voted for "${option}".`;
  } else {
   voters.add(voterId);
   return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
   results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```

