---
id: 66faac4139dbbd5dd632c860
title: 构建一个求职申请表单
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个带有 `container` 类的 `div` 元素。
1. 在该 `div` 元素内，你应该有一个 `form` 表单元素。
1. 表单中应包含一个类型为 `text`、id 为 `name` 的 `input` 元素，用于输入用户的全名。
1. 你还应有一个类型为 `email`、id 为 `email` 的 `input` 元素，用于输入用户的电子邮箱地址。
1. 表单应包含一个 id 为 `position` 的 `select` 元素，允许用户选择职位。
1. 你应该有一个带有 `radio-group` 类的 `fieldset` 元素。
1. 在 `.radio-group` 内应有一组类型为 `radio` 的 `input` 元素及相关标签，用于选择可用性选项（如全职、兼职）。该组的 `name` 应为 `availability`。
1. 你应该有一个 id 为 `message` 的 `textarea` 元素，用于输入留言。
1. 你应该有一个类型为 `submit` 的 `button` 元素，用于提交表单。
1. 为 `input` 和 `textarea` 元素添加 `:focus` 伪类，在获得焦点时改变其边框颜色。
1. `input`、`select` 和 `textarea` 元素应有 `:invalid` 伪类，在输入无效时将边框颜色变为红色。
1. `input`、`select` 和 `textarea` 元素应有 `:valid` 伪类，在输入有效时将边框颜色变为绿色。
1. `button` 元素应有 `:hover` 伪类，在悬停时改变背景色。
1. 对 `.radio-group input[type="radio"]:checked` 使用 `:checked` 伪类，在单选按钮被选中时添加边框颜色、背景色和阴影。
1. 对被选中的单选按钮，使用 `:checked` 伪类改变其关联 `label` 的文本颜色。
1. 对 `input` 元素使用 `:first-of-type` 伪类，为第一个输入字段添加不同的样式（如圆角）。

**注意：** 请确保在 HTML 中正确链接样式表并应用 CSS。

# --hints--

你应该有一个带有 `container` 类的 `div` 元素。

```js
assert.exists(document.querySelector("div.container"));
```

在该 `div` 元素内，你应该有一个 `form` 表单元素。

```js
assert.exists(document.querySelector("div.container > form"));
```

表单中应包含一个类型为 `text`、id 为 `name` 的 `input` 元素，用于输入用户的全名。

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

你还应有一个类型为 `email`、id 为 `email` 的 `input` 元素，用于输入用户的电子邮箱地址。

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

表单应包含一个 id 为 `position` 的 `select` 元素，并包含一些 `option` 选项。

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

你应该有一个带有 `.radio-group` 类的 `fieldset` 或 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

在 `.radio-group` 内应有一组类型为 `radio` 的 `input` 元素用于选择可用性选项。该组的 `name` 应为 `availability`。

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

你应该有一个 id 为 `message` 的 `textarea` 元素，用于输入留言。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

你应该有一个类型为 `submit` 的 `button` 元素，用于提交表单。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

你应该为 `input` 和 `textarea` 元素添加 `:focus` 伪类，在获得焦点时改变其边框颜色。请按给定顺序使用列表选择器。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

`input`、`select` 和 `textarea` 元素应有 `:invalid` 伪类，在输入无效时将边框颜色变为红色。请按给定顺序使用列表选择器。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

`input`、`select` 和 `textarea` 元素应有 `:valid` 伪类，在输入有效时将边框颜色变为绿色。请按给定顺序使用列表选择器。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

`button` 元素应有 `:hover` 伪类，在悬停时改变背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

你应对 `.radio-group input[type="radio"]:checked` 使用 `:checked` 伪类，在单选按钮被选中时添加边框颜色、背景色和阴影。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
assert.isNotEmpty(style.getPropertyValue('background-color'));
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

你应对被选中的单选按钮，使用 `:checked` 伪类改变其关联 `label` 的文本颜色。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

对 `input` 元素使用 `:first-of-type` 伪类，为第一个输入字段添加不同的样式。

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```

