---
id: 66b62d0ad68488dd76228d6c
title: 调试一个随机背景色切换器
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot 正在学习 JavaScript，并尝试构建自己的随机背景色切换器。但在过程中出现了一些错误。

请完成以下用户需求，并通过所有测试，使实验室功能正常。

**用户需求：**

1. 你需要修复 `darkColorsArr` 数组中的语法错误。
1. 你需要修复在 `getRandomIndex` 函数中引用 `math` 对象时抛出的 `ReferenceError`。
1. 你需要将 `randomIndex` 向下取整为最接近的整数。
1. 你需要更新 `body` 变量的赋值，使用正确的方法名，并修复相关的 `TypeError`。
1. 你需要修复 `bgHexCodeSpanElement`，使其能正确选择 `id` 为 `bg-hex-code` 的元素。
1. 你需要修改 `changeBackgroundColor` 函数中的 `color` 变量，使其被设置为 `darkColorsArr` 数组中的一个随机颜色。
1. 你需要更新 `btn` 变量，使其选中按钮元素。请打开 `index.html` 查看该按钮元素的正确 `id` 名称。
1. 你需要修复为 `btn` 添加 `click` 事件监听器的代码，传递 `changeBackgroundColor` 函数的引用给该方法。
1. 每次点击 `#btn` 时，背景应随机切换为 `darkColorsArr` 中的某个颜色。

# --hints--

你需要解决 `darkColorsArr` 数组中的语法错误。记住数组元素之间要用逗号分隔。

```js
assert.isArray(darkColorsArr);
```

你需要修复 `math.random()` 行中的大小写错误。

```js
assert.match(getRandomIndex.toString(), /\s*darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)\s*/);
```

你需要将 `darkColorsArr.length * Math.random()` 向下取整为最接近的整数。

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

你需要更新 `body` 变量的赋值，使 `queryselector` 使用驼峰命名法。

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

你需要修复 `"bg-hex-code"` 的引用，使其正确指向 `bg-hex-code` 的 `id`。记住 CSS 中的 `id` 选择器以井号（`#`）开头。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

你需要在 `darkColorsArr[getRandomIndex]` 这一行中调用 `getRandomIndex` 函数。

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

你需要修正 `"#click-btn"` 这一行的 `id` 名称，使其与 `index.html` 文件中的正确 `id` 名称一致。

```js
assert.match(btn?.id, /btn/);
```

你不应调用 `changeBackgroundColor` 函数，而应将函数引用传递给 `addEventListener` 方法。

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];

  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())

  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)

  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];

    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)

  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```

