---
id: 66b205e6eacba4c4e54ea434
title: 构建一个邮箱掩码器
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

在本实验中，你将用星号掩盖邮箱地址的用户名部分。掩码是一种用星号或其他字符隐藏或替换敏感信息的方式。

例如，如果邮箱地址是 `myEmail@email.com`，那么掩码后的邮箱地址将是 `m*****l@email.com`。

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 创建一个名为 `maskEmail` 的函数，接收 `email` 作为参数。
2. 在函数内部，你需要对 `email` 进行掩码处理，并将域名部分拼接上。你可以使用如 `slice`、`repeat`、`indexOf` 或 `replace` 等方法来帮助实现。
3. 在函数外部，声明一个名为 `email` 的变量，用于存储你想要掩码的邮箱地址。
4. 使用 `email` 变量调用 `maskEmail` 函数，并将结果输出到控制台。
5. `maskEmail("apple.pie@example.com")` 应返回 `"a*******e@example.com"`。
6. `maskEmail("freecodecamp@example.com")` 应返回 `"f**********p@example.com"`。

# --hints--

你应该定义一个名为 `maskEmail` 的函数。

```js
assert.isFunction(maskEmail);
```

`maskEmail` 函数应接收一个字符串类型的参数 `email`。

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

在函数外部，你应该有一个 `email` 变量。

```js
assert.isDefined(email);
```

你应该为 `email` 变量赋值一个有效的邮箱地址。

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` 应返回 `"a*******e@example.com"`。

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` 应返回 `"f**********p@example.com"`。

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

你的 `maskEmail` 应产生正确的结果。

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

你应该将调用 `maskEmail` 并传入 `email` 的结果输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```

