---
id: 587d8250367417b2b2512c5f
title: 构建一个栈类
challengeType: 14
dashedName: build-a-stack-class
---

# --description--

栈（Stack）是一种存储有序元素集合的数据结构。它遵循*后进先出*（Last-In-First-Out, LIFO）原则，即最后插入的元素最先被移除。

请完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要定义一个 `Stack` 类，并包含一个初始化为空数组的 `collection` 属性。你将使用这个数组来模拟栈结构。
1. `Stack` 类应有一个 `push` 方法，用于将元素添加到栈顶。
1. `Stack` 类应有一个 `pop` 方法，用于移除并返回栈顶元素。
1. `Stack` 类应有一个 `peek` 方法，用于返回栈顶元素。
1. `Stack` 类应有一个 `isEmpty` 方法，当栈为空时返回 `true`，否则返回 `false`。
1. `Stack` 类应有一个 `clear` 方法，用于清空栈。

# --hints--

你的 `Stack` 类应有一个初始化为空数组的 `collection` 属性。

```js
const test = new Stack();
assert.isArray(test.collection);
assert.lengthOf(test.collection, 0);
```

你的 `Stack` 类应有一个 `push` 方法。

```js
const test = new Stack();
assert.isFunction(test.push);
```

`push` 方法应将元素添加到 `collection` 数组的末尾。

```js
const test = new Stack();
test.push('CS61');
assert.lengthOf(test.collection, 1);
assert.strictEqual(test.collection[0], 'CS61');
test.push('CS50');
assert.lengthOf(test.collection, 2);
assert.strictEqual(test.collection[0], 'CS61');
assert.strictEqual(test.collection[1], 'CS50');
```

你的 `Stack` 类应有一个 `pop` 方法。

```js
const test = new Stack();
assert.isFunction(test.pop);
```

你的 `Stack` 类应有一个 `peek` 方法。

```js
const test = new Stack();
assert.isFunction(test.peek);
```

你的 `Stack` 类应有一个 `isEmpty` 方法。

```js
const test = new Stack();
assert.isFunction(test.isEmpty);
```

你的 `Stack` 类应有一个 `clear` 方法。

```js
const test = new Stack();
assert.isFunction(test.clear);
```

`peek` 方法应返回栈顶元素。

```js
const test = new Stack();
test.push('CS61');
test.push('CS50');
assert.strictEqual(test.peek(), 'CS50');
assert.strictEqual(test.peek(), 'CS50');
```

`pop` 方法应移除并返回栈顶元素。

```js
const test = new Stack();
test.push('CS61');
test.push('CS50');
assert.strictEqual(test.pop(), 'CS50');
assert.lengthOf(test.collection, 1);
assert.strictEqual(test.collection[0], 'CS61');
assert.strictEqual(test.pop(), 'CS61');
assert.lengthOf(test.collection, 0);
```

`isEmpty` 方法在栈为空时应返回 `true`。

```js
const test = new Stack();
assert.isTrue(test.isEmpty());
```

`isEmpty` 方法在栈不为空时应返回 `false`。

```js
const test = new Stack();
test.push('CS61');
assert.isFalse(test.isEmpty());
```

`clear` 方法应移除栈中的所有元素。

```js
const test = new Stack();
test.push('CS61');
test.push('CS50');
test.clear();
assert.isTrue(test.isEmpty());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class Stack {
    constructor() {
        this.collection = [];
    }

    push = function (item) {
        this.collection.push(item);
    }

    pop = function () {
        return this.collection.pop();
    }

    peek = function () {
        return this.collection[this.collection.length - 1];
    }

    isEmpty = function () {
        return this.collection.length === 0;
    }

    clear = function () {
        this.collection.length = 0;
    }
}
```

