---
id: 66ec4c8e9878d8441956516f
title: 构建图书目录表格
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应创建一个 `table` 元素来列出图书信息。
1. 你的表格应包含一个表头（`thead`）元素，并且其中有一行。
1. 表头中的这一行应包含四个表头单元格（`th`），依次为 `Title`、`Author`、`Genre` 和 `Publication Year`。
1. 你的表格应包含一个表体（`tbody`）元素，并且其中至少有五行。
1. 表体中的每一行应包含四个数据单元格（`td`），分别显示图书的标题、作者、类别和出版年份。
1. 你的表格应包含一个表尾（`tfoot`）元素，并且其中有一行。
1. 表尾中的这一行应包含一个跨四列的单元格（`td`），内容为 `Total Books: N`，其中 `N` 应替换为表格中的图书数量。

# --hints--

你应有一个 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

你的 `table` 元素中应有一个 `thead` 元素。

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

你的 `thead` 元素中应有一个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

你的 `thead` 行中应有四个 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

这四个 `th` 元素的文本应依次为 `Title`、`Author`、`Genre` 和 `Publication Year`。

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

你的 `table` 元素中应有一个 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

你的 `tbody` 元素中应至少有五行。

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

`tbody` 中的每一行应正好有四个 `td` 子元素。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

表体中的每个 `td` 元素应包含图书信息文本。

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

你的 `table` 元素中应有一个 `tfoot` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

你的 `tfoot` 元素中应正好有一个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

`tfoot` 行中的 `td` 元素应设置 `colspan` 属性为 `4`。

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

`tfoot` 行中的 `td` 元素文本应为 `Total Books: [number of books in your table]`。

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent, `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>图书目录</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh-CN">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>图书目录</title>
</head>
<body>
    <h1>图书目录</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```

