---
id: 66e028680eca7d21db7e1aee
title: 构建一个正则表达式沙盒
challengeType: 25
dashedName: lab-regex-sandbox
demoType: onClick
---

# --description--

在本实验中，你将获得已经编写好的 CSS 和 HTML。你需要使用 JavaScript，使正则表达式沙盒能够测试正则表达式是否匹配某个字符串，并高亮显示匹配结果。

请完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要获取 `#pattern`、`#test-string`、`#test-btn` 和 `#result` 元素，并分别保存在变量 `regexPattern`、`stringToTest`、`testButton` 和 `testResult` 中。
1. 你需要获取 `#i` 和 `#g` 标志复选框，并分别保存在变量 `caseInsensitiveFlag` 和 `globalFlag` 中。
1. 你需要有一个名为 `getFlags` 的函数，该函数返回一个包含已勾选标志复选框的字符串，如果都未勾选则返回空字符串。
1. 当点击 `testButton` 按钮时，如果正则表达式模式匹配测试字符串，则应高亮显示匹配的文本。要高亮显示匹配内容，请用带有 `highlight` 类的 `span` 元素包裹每个匹配项。相关 CSS 已为你添加好。
1. 匹配成功时，所有匹配的文本应显示在 `#result` 中。如果有多个匹配项，每个匹配项之间用逗号和空格分隔。
1. 如果没有匹配项，则在 `#result` 中显示 `no match`。

# --hints--

你需要获取 `#pattern` 元素，并保存在 `regexPattern` 变量中。

```js
assert.equal(regexPattern, document.querySelector('#pattern'))
```

你需要获取 `#test-string` 元素，并保存在 `stringToTest` 变量中。

```js
assert.equal(stringToTest, document.querySelector('#test-string'))
```

你需要获取 `#test-btn` 元素，并保存在 `testButton` 变量中。

```js
assert.equal(testButton, document.querySelector('#test-btn'))
```

你需要获取 `#result` 元素，并保存在 `testResult` 变量中。

```js
assert.equal(testResult, document.querySelector('#result'))
```

你需要获取 `#i` 标志复选框，并保存在 `caseInsensitiveFlag` 变量中。

```js
assert.equal(caseInsensitiveFlag, document.querySelector('#i'))
```

你需要获取 `#g` 标志复选框，并保存在 `globalFlag` 变量中。

```js
assert.equal(globalFlag, document.querySelector('#g'))
```

你需要有一个 `getFlags` 函数。

```js
assert.isFunction(getFlags)
```

当 `#i` 复选框被勾选时，`getFlags` 函数应返回 `"i"`。

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = false;
assert.equal(getFlags(), "i")
```

当 `#g` 复选框被勾选时，`getFlags` 函数应返回 `"g"`。

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
assert.equal(getFlags(), "g")
```

当两个复选框都被勾选时，`getFlags` 函数应返回 `"ig"` 或 `"gi"`。

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const flags = getFlags();
assert.lengthOf(flags, 2);
assert.include(flags, "i");
assert.include(flags, "g");
```

当两个复选框都未勾选时，`getFlags` 函数应返回空字符串。

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
assert.equal(getFlags(), "")
```

当 `stringToTest` 的 innerHTML 为 `Gu1n34 P1g5`，`regexPattern` 的值为 `\d+`，且未勾选任何标志时，点击 `testButton` 按钮后，`stringToTest.innerHTML` 应变为 `Gu<span class="highlight">1</span>n34 P1g5`。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n34 P1g5';
assert.equal(stringToTest.innerHTML, expected);
```

当 `stringToTest` 的 innerHTML 为 `Gu1n34 P1g5`，`regexPattern` 的值为 `\d+`，且全局标志被勾选时，点击 `testButton` 按钮后，`stringToTest.innerHTML` 应变为 `Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>`。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>';
assert.equal(stringToTest.innerHTML, expected);
```

当 `stringToTest` 的 innerHTML 为 `Gu1n34 P1g5`，`regexPattern` 的值为 `G`，且全局和不区分大小写标志都被勾选时，点击 `testButton` 按钮后，`stringToTest.innerHTML` 应变为 `<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5`。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "G";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5';
assert.equal(stringToTest.innerHTML, expected);
```

点击 `testButton` 按钮时，如果正则表达式模式匹配测试字符串，则匹配的文本应被带有 `highlight` 类的 `span` 元素包裹。

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">Sh</span>e <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = '<span class="highlight">Sh</span>e sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

没有匹配时，测试字符串不应被修改。

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'She sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

当 `stringToTest` 的 innerHTML 为 `Gu1n34 P1g5`，`regexPattern` 的值为 `\d+`，且未勾选任何标志时，点击 `testButton` 按钮后，`#result` 的 innerText 应变为 `1`。

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1"
assert.strictEqual(out.innerText, expected)
```

当 `stringToTest` 的 innerHTML 为 `Gu1n34 P1g5`，`regexPattern` 的值为 `\d+`，且全局标志被勾选时，点击 `testButton` 按钮后，`#result` 的 innerText 应变为 `1, 34, 1, 5`。

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1, 34, 1, 5"
assert.strictEqual(out.innerText, expected)
```

有匹配时，所有匹配的文本应显示在 `#result` 中。

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "Sh, se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = "Sh"
assert.strictEqual(out.innerText, expected)
```

没有匹配时，`#result` 中应显示 `no match`。

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "no match";
assert.strictEqual(out.innerText, expected);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
    <title>正则表达式沙盒</title>

</head>

<body>
    <h1>正则表达式沙盒</h1>
    <main>
        <div id="regex-container">
            <label for="pattern">正则表达式模式:
                <div id="pattern-container">/<input type="text" id="pattern" name="pattern"
                        placeholder="请输入正则表达式">/</div>
            </label>
            <div id="flags-container">
                <p>标志: </p>
                <label for="i">
                    <input type="checkbox" name="flags" id="i"> i
                </label>
                <label for="g">
                    <input type="checkbox" name="flags" id="g"> g
                </label>
            </div>
        </div>
        <div id="test-container">
            <p>测试字符串:</p>
            <div id="test-string" placeholder="请输入测试字符串" contenteditable="true"></div>
        </div>
        <button class="btn" id="test-btn" type="button">测试正则</button>
        <div id="result-container">
            <h2>结果:</h2>
            <p id="result">
            </p>
        </div>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

:root {
    --dark-grey: #1b1b32;
    --light-grey: #f5f6f7;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
    --border: 0.2rem solid darkgrey;
    --padding: 0.3rem;
}

body {
    background-color: var(--dark-grey);
    color: var(--light-grey);
    font-size: 20px;
    font-family: "Lato", Helvetica, Arial, sans-serif;
    padding: 5px;
}

h1 {
    margin: 5rem auto 2rem;
    text-align: center;
}

p {
    padding: var(--padding);
}

#regex-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
    border: var(--border);
}

#regex-container>label {
    padding: var(--padding);
    flex: 1 1 auto;
}

#pattern-container {
    display: inline-block;
    color: var(--dark-grey);
    background-color: var(--light-grey);
    margin: 5px;
    border: var(--border);
}

#pattern {
    margin: 0.2rem;
    border: 0;
    font-size: 1rem;
    width: calc(100% - 1.2rem);
}

#pattern:focus {
    outline: none;
}

#flags-container {
    display: flex;
    align-items: center;
    flex: 1 1 auto;
}

#flags-container>label {
    padding: var(--padding);
    margin-right: 0.3rem;
}

#test-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    flex-direction: column;
    flex: 0 0 auto;
    border: var(--border);
}

#test-string {
    background-color: var(--light-grey);
    min-height: 5rem;
    color: var(--dark-grey);
    border-top: var(--border);
    font-size: 1.2rem;
}

[contenteditable=true]:empty:before {
    content: attr(placeholder);
    pointer-events: none;
    color: var(--dark-grey);
}

::placeholder {
    color: var(--dark-grey);
}


button {
    display: block;
    cursor: pointer;
    width: 8rem;
    margin: 0.2rem auto;
    color: var(--dark-grey);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 0.2rem;
    font-size: 1.1rem;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
}

h2 {
    align-self: flex-start;
    margin: 0.4rem 0.2rem 0.2rem;
    flex: 0 1 auto;
}

#result {
    color: var(--dark-grey);
    background-color: var(--light-grey);
    font-size: 1.5rem;
    flex: 1 1 auto;
    margin: 0.2rem;
    border: var(--border);
    min-height: 3rem;
}

.highlight {
    background-color: lightgreen;
}
```

```js

```

