---
id: 66e45c8140f9fda5c105ae26
title: 构建一个活动宣传单页面
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该在 body 内包含一个 `header` 元素。
1. 你的 `header` 元素中应依次包含一张活动图片和一个 `h1`，`h1` 用于显示活动标题。你可以使用这张图片：`https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. 你应该在 body 内包含一个 `main` 元素。
1. 你的 `main` 元素中应至少包含两个 `section` 元素，用于展示活动特色。
1. 每个 `section` 元素中都应包含一个 `h2`。
1. 你的 body 顶部和底部应有 `50px` 的内边距（padding）。
1. 你的 body 顶部和底部外边距（margin）应为 `0`，左右外边距应能使其居中。
1. 你的 body 宽度应相对于视口宽度设置。
1. 你的 body 应使用 `calc` 函数将其 `min-height` 属性设置为视口高度的 100% 减去 body 顶部和底部所有内边距。
1. 你的宣传单中应至少有一个 `hr` 元素。
1. 你应将所有 `hr` 和 `section` 元素的宽度设置为相对于其父元素的百分比值。

**注意：** 请确保在 HTML 中链接你的样式表并应用 CSS。

# --hints--

你应该在 `body` 元素内有一个 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你应该在 `header` 的第一个子元素为 `img`。

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

你应该在 `header` 的第二个子元素为 `h1`。

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

你只能有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 应包含活动标题。

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

你应该在 `body` 元素内有一个 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

你的 `main` 元素内应至少有两个 `section` 元素。

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

每个 `section` 元素中都应包含一个 `h2`。

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

你的 `h2` 元素不能为空。

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

你的 `body` 元素顶部和底部应有 `50px` 的内边距。

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

你的 `body` 元素顶部和底部外边距应为 `0`。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

你的 `body` 元素左右外边距应为 `auto`。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

你的 `body` 元素宽度应使用 `vw` 设置。

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

你的 `body` 应使用 `calc` 将 `min-height` 设置为 `100vh - 100px`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

你应至少有一个 `hr` 元素。

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

你的 `hr` 元素宽度应使用百分比值设置。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

你的 `section` 元素宽度应使用百分比值设置。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>构建一个活动宣传单页面</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>构建一个活动宣传单页面</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>宠物石头节</h1>
    <p>欢迎参加首届宠物石头节！</p>
    <p><b>时间：</b> <time datetime="2025-10-03">10月3日</time> <b>地点：</b> Boulder 公园</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>活动亮点</h2>
      <ul>
        <li>宠物石头服从表演</li>
        <li>最佳着装宠物石头比赛</li>
        <li>宠物石头爬坡赛</li>
      </ul>
    </section>

    <section>
      <h2>特邀嘉宾</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>美食车</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 宠物石头节</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```

