---
id: 67051431a73f1ca25d9a6f25
title: 构建项目创意板
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该定义一个名为 `projectStatus` 的对象常量，包含三个键：`PENDING`、`SUCCESS` 和 `FAILURE`。每个状态都应分配一个包含 `description` 键的对象，值分别为 `待执行`、`执行成功` 和 `执行失败`。
2. 你应该定义一个名为 `ProjectIdea` 的类，其 `constructor` 接受 `title` 和 `description` 作为参数。用传入的参数初始化 `title` 和 `description` 属性。该类还应有一个名为 `status` 的属性，默认值为 `projectStatus.PENDING`。
3. 你应该在 `ProjectIdea` 类中定义一个名为 `updateProjectStatus` 的方法。该方法接受一个 `newStatus` 参数，并将 `status` 属性更新为该值。
4. 你应该定义一个名为 `ProjectIdeaBoard` 的类，其 `constructor` 接受一个 `title`，并初始化一个空数组 `ideas` 用于存放 `ProjectIdea` 类的实例。
5. 你应该在 `ProjectIdeaBoard` 类中定义一个名为 `pin` 的方法，接受一个 `ProjectIdea` 实例并将其添加到 `ideas` 数组中。
6. 你应该在 `ProjectIdeaBoard` 类中定义一个名为 `unpin` 的方法。该方法接受一个 `ProjectIdea` 实例并将其从 `ideas` 数组中移除。
7. 你应该定义一个名为 `count` 的方法，返回当前 `ProjectIdeaBoard` 中项目创意的数量。
8. 你应该定义一个名为 `formatToString` 的方法，返回项目创意的名称、描述和状态，格式如下：

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

你应该定义一个名为 `projectStatus` 的对象常量，包含正确的项目状态和描述。

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "待执行" });
assert.deepEqual(projectStatus.SUCCESS, { description: "执行成功" });
assert.deepEqual(projectStatus.FAILURE, { description: "执行失败" });
```

你应该有一个 `ProjectIdea` 类。

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

你的 `ProjectIdea` 类应根据传入参数初始化 `title` 和 `description` 属性，并将 `status` 设置为 `projectStatus.PENDING`。

```js
const test = new ProjectIdea("示例标题", "示例描述");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "待执行");
```

你的 `ProjectIdea` 类应有一个 `updateProjectStatus` 方法。

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

你应该有一个 `ProjectIdeaBoard` 类。

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

你的 `ProjectIdeaBoard` 应根据传入参数初始化 `title` 属性，并初始化一个空数组 `ideas` 用于存放 `ProjectIdea` 实例。

```js
const test = new ProjectIdeaBoard("示例标题");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

你的 `ProjectIdeaBoard` 类应有一个 `pin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

你的 `ProjectIdeaBoard` 类应有一个 `unpin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

你的 `ProjectIdeaBoard` 类应有一个 `count` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

你的 `ProjectIdeaBoard` 类应有一个 `formatToString` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("智能窗锁", "一个自动化项目，允许用户根据天气状况自动锁定和解锁窗户。")` 应返回 `{ title: '智能窗锁', description: '一个自动化项目，允许用户根据天气状况自动锁定和解锁窗户。', status: { description: '待执行' } }`。

```js
const windowLocks = new ProjectIdea(
  "智能窗锁",
  "一个自动化项目，允许用户根据天气状况自动锁定和解锁窗户。"
);
assert.deepEqual(windowLocks, {
  title: "智能窗锁",
  description: "一个自动化项目，允许用户根据天气状况自动锁定和解锁窗户。",
  status: { description: "待执行" },
});
```

在 `new ProjectIdea("健身追踪应用", "一个可以追踪用户锻炼、饮食和睡眠模式的应用。")` 上调用 `updateProjectStatus(projectStatus.SUCCESS)` 应将状态更新为 `{ description: '执行成功' }`。

```js
let mobileApp = new ProjectIdea(
  "健身追踪应用",
  "一个可以追踪用户锻炼、饮食和睡眠模式的应用。"
);

// 更新项目状态为 SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "健身追踪应用",
  description: "一个可以追踪用户锻炼、饮食和睡眠模式的应用。",
  status: { description: "执行成功" },
});
```

在 `new ProjectIdea("早餐厨师机器人", "一个可以根据给定指令列表为用户准备早餐并通过手机通知用户的机器人。")` 上调用 `updateProjectStatus(projectStatus.FAILURE)` 应将状态更新为 `{ description: '执行失败' }`。

```js
let breakfastChef = new ProjectIdea(
  "早餐厨师机器人",
  "一个可以根据给定指令列表为用户准备早餐并通过手机通知用户的机器人。"
);

// 更新项目状态为 FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "早餐厨师机器人",
  description: "一个可以根据给定指令列表为用户准备早餐并通过手机通知用户的机器人。",
  status: { description: "执行失败" },
});
```

在 `new ProjectIdea("二手游戏商店", "一个在线平台，用户可以从其他用户那里购买二手实体游戏。")` 上调用 `updateProjectStatus(projectStatus.SUCCESS)` 应将状态更新为 `{ description: '执行成功' }`。

```js
let usedVideoGamesStore = new ProjectIdea(
  "二手游戏商店",
  "一个在线平台，用户可以从其他用户那里购买二手实体游戏。"
);

// 更新项目状态为 SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "二手游戏商店",
  description: "一个在线平台，用户可以从其他用户那里购买二手实体游戏。",
  status: { description: "执行成功" },
});
```

你应该可以使用 `pin` 方法将 `ProjectIdea` 对象固定到你的 `ProjectIdeaBoard` 上。

```js
let smartHome = new ProjectIdea(
  "智能家居系统",
  "一个集成系统，可远程控制照明、温度和安全设备。"
);

// 创建项目板并固定创意
let techProjects = new ProjectIdeaBoard("科技项目板");
techProjects.pin(smartHome);

// 正确断言以比较数组中的第一个创意
assert.deepEqual(techProjects.ideas[0], {
    title: '智能家居系统',
    description: '一个集成系统，可远程控制照明、温度和安全设备。',
    status: { description: '待执行' }
});
```

你应该可以使用 `unpin` 方法将 `ProjectIdea` 对象从你的 `ProjectIdeaBoard` 上移除。

```js
let smartHome = new ProjectIdea(
  "智能家居系统",
  "一个集成系统，可远程控制照明、温度和安全设备。"
);

// 创建项目板并固定创意
let techProjects = new ProjectIdeaBoard("科技项目板");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

当 `new ProjectIdeaBoard("空板")` 为空时，`emptyBoard.formatToString()` 应返回 `空板 has 0 idea(s)\n`。

```js

let smartHome = new ProjectIdea(
  "智能家居系统",
  "一个集成系统，可远程控制照明、温度和安全设备。"
);

// 创建项目板并固定创意
let techProjects = new ProjectIdeaBoard("科技项目板");

// 创建一个空板
let emptyBoard = new ProjectIdeaBoard("空板");

// 尝试从空板移除项目
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "空板 has 0 idea(s)\n");
```

当你将 `new ProjectIdea("智能家居系统", "一个集成系统，可远程控制照明、温度和安全设备。")` 固定到 `new ProjectIdeaBoard("科技项目板")` 时，`techProjects.formatToString()` 应返回 `科技项目板 has 1 idea(s)\n智能家居系统 (待执行) - 一个集成系统，可远程控制照明、温度和安全设备。\n`。

```js
let techProjectsBoard = new ProjectIdeaBoard("科技项目板");
let smartHome = new ProjectIdea("智能家居系统", "一个集成系统，可远程控制照明、温度和安全设备。");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "科技项目板 has 1 idea(s)\n智能家居系统 (待执行) - 一个集成系统，可远程控制照明、温度和安全设备。\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "待执行" },
  SUCCESS: { description: "执行成功" },
  FAILURE: { description: "执行失败" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "智能窗锁",
  "一个自动化项目，允许用户根据天气状况自动锁定和解锁窗户。"
);

let breakfastChef = new ProjectIdea(
  "早餐厨师机器人",
  "一个可以根据给定指令列表为用户准备早餐并通过手机通知用户的机器人。"
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "二手游戏商店",
  "一个在线平台，用户可以从其他用户那里购买二手实体游戏。"
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("硬件项目板");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("软件项目板");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```

