---
id: 66a207974c806a19d6607073
title: 构建一本图书库存应用
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个文本为 `Book Inventory` 的 `h1` 元素。
1. 你应该有一个包含 `Title`、`Author`、`Category`、`Status` 和 `Rate` 列的 `table` 元素。
1. 你的表格应至少有四行，第一行为列标题，其余为图书信息。
1. 表格主体中的每一行都应有 `read`、`to-read` 或 `in-progress` 其中之一的类名。
1. `Status` 列的 `td` 元素应包含一个 `class` 为 `status` 的 `span` 元素，内容为 `Read`、`To Read` 或 `In Progress`，具体取决于该行的类名。
1. `Rate` 列的 `td` 元素应包含一个 `class` 为 `rate` 的 `span` 元素，内部包裹三个空的 `span` 元素。
1. 位于 `.read` 行内的 `.rate` 元素应在 `rate` 类后额外包含 `one`、`two` 或 `three` 之一的类名，具体取决于个人评分。
1. 你应创建三个属性选择器，分别针对 `read`、`to-read` 和 `in-progress` 类的元素，并为它们设置你选择的 `linear-gradient` 线性渐变背景图像。
1. 你应将每个 `span` 元素的 `display` 属性设置为 `inline-block`。
1. 你应使用属性选择器，针对属于 `to-read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素，设置其 `border` 和 `background-image` 属性。
1. 你应使用属性选择器，针对属于 `read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素，设置其 `border` 和 `background-image` 属性。
1. 你应使用属性选择器，针对属于 `in-progress` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素，设置其 `border` 和 `background-image` 属性。
1. 你应使用属性选择器，针对 `class` 为 `status` 的 `span` 元素和 `class` 以 `rate` 开头的 `span` 元素，设置其 `height`、`width` 和 `padding` 属性。
1. 你应使用属性选择器，针对 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span`，设置其 `border`、`border-radius`、`margin`、`height`、`width` 和 `background-color` 属性。
1. 你应使用属性选择器，针对 `class` 包含 `one` 的 `span` 元素的第一个后代，设置其 `background-image` 属性为线性渐变。
1. 你应使用属性选择器，针对 `class` 包含 `two` 的 `span` 元素的前两个后代，设置其 `background-image` 属性为线性渐变。
1. 你应使用属性选择器，针对 `class` 包含 `three` 的 `span` 元素的三个后代 `span`，设置其 `background-image` 属性为线性渐变。

# --hints--

你应该有一个文本为 `Book Inventory` 的 `h1` 元素。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Book Inventory');
```

你应该只有一个 `h1` 元素。

```js
assert.equal(document.querySelectorAll('h1')?.length, 1);
```

你的 `h1` 元素后面应该有一个 `table` 元素。

```js
assert.equal(document.querySelector('table')?.previousElementSibling?.tagName, 'H1')
```

你的 `table` 元素应有五列。

```js
assert.equal(document.querySelectorAll('th')?.length, 5);
```

你的第一列标题应为 `Title`。

```js
assert.equal(document.querySelectorAll('th')[0]?.innerText, 'Title');
```

你的第二列标题应为 `Author`。

```js
assert.equal(document.querySelectorAll('th')[1]?.innerText, 'Author');
```

你的第三列标题应为 `Category`。

```js
assert.equal(document.querySelectorAll('th')[2]?.innerText, 'Category');
```

你的第四列标题应为 `Status`。

```js
assert.equal(document.querySelectorAll('th')[3]?.innerText, 'Status');
```

你的第五列标题应为 `Rate`。

```js
assert.equal(document.querySelectorAll('th')[4]?.innerText, 'Rate');
```

你的表格应至少有四行。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

除标题行外，每一行都应有 `read`、`to-read` 或 `in-progress` 其中之一的类名。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`Status` 列的 `td` 元素应包含一个 `span` 元素。

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isAbove(statusColumnData.length, 0);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

每个 `Status` 列的 `span` 元素应有 `status` 类。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isAbove(statusSpans.length, 0);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

每个 `.status` 元素的文本应为 `Read`、`To Read` 或 `In Progress`，具体取决于其所在行的类名。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isAbove(statusSpans.length, 0);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`Rate` 列的 `td` 元素应包含一个 `span` 元素。

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isAbove(rateColumnData.length, 0);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

每个 `Rate` 列的 `span` 元素应以 `rate` 为首个类。

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

每个 `.rate` 元素应包含三个空的 `span` 元素。

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

位于 `.read` 行内的 `.rate` 元素应在 `rate` 类后额外包含 `one`、`two` 或 `three` 之一的类名。

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isAbove(readBooksRates.length, 0);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

你应有属性选择器针对 `read` 类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

你应使用属性选择器为 `read` 类的行设置线性渐变背景图像。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

你应有属性选择器针对 `to-read` 类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

你应使用属性选择器为 `to-read` 类的行设置线性渐变背景图像。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

你应有属性选择器针对 `in-progress` 类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

你应使用属性选择器为 `in-progress` 类的行设置线性渐变背景图像。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

你应将每个 `span` 元素的 `display` 属性设置为 `inline-block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

你应有属性选择器针对 `to-read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

你应使用属性选择器为 `to-read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

你应使用属性选择器为 `to-read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

你应有属性选择器针对 `read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

你应使用属性选择器为 `read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

你应使用属性选择器为 `read` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

你应有属性选择器针对 `in-progress` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

你应使用属性选择器为 `in-progress` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

你应使用属性选择器为 `in-progress` 类的 `tr` 元素后代中 `class` 为 `status` 的 `span` 元素设置 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

你应有属性选择器针对 `class` 为 `status` 的 `span` 元素和 `class` 以 `rate` 开头的 `span` 元素。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

你应使用属性选择器为 `class` 为 `status` 的 `span` 元素和 `class` 以 `rate` 开头的 `span` 元素设置 `height` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

你应使用属性选择器为 `class` 为 `status` 的 `span` 元素和 `class` 以 `rate` 开头的 `span` 元素设置 `width` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

你应使用属性选择器为 `class` 为 `status` 的 `span` 元素和 `class` 以 `rate` 开头的 `span` 元素设置 `padding` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

你应有属性选择器针对 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `border-radius` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `margin` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

你应使用属性选择器为 `class` 以 `rate` 开头的 `span` 元素的直接子元素 `span` 设置 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

你应有属性选择器针对 `class` 包含 `one` 的 `span` 元素的第一个后代。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应使用属性选择器为 `class` 包含 `one` 的 `span` 元素的第一个后代设置线性渐变背景图像。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

你应有属性选择器针对 `class` 包含 `two` 的 `span` 元素的前两个后代。

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应使用属性选择器为 `class` 包含 `two` 的 `span` 元素的前两个后代设置线性渐变背景图像。

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

你应有属性选择器针对 `class` 包含 `three` 的 `span` 元素的所有后代 `span`。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应使用属性选择器为 `class` 包含 `three` 的 `span` 元素的所有后代 `span` 设置线性渐变背景图像。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <table>
      <thead>
        <tr>
          <th>Title</th>
          <th>Author</th>
          <th>Category</th>
          <th>Status</th>
          <th>Rate</th>
        </tr>
      </thead>
      <tbody>
        <tr class="read">
          <td>The Three Musketeers</td>
          <td>Alexandre Dumas</td>
          <td>Historical Novel</td>
          <td><span class="status">Read</span></td>
          <td>
            <span class="rate three">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
        <tr class="to-read">
          <td>The Plague</td>
          <td>Albert Camus</td>
          <td>Philosophical Novel</td>
          <td><span class="status">To Read</span></td>
          <td>
            <span class="rate">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
        <tr class="read">
          <td>The Metamorphosis</td>
          <td>Franz Kafka</td>
          <td>Novella</td>
          <td><span class="status">Read</span></td>
          <td>
            <span class="rate one">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
        <tr class="read">
          <td>Dead Souls</td>
          <td>Nikolai Gogol</td>
          <td>Picaresque</td>
          <td><span class="status">Read</span></td>
          <td>
            <span class="rate two">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
        <tr class="in-progress">
          <td>Lord of the Flies</td>
          <td>William Golding</td>
          <td>Allegorical Novel</td>
          <td><span class="status">In Progress</span></td>
          <td>
            <span class="rate">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
        <tr class="read">
          <td>Do Androids Dream of Electric Sheep?</td>
          <td>Philip K. Dick</td>
          <td>Science Fiction</td>
          <td><span class="status">Read</span></td>
          <td>
            <span class="rate two">
              <span></span><span></span><span></span>
            </span>
          </td>
        </tr>
      </tbody>
    </table>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```

