---
id: 66c057041df6394ca796bf33
title: 构建一个句子生成器
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

在本实验中，你将使用句子模板创建两个不同的故事。你需要使用变量来存储故事的不同部分，然后将故事输出到控制台。

请完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你应该使用 `let` 声明以下变量：

    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

2. 你应该为上述变量赋予你选择的字符串值。

3. 你应该声明一个 `firstStory` 变量。

4. 你应该使用以下故事模板创建第一个故事，并将其赋值给 `firstStory` 变量：  
   `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`

5. 你应该使用消息 `"First story: [firstStory]"` 将第一个故事输出到控制台。

6. 你应该为 `adjective`、`noun`、`verb`、`place`、`adjective2` 和 `noun2` 变量赋予新的值。

7. 你应该声明一个 `secondStory` 变量。

8. 使用相同的模板创建另一个故事，并将其赋值给 `secondStory` 变量。

9. 你应该使用消息 `"Second story: [secondStory]"` 将第二个故事输出到控制台。


# --hints--

你应该声明一个 `adjective` 变量。

```js
assert.isNotNull(adjective);
```

你应该声明一个 `noun` 变量。

```js
assert.isNotNull(noun);
```

你应该声明一个 `verb` 变量。

```js
assert.isNotNull(verb);
```

你应该声明一个 `place` 变量。

```js
assert.isNotNull(place);
```

你应该声明一个 `adjective2` 变量。

```js
assert.isNotNull(adjective2);
```

你应该声明一个 `noun2` 变量。

```js
assert.isNotNull(noun2);
```

你应该为 `adjective` 变量赋值为字符串。

```js
assert.isString(adjective);
```

你应该为 `noun` 变量赋值为字符串。

```js
assert.isString(noun);
```

你应该为 `verb` 变量赋值为字符串。

```js
assert.isString(verb);
```

你应该为 `place` 变量赋值为字符串。

```js
assert.isString(place);
```

你应该为 `adjective2` 变量赋值为字符串。

```js
assert.isString(adjective2);
```

你应该为 `noun2` 变量赋值为字符串。

```js
assert.isString(noun2);
```

你应该声明一个 `firstStory` 变量。

```js
assert.isNotNull(firstStory);
```

你应该为第一个故事使用正确的模板格式：  
`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`  
请注意空格。

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

你应该使用消息 `"First story: [firstStory]"` 输出第一个故事。

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

你应该声明一个 `secondStory` 变量。

```js
assert.isNotNull(secondStory);
```

你应该为第二个故事重新赋值 `adjective` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

你应该为第二个故事重新赋值 `noun` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

你应该为第二个故事重新赋值 `verb` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

你应该为第二个故事重新赋值 `place` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

你应该为第二个故事重新赋值 `adjective2` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

你应该为第二个故事重新赋值 `noun2` 变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

你应该为第二个故事使用正确的模板格式：  
`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`  
请注意空格。

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

你应该使用消息 `"Second story: [secondStory]"` 输出第二个故事。

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` 不应与 `secondStory` 相同。

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
