---
id: 66dadcf18df3a76104054d95
title: 构建一个链表类
challengeType: 14
dashedName: build-a-linked-list-class
---

# --description--

在本实验中，你将构建一个链表。链表是一种线性数据结构，由称为<dfn>节点</dfn>的数据元素组成，每个节点都指向链表中的下一个节点。链表中的每个节点包含两个关键信息：元素本身，以及对下一个节点的引用。

请完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要创建一个名为 `LinkedList` 的类，用于表示链表数据结构。

1. 每当实例化一个新的 `LinkedList` 对象时，其 `length` 和 `head` 属性应分别设置为 `0` 和 `null`。

1. `LinkedList` 类应有一个属性 `Node`，它是一个类，实例化时 `element` 属性被设置为传入的值，`next` 属性被设置为 `null`。

1. `LinkedList` 类应有一个方法 `isEmpty`，当链表为空时返回 `true`，否则返回 `false`。

1. `LinkedList` 类应有一个方法 `add`，该方法从其参数创建一个 `Node` 对象并将其添加到链表中。

1. 当添加的节点是第一个节点时，`add` 方法应将 `head` 设置为新节点，否则应将链表中最后一个节点的 `next` 属性设置为要添加的节点。

1. 每次向链表添加新节点时，链表的长度应加一。

1. `LinkedList` 类应有一个方法 `remove`，该方法接受一个参数并从链表中移除对应的节点。

1. 每当从链表中移除一个节点时，前一个节点的 `next` 属性应指向被移除节点的下一个节点，并且链表长度应减一。

# --hints--

你的 `LinkedList` 类实例化对象时，`length` 属性应为 `0`。

```js
const test = new LinkedList();
assert.strictEqual(test.length, 0);
```

你的 `LinkedList` 类实例化对象时，`head` 属性应为 `null`。

```js
const test = new LinkedList();
assert.strictEqual(test.head, null);
```

你的 `LinkedList` 类应有一个属性 `Node`。

```js
const test = new LinkedList();
assert.property(test, "Node");
```

你的 `Node` 属性应为一个类，实例化时 `element` 属性被设置为传入的值。

```js
const test = new LinkedList();
const node = new test.Node("pig");
assert.strictEqual(node.element, "pig");
```

你的 `Node` 属性应为一个类，实例化时 `next` 属性被设置为 `null`。

```js
const test = new LinkedList();
const node = new test.Node("pig");
assert.strictEqual(node.next, null);
```

你的 `LinkedList` 类应有一个 `add` 方法。

```js
const test = new LinkedList();
assert.typeOf(test.add, 'function');
```

你的 `add` 方法应将 `head` 分配为添加到链表中的第一个节点对象。

```js
const test = new LinkedList();
test.add("pig");
const expected = { element: "pig", next: null }
assert.deepEqual(test.head, expected);
```

你的 `add` 方法应将链表中最后一个节点的 `next` 属性设置为要添加的节点。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog")
const expected = { element: "pig", next: {element: "dog", next: null} }
assert.deepEqual(test.head, expected);
```

每次向链表添加新节点时，链表的 `length` 应加一。

```js
const test = new LinkedList();
test.add("pig");
assert.strictEqual(test.length, 1);
test.add("dog");
assert.strictEqual(test.length, 2);
test.add("sheep");
assert.strictEqual(test.length, 3);
```

你的 `LinkedList` 类应有一个 `isEmpty` 方法。

```js
const test = new LinkedList();
assert.typeOf(test.isEmpty, 'function');
```

当链表中至少有一个元素时，`isEmpty` 方法应返回 `false`。

```js
const test = new LinkedList();
test.add('pig');
test.add('dog');
test.add('kitten');
assert.isFalse(test.isEmpty());
```

当链表中没有元素时，`isEmpty` 方法应返回 `true`。

```js
const test = new LinkedList();
assert.isTrue(test.isEmpty());
```

你的 `LinkedList` 类中，前一个节点应引用新创建的节点。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog")
test.add("sheep");
assert(test.head.next.element === 'dog' && test.head.next.next.element === "sheep");
```

你的 `LinkedList` 类应有一个 `remove` 方法。

```js
const test = new LinkedList();
assert.typeOf(test.remove, 'function');
```

当移除第一个节点时，`remove` 方法应将 `head` 重新分配为第二个节点。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.remove("pig");
const expected = { element: "dog", next: {element: "sheep", next: null} }
assert.deepEqual(test.head, expected);
```

每移除一个节点，链表的 `length` 应减一。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.remove("pig");
assert.strictEqual(test.length, 2);
test.remove("dog");
assert.strictEqual(test.length, 1);
test.remove("sheep");
assert.strictEqual(test.length, 0);
```

你的 `remove` 方法应将前一个节点的引用重新指向被移除节点的下一个节点。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.add("cat");
test.remove("sheep");
const expected = { element: "pig", next: {element: "dog", next: {element: "cat", next: null}} }
assert.deepEqual(test.head, expected);
```

如果链表中不存在该元素，`remove` 方法不应改变链表。

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.add("cat");
test.remove("elephant");
const expected = { element: "pig", next: {element: "dog", next: {element: "sheep", next: {element: "cat", next: null}}} }
assert.deepEqual(test.head, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class LinkedList {
    constructor() {
        this.length = 0;
        this.head = null;
    }

    Node = class {
        constructor(element) {
            this.element = element;
            this.next = null;
        }
    }

    isEmpty() {
        return this.length === 0;
    }

    add(element) {
        const node = new this.Node(element)
        if (this.isEmpty()) {
            this.head = node;
        } else {
            let currentNode = this.head;
            while (currentNode.next !== null) {
                currentNode = currentNode.next;
            }
            currentNode.next = node;
        }
        this.length++;
    }

    remove(element) {
        if (this.isEmpty()) {
            return
        }
        let previousNode;
        let currentNode = this.head;
        while (currentNode.next !== null && currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }
        if (currentNode.next === null && currentNode.element !== element) {
            return;
        } else if (previousNode) {
            previousNode.next = currentNode.next
        } else {
            this.head = currentNode.next;
        }
        this.length--
    }
}
```

