---
id: 66f26c32ec6f90df01a44f60
title: 设计一个联系表单
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你需要有一个 `class` 为 `form-container` 的 `div` 元素来包裹你的表单。

2. 在 `.form-container` 元素内需要有一个 `form` 元素。

3. 在 `form` 元素内，你需要包含以下元素和输入字段：

  - 一个带有文本的 `h2` 元素。
  - 一个用于姓名的文本输入框，`type` 设为 `text`，并包含 `id`、`name`、`required` 属性。
  - 一个用于邮箱的输入框，`type` 设为 `email`，并包含 `id`、`name`、`required` 属性。
  - 一个用于留言的 `textarea`，包含 `id`、`name` 和 `required` 属性。
  - 一个文本为 `"Submit"` 且 `type` 属性为 `submit` 的 `button` 元素。

4. 每个 `input` 和 `textarea` 元素都应有对应的 `label`，其 `for` 属性值与对应元素的 `id` 相同。

5. `.form-container` 元素应设置背景色，并有 `border-radius`、`padding` 和 `width` 的值。

6. `label` 元素应设置外边距和字体颜色。

7. `input` 和 `textarea` 元素应设置：

  - `width`。
  - `padding`。
  - 下外边距（margin-bottom）。

8. `button` 元素应设置：

  - 背景色
  - 非默认的字体大小

9. 按钮应有悬停（hover）时改变背景色的效果。

**注意：** 请确保在 HTML 中正确引入样式表并应用 CSS。

# --hints--

你需要有一个 `class` 为 `form-container` 的 `div` 元素。

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

你需要在 `.form-container` 元素内有一个 `form` 元素。

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

在 `form` 内应有一个带文本的 `h2` 元素。

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

在 `form` 内应有一个 `type="text"` 的文本输入框。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

你的 `type=text` 输入框应有 `id`、`name` 和 `required` 属性。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form` 内应有一个 `type="email"` 的邮箱输入框。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

你的 `type=email` 输入框应有 `id`、`name` 和 `required` 属性。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form` 内应有一个带有 `id`、`name` 和 `required` 属性的 `textarea`。

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form` 内应有一个 `type="submit"` 的 `button` 元素。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

你的按钮文本应为 `"Submit"`。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

你应为 `#name` 元素设置一个 `label`，其 `for` 属性与对应输入框的 `id` 相同。

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
   if (label.getAttribute('for') === nameInput) {
      found = true;
      break;
   }
}

assert.isTrue(found);
```

你应为 `#email` 元素设置一个 `label`，其 `for` 属性与对应输入框的 `id` 相同。

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
   if (label.getAttribute('for') === emailInput) {
      found = true;
      break;
   }
}

assert.isTrue(found);
```

你应为 `#message` 元素设置一个 `label`，其 `for` 属性与对应 `textarea` 的 `id` 相同。

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

所有 `label` 不应为空。

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
   const labelText = label.textContent.trim();
   assert.isNotEmpty(labelText);
}

```

`.form-container` 元素应设置背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

`.form-container` 元素应设置 `border-radius`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

`.form-container` 元素应设置 `padding`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

`.form-container` 元素应设置 `width`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

`label` 元素应设置外边距和字体颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

`input` 元素应设置宽度、内边距和下外边距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
   result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

`textarea` 元素应设置宽度、内边距和下外边距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
   result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

`button` 元素应设置背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

`button` 元素应设置字体大小。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

`button` 元素应有悬停时改变背景色的效果。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
   <meta charset="utf-8">
   <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
   <meta charset="UTF-8">
   <meta name="viewport" content="width=device-width, initial-scale=1.0">
   <title>Styled Form Project</title>
   <link rel="stylesheet" href="styles.css">
</head>
<body>
   <div class="form-container">
      <form>
        <h2>Contact Us</h2>
        <label for="name">Name:</label>
        <input type="text" id="name" name="name" required>

        <label for="email">Email:</label>
        <input type="email" id="email" name="email" required>

        <label for="message">Message:</label>
        <textarea id="message" name="message" required></textarea>

        <button type="submit">Submit</button>
      </form>
   </div>
</body>
</html>
```

```css
body {
   display: flex;
   justify-content: center;
   align-items: center;
   height: 100vh;
   margin: 0;
   background: linear-gradient(45deg, #83a4d4, #b6fbff);
   font-family: Arial, sans-serif;
}

.form-container {
   background-color: white;
   border-radius: 10px;
   padding: 20px;
   width: 300px;
   text-align: center;
}

form h2 {
   margin-bottom: 20px;
   color: #333;
}

label {
   display: block;
   margin: 10px 0 5px;
   color: #555;
}

input{
   width: 100%;
   padding: 10px;
   margin-bottom: 10px;
   border: 1px solid #ddd;
   border-radius: 5px;
   box-sizing: border-box;
   resize: vertical;
}

textarea {
   width: 100%;
   padding: 10px;
   margin-bottom: 10px;
   border: 1px solid #ddd;
   border-radius: 5px;
   box-sizing: border-box;
   resize: vertical;
}

button {
   background-color: #4CAF50;
   color: white;
   padding: 10px 15px;
   border: none;
   border-radius: 5px;
   cursor: pointer;
   font-size: 1em;
}

button:hover {
   background-color: #45a049;
}
```

