---
id: 613297a923965e0703b64796
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

你可能已经熟悉 `meta` 元素了；它用于指定关于页面的信息，例如标题、描述、关键字和作者。

为你的页面提供一个带有适当 `charset` 值的 `meta` 元素。

`charset` 属性指定页面的字符编码，如今，`UTF-8` 是大多数浏览器唯一支持的编码。

# --hints--

你应该在 `head` 元素内创建一个 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你应该给 `meta` 标签一个 `UTF-8` 的 `charset`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

