---
id: 6133acc353338c0bba9cb553
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

最后，在 `head` 中，`title` 元素对于屏幕阅读器理解页面内容非常有用。此外，它也是 _SEO_ 的重要组成部分。

请为你的页面添加一个简洁且具有描述性的 `title`。

# --hints--

你应该在 `head` 中添加一个 `title` 元素。

```js
assert.exists(document.querySelector('head > title'));
```

你的 `title` 不应超过 60 个字符。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

尝试让你的 `title` 更具描述性。_提示：至少 15 个字符_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍测验练习项目" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

