---
id: 613e275749ebd008e74bb62e
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

目前，`img` 元素采用了默认尺寸，显示过大。CSS 提供了一个 `max` 函数，可以返回一组逗号分隔值中的最大值。例如：

```css
img {
  width: max(250px, 25vw);
}
```

在上面的例子中，如果视口宽度小于 1000 像素，图片宽度为 250px；如果视口宽度大于 1000 像素，图片宽度为 25vw（即视口宽度的 25%）。

请使用图片的 `id` 作为选择器，将其 `width` 设置为 `10rem` 和 `18vw` 中的最大值。

# --hints--

你应该使用 `#logo` 选择器来选中 `img` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

你应该将 `img` 的 `width` 设置为 `max(10rem, 18vw)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍测验练习项目" />
    <title>无障碍测验</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS 测验</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```

