---
id: 613e2546d0594208229ada50
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `header` 中，通过嵌套一个 `img`、`h1` 和 `nav` 元素，为页面提供上下文信息。

`img` 应该指向 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`，`id` 设置为 `logo`，`alt` 文本为 `freeCodeCamp`。

`h1` 应包含文本 `HTML/CSS Quiz`。

# --hints--

你应该在 `header` 元素中添加一个 `img` 元素。

```js
assert.exists(document.querySelector('header > img'));
```

你应该在 `header` 元素中添加一个 `h1` 元素。

```js
assert.exists(document.querySelector('header > h1'));
```

你应该在 `header` 元素中添加一个 `nav` 元素。

```js
assert.exists(document.querySelector('header > nav'));
```

你应该按顺序放置 `img`、`h1` 和 `nav` 元素。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

你应该为 `img` 元素设置 `src` 属性为 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

你应该为 `img` 元素设置 `id` 属性为 `logo`。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

你应该为 `img` 元素设置 `alt` 属性为 `freeCodeCamp`。

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

你应该为 `h1` 元素设置文本 `HTML/CSS Quiz`。

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

