---
id: 6141fabd6f75610664e908fd
title: 第14步
challengeType: 0
dashedName: step-14
---

# --description--

由于这是一个测验，你需要为用户提交答案创建一个表单。你可以使用 `section` 元素在表单内语义化地分隔内容。

在 `main` 元素内，创建一个包含三个嵌套 `section` 元素的表单。然后，让该表单提交到 `https://freecodecamp.org/practice-project/accessibility-quiz`，并使用正确的方法。

# --hints--

你应该在 `main` 元素内嵌套一个 `form` 元素。

```js
assert.exists(document.querySelector('main > form'));
```

你应该在 `form` 元素内嵌套三个 `section` 元素。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

你应该为 `form` 元素添加一个 `action` 属性。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

你应该将 `action` 属性的值设置为 `https://freecodecamp.org/practice-project/accessibility-quiz`。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

你应该为 `form` 元素添加一个 `method` 属性。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

你应该将 `form` 元素的 `method` 属性值设置为 `post`。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍测验练习项目" />
    <title>无障碍测验</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS 测验</h1>
      <nav>
        <ul>
          <li><a>信息</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
