---
id: 614206033d366c090ca7dd42
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

字体在页面的可访问性中扮演着重要角色。有些字体比其他字体更易于阅读，在低分辨率屏幕上尤其如此。

将 `h1` 和 `h2` 元素的字体都更改为 `Verdana`，并使用 sans-serif 系列中的另一种网页安全字体作为备选。

另外，为 `h2` 元素添加 `4px solid #dfdfe2` 的 `border-bottom`，以使各个区域更加分明。

# --hints--

你应该使用多元素选择器来选中 `h1` 和 `h2` 元素。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
assert.exists(gs('h1, h2') || gs('h2, h1'));
```

你应该将 `font-family` 属性的第一个值设置为 `Verdana`。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.include(style?.fontFamily, 'Verdana');
```

你应该将 `font-family` 属性的第二个值设置为另一种 sans-serif 网页安全字体。_提示：我会选择 Tahoma_。

```js
// 可接受的字体：Arial、sans-serif、Helvetica、Tahoma、Trebuchet MS。
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.match(style?.fontFamily, /(Tahoma)|(Arial)|(sans-serif)|(Helvetica)|(Trebuchet MS)/);
```

你应该使用 `h2` 元素选择器来选中 `h2` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h2'));
```

你应该给 `h2` 元素添加一个值为 `4px solid #dfdfe2` 的 `border-bottom` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h2')?.borderBottom, '4px solid rgb(223, 223, 226)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍测验练习项目" />
    <title>无障碍测验</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS 测验</h1>
      <nav>
        <ul>
          <li><a>信息</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">学员信息</h2>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--

```

