---
id: 6145e6eeaa66c605eb087fe9
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

为所有的单选 `input` 添加 `id` 属性，以便将标签链接到它们。第一个 `input` 的 `id` 为 `q1-a1`。其余的 `id` 分别为 `q1-a2`、`q2-a1` 和 `q2-a2`。

# --hints--

你应该给第一个 `input` 元素一个 `id`，其值为 `q1-a1`。

```js
assert.equal(document.querySelectorAll('ul.answers-list > li > label > input')?.[0]?.id, "q1-a1");
```

你应该给第二个 `input` 元素一个 `id`，其值为 `q1-a2`。

```js
assert.equal(document.querySelectorAll('ul.answers-list > li > label > input')?.[1]?.id, "q1-a2");
```

你应该给第三个 `input` 元素一个 `id`，其值为 `q2-a1`。

```js
assert.equal(document.querySelectorAll('ul.answers-list > li > label > input')?.[2]?.id, "q2-a1");
```

你应该给第四个 `input` 元素一个 `id`，其值为 `q2-a2`。

```js
assert.equal(document.querySelectorAll('ul.answers-list > li > label > input')?.[3]?.id, "q2-a2");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍访问性测验练习项目" />
    <title>无障碍访问性测验</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS 测验</h1>
      <nav>
        <ul>
          <li><a href="#student-info">信息</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">学生信息</h2>
          <div class="info">
            <label for="student-name">姓名：</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">邮箱：</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">出生日期：</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
            <h3><span class="sr-only">问题</span>1</h3>
            <fieldset class="question" name="html-question-one">
              <legend>
                legend 元素代表其父 fieldset 元素内容的标题。
              </legend>
--fcc-editable-region--
              <ul class="answers-list">
                <li>
                  <label>
                    <input type="radio" />
                  </label>
                </li>
                <li>
                  <label>
                    <input type="radio" />
                  </label>
                </li>
              </ul>
            </fieldset>
          </div>
          <div class="question-block">
            <h3><span class="sr-only">问题</span>2</h3>
            <fieldset class="question" name="html-question-two">
              <legend>
                嵌套 input 元素的 label 元素需要有一个 for 属性，其值与 input 的 id 相同。
              </legend>
              <ul class="answers-list">
                <li>
                  <label>
                    <input type="radio" />
                  </label>
                </li>
                <li>
                  <label>
                    <input type="radio" />
                  </label>
                </li>
              </ul>
--fcc-editable-region--
            </fieldset>
          </div>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  padding: 0;
  margin: -1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  white-space: nowrap;
  border: 0;
}

```

