---
id: 6145f14f019a4b0adb94b051
title: 第 37 步
challengeType: 0
dashedName: step-37
---

# --description--

在第一个 `div.answer` 元素中，嵌套一个必需的 `select` 元素，其中包含三个 `option` 元素。

为第一个 `option` 元素设置 `value` 为 `""`，文本为 `选择一个选项`。为第二个 `option` 元素设置 `value` 为 `yes`，文本为 `是`。为第三个 `option` 元素设置 `value` 为 `no`，文本为 `否`。

# --hints--

你应该在第一个 `div.answer` 元素中嵌套一个 `select` 元素。

```js
assert.exists(document.querySelector('div.answer > select'));
```

你应该在 `select` 元素中嵌套三个 `option` 元素。

```js
assert.equal(document.querySelectorAll('div.answer > select > option')?.length, 3);
```

你应该给第一个 `option` 元素的 `value` 属性设置为 `""`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(1)')?.value, '');
```

你应该给第一个 `option` 元素的文本内容设置为 `选择一个选项`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(1)')?.textContent, '选择一个选项');
```

你应该给第二个 `option` 元素的 `value` 属性设置为 `yes`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(2)')?.value, 'yes');
```

你应该给第二个 `option` 元素的文本内容设置为 `是`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(2)')?.textContent, '是');
```

你应该给第三个 `option` 元素的 `value` 属性设置为 `no`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(3)')?.value, 'no');
```

你应该给第三个 `option` 元素的文本内容设置为 `否`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(3)')?.textContent, '否');
```

你应该给 `select` 元素一个 `required` 属性。

```js
assert.isTrue(document.querySelector('div.answer > select')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp 无障碍化测验练习项目" />
    <title>无障碍化测验</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS 测验</h1>
      <nav>
        <ul>
          <li><a href="#student-info">信息</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">学员信息</h2>
          <div class="info">
            <label for="student-name">姓名：</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">邮箱：</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">出生日期：</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
            <h3><span class="sr-only">问题</span>1</h3>
            <fieldset class="question" name="html-question-one">
              <legend>
                legend 元素代表其父 fieldset 元素内容的标题
              </legend>
              <ul class="answers-list">
                <li>
                  <label for="q1-a1">
                    <input type="radio" id="q1-a1" name="q1" value="true" />
                    正确
                  </label>
                </li>
                <li>
                  <label for="q1-a2">
                    <input type="radio" id="q1-a2" name="q1" value="false" />
                    错误
                  </label>
                </li>
              </ul>
            </fieldset>
          </div>
          <div class="question-block">
            <h3><span class="sr-only">问题</span>2</h3>
            <fieldset class="question" name="html-question-two">
              <legend>
                嵌套 input 元素的 label 元素需要具有与 input 的 id 值相同的 for 属性
              </legend>
              <ul class="answers-list">
                <li>
                  <label for="q2-a1">
                    <input type="radio" id="q2-a1" name="q2" value="true" />
                    正确
                  </label>
                </li>
                <li>
                  <label for="q2-a2">
                    <input type="radio" id="q2-a2" name="q2" value="false" />
                    错误
                  </label>
                </li>
              </ul>
            </fieldset>
          </div>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
          <div class="formrow">
            <div class="question-block">
              <label>CSS margin 属性可以接受负值吗？</label>
            </div>
--fcc-editable-region--
            <div class="answer">

            </div>
--fcc-editable-region--
            <div class="question-block">
              <label>你有什么问题吗：</label>
            </div>
            <div class="answer">
            </div>
          </div>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

h3::before {
  content: "问题 #";
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  padding: 0;
  margin: -1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  white-space: nowrap;
  border: 0;
}

```

