---
id: 614ccc21ea91ef1736b9b578
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次练习中，你将获得一个基础的样板文件。

首先，请为你的 `html` 元素添加一个 `lang` 属性。这将帮助屏幕阅读器识别页面的语言。

另外，记得将样式表链接到页面中。

# --hints--

你应该为 `html` 元素添加一个 `lang` 属性。_提示：你可以使用 `en` 作为英文的值。_

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

你应该在 `head` 元素中包含一个 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素应包含 `rel="stylesheet"` 和 `href="styles.css"`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

