---
id: 66db57ad34c7089b9b41bfd6
title: 构建一个灯箱查看器
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

灯箱（Lightbox）常用于网站上，通过弹出窗口或模态窗口在页面主内容之上展示多媒体内容，使内容展示更加生动。

在本实验中，你将创建一个灯箱画廊，当点击缩略图时显示全尺寸图片。每张图片都提供了两个版本：缩略图和全尺寸图片。全尺寸图片与缩略图相同，只是去掉了 `-thumbnail` 后缀。

完成以下用户故事并通过所有测试即可完成本实验。

**用户故事：**

1. 你需要在 `body` 内部有一个 `class` 为 `gallery` 的 `div`。
1. 在 `.gallery` 元素内，你需要有三张缩略图图片，每张图片的 `class` 都为 `gallery-item`。请使用以下缩略图链接：

  - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
  - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
  - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. 你需要在 `body` 内部有一个 `class` 为 `lightbox` 的 `div`。
1. 你需要在 `.lightbox` 元素内有一个 `id` 为 `close-btn` 的 `button`。你可以将其文本设置为 `&times;`。
1. 你需要在 `.lightbox` 元素内有一个 `id` 为 `lightbox-image` 的 `img`。
1. 你的 `.lightbox` 元素应使用固定定位（fixed），以便预览窗口显示在当前图片之上。
1. 你的 `.lightbox` 元素应通过设置高度和宽度为容器的 100% 来覆盖整个视口，并确保 `.lightbox` 元素从容器的左上角开始。
1. `.lightbox` 应有背景色。初始时，其 `display` 属性应设置为 `none`，以隐藏它。
1. 当点击任意 `.gallery-item` 元素时，`.lightbox` 元素的 `display` 属性应设置为 `flex`，使 `.lightbox` 及其内部两个元素可见。
1. 当点击任意 `.gallery-item` 元素时，`#lightbox-image` 元素的 `src` 应设置为所点击图片的全尺寸版本（通过去除图片 `src` 属性中的 `-thumbnail`）。全尺寸图片链接如下：

  - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
  - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
  - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`

1. 当 `.lightbox` 元素可见且点击 `#close-btn` 或 `.lightbox` 元素时，`.lightbox` 元素的 `display` 应恢复为 `none`。

**注意：** 请确保在 HTML 中正确链接样式表和 JavaScript 文件。

# --hints--

你需要在 `body` 元素内有一个 `class` 为 `gallery` 的 `div`。

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

在 `.gallery` 元素内，你需要有三个带有 `gallery-item` 类的 `img` 元素。

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

在 `.gallery` 元素内，你需要有一个 `img` 元素，其 `src` 为 `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
  srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 `.gallery` 元素内，你需要有一个 `img` 元素，其 `src` 为 `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
  srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 `.gallery` 元素内，你需要有一个 `img` 元素，其 `src` 为 `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
  srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

你需要在 `body` 元素内有一个 `class` 为 `lightbox` 的 `div`。

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

在 `.lightbox` 元素内，你需要有一个 `id` 为 `close-btn` 的 `button` 元素。

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

在 `.lightbox` 元素内，你需要有一个 `id` 为 `lightbox-image` 的 `img` 元素。

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

你的 `.lightbox` 元素应为固定定位（fixed）。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

你的 `.lightbox` 元素应覆盖整个视口。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.width, "100%");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.height, "100%");
```

你的 `.lightbox` 元素应与容器的左上角对齐。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

你的 `.lightbox` 元素应有背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

你的 `.lightbox` 元素初始应为隐藏状态。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

当点击任意 `.gallery-item` 元素时，`.lightbox` 元素的 `display` 属性应设置为 `flex`，使 `.lightbox` 及其内部两个元素可见。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

当点击任意 `.gallery-item` 元素时，`#lightbox-image` 元素的 `src` 应设置为所点击图片的全尺寸版本（去除图片 `src` 属性中的 `-thumbnail`）。

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

当 `.lightbox` 元素可见且点击 `#close-btn` 按钮时，`.lightbox` 元素的 `display` 应恢复为 `none`。

```js
const lightbox = document.querySelector(".lightbox");
const background = document.getElementById("close-btn"); 

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

background.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

当 `.lightbox` 元素可见且点击 `.lightbox` 元素时，`.lightbox` 元素的 `display` 应恢复为 `none`。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

lightbox.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Lightbox Viewer</title>
  <link rel="stylesheet" href="styles.css">
  </head>

  <body>
  <h1 class="title">Museum Gallery</h1>

  <div class="gallery">
    <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
    <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
    <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
  </div>
  <div id="lightbox" class="lightbox">
    <button id="close-btn" class="close-btn">&times;</button>
    <img id="lightbox-image" src="" alt="Lightbox Image">
  </div>
  <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
  closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
  openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```

