---
id: 669e2f60e83c011754f711f9
title: 构建一个旅行社页面
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个 `DOCTYPE` 声明。
1. 你应该有一个 `html` 元素，并将 `lang` 设置为 `en`。
1. 你应该有一个 `head` 元素，包含一个 `charset` 设置为 `utf-8` 的 `meta` 空元素，以及一个文本为 `Travel Agency Page` 的 `title` 元素。
1. 你应该在 `head` 元素中有一个 `meta` 标签，为你的网站提供简短的描述以用于 SEO。
1. 你应该有一个 `h1` 元素来展示你的旅游目的地。
1. 你应该在 `h1` 元素下方有一个介绍旅游机会的段落。
1. 你应该有一个文本为 `Packages` 的 `h2` 元素。
1. 你应该有一个简要介绍各种套餐的 `p` 元素。
1. 你应该有一个无序列表元素，包含两个列表项。两个列表项的文本分别为 `Group Travels` 和 `Private Tours`。每个列表项的文本都应由锚点元素包裹。
1. 你应该有一个文本为 `Top Itineraries` 的 `h2` 元素。
1. 你应该至少有三个 `figure` 元素，每个都包含一个锚点元素和一个 `figcaption` 元素。
1. 这三个锚点元素都应包含一个带有合适 `alt` 属性和有效 `src` 属性的 `img` 元素。你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` 和 `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`。
1. 你的五个锚点元素都应有 `href="https://www.freecodecamp.org/learn"` 和 `target="_blank"` 属性。

# --hints--

你的旅行社页面应有 `!DOCTYPE html` 声明。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `lang` 设置为 `en` 的 `html` 元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

你应该在 `html` 元素内有一个 `head` 元素。

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你的 `head` 元素内应有两个 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

其中一个 `meta` 元素应有 `name="description"` 属性，且 `content` 属性非空。

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

另一个 `meta` 元素应有 `charset="UTF-8"` 属性。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你应该在 `head` 元素内有 `title` 元素。

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

你的 `title` 元素应包含你的旅行社名称。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你应该在 `html` 元素内有 `body` 元素。

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你应该有一个用于展示旅游目的地的 `h1` 元素。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你只能有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该在 `h1` 元素下方有一个 `p` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

你的第一个 `p` 元素应介绍旅游机会。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一个 `h2` 元素文本应为 `Packages`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Packages');
```

你的第一个 `h2` 元素下方应有一个 `p` 元素。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

你的第二个 `p` 元素应简要介绍各种套餐。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

你的第二个 `p` 元素下方应有一个无序列表元素。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

你的无序列表应有两个列表项。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

你的两个列表项都应包含一个锚点元素。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

第一个列表项的锚点元素应包裹文本 `Group Travels`。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText, 'Group Travels');
```

第二个列表项的锚点元素应包裹文本 `Private Tours`。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText, 'Private Tours');
```

你的无序列表后应有一个 `h2` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

你的第二个 `h2` 元素文本应为 `Top Itineraries`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Top Itineraries');
```

你应至少有三个 `figure` 元素。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

每个 `figure` 元素的第一个子元素应为锚点元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

每个 `figure` 元素的第二个子元素应为 `figcaption` 元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

每个 `figcaption` 应包含一些文本。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

每个作为 `figure` 子元素的 `a` 元素都应包含一个图片。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

每个 `img` 元素应有有效的 `src` 属性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

每个 `img` 元素应有合适的 `alt` 属性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

每个 `a` 元素都应有 `href="https://www.freecodecamp.org/learn"` 属性。不要忘记列表项中的链接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

每个 `a` 元素都应有 `target="_blank"` 属性。不要忘记列表项中的链接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```

