---
id: 66bb3e20d3dc5b6d0a21f5dd
title: 构建一个实时计数器
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

在本实验中，你将使用 JavaScript 构建一个实时字符计数器。该计数器会显示在 `textarea` 元素中输入的字符数，并且会随着用户输入实时更新。

请完成以下用户故事，并通过所有测试以完成本实验。**不要复制此演示项目**。

**用户需求：**

1. 你应该有一个 `id` 为 `text-input` 的 `textarea` 元素。
2. 应该有一个 `id` 为 `char-count` 的 `p` 元素。它初始内容为 `Character Count: 0/50`。该占位文本应被 `textarea` 元素中的当前字符数替换。
3. 当 `#text-input` 元素内容为 `hello coder` 时，`#char-count` 元素内容应为 `"Character Count: 11/50"`。
4. 当 `#text-input` 元素内容为 `i am a programmer` 时，`#char-count` 元素内容应为 `"Character Count: 17/50"`。
5. 当 `#text-input` 元素内容为 `hello world` 时，`#char-count` 元素内容应为 `"Character Count: 11/50"`。
6. 当 `#text-input` 元素内容为 `I am learning a new language and it is called c++.` 时，`#char-count` 元素内容应为 `"Character Count: 50/50"`。**注意**：虽然 `maxlength` 属性可以实现该功能，但本实验要求仅用 JavaScript 实现。
7. 用户不能在 `textarea` 元素中输入超过 `50` 个字符。当字符数达到 `50` 时，`char-count` 中的文本应显示为红色。

**注意：** 请确保在 HTML 中正确链接样式表和 JavaScript 文件。

# --hints--

你应该有一个 `id` 为 `text-input` 的 `textarea` 元素。

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

你应该有一个 `id` 为 `char-count` 的 `p` 元素。

```js
const charCount = document.querySelector('p');
assert.equal(charCount.id, 'char-count');
```

`#char-count` 元素初始内容应为 `Character Count: 0/50`。

```js
const charCount = document.querySelector('p')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

当 `#text-input` 元素内容为 `hello coder` 时，`#char-count` 元素内容应为 `"Character Count: 11/50"`。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

当字符数为 `50` 时，文本应显示为红色。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
    "red",
    "rgb(255, 0, 0)",
    "rgb(100%, 0%, 0%)",
    "rgba(255, 0, 0, 1)",
    "rgba(100%, 0%, 0%, 1)",
    "hsl(0, 100%, 50%)",
    "hsla(0, 100%, 50%, 1)"
]);
```

如果字符数大于或等于 `50`，用户不能再输入更多字符。

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
        <meta charset="utf-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>Real-time Character Counter</title>
        <link rel="stylesheet" href="styles.css">

<body>
        <div class="container">
                <h1>Real-Time Character Counter</h1>
                <textarea id="text-input" placeholder="Type something..."></textarea>
                <p id="char-count">Character Count: 0/50</p>
        </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
        display: flex;
        justify-content: center;
        align-items: center;
        height: 100vh;
        font-family: Arial, sans-serif;
        background-color: #c0c1ec;
        margin: 0;
}

h1{
        background-color: #f2f48b;
        padding-left: 10px;
        padding-right: 10px;
}
.container {
        text-align: center;
}

textarea {
        width: 300px;
        height: 100px;
        font-size: 18px;
        padding: 10px;
        border: 1px solid #ccc;
        border-radius: 5px;
        resize: none;
        display: block;
        margin: 0 auto;
        border: 1px solid slategray;
        box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
        outline: none;
        border-color: #f2f48b;
}

#char-count {
        margin-top: 10px;
        font-size: 18px;
        color: darkslategray;
        font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
    let textInput = document.getElementById("text-input");
    let charCount = textInput.value.length;

    if (charCount > 50) {
        textInput.value = textInput.value.substring(0, 50);
        charCount = 50;
    }

    let charCountDisplay = document.getElementById("char-count");
    charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

    if (charCount === 50) {
        charCountDisplay.style.color = "red";
    } 

});
```

