---
id: 66eaddd04a9e533fba689001
title: 设计一个博客文章卡片
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

在本实验中，你将通过创建一个博客文章卡片来练习如何为背景和边框添加样式。

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个带有 `blog-post-card` 类的 `div` 元素，用于包含所有卡片内容。
2. 在 `.blog-post-card` 元素内，你应该有一个带有有效 `alt` 属性和文本的图片元素，并且类名为 `post-img`。你可以将图片的 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg`。
3. 在 `.blog-post-card` 元素内，还应有一个类名为 `post-content` 的 `div`，其中包含：
    - 一个带有 `post-title` 类的 `h2` 元素，作为博客文章的标题。
    - 一个带有 `post-excerpt` 类的 `p` 元素，用于总结你的博客文章。
    - 一个带有 `read-more` 类的 `a` 元素，文本为 `Read More`。
4. 你应为 `.blog-post-card` 元素应用以下样式：
    - 白色背景。
    - 圆角。
    - 你自定义的宽度。
    - 你自定义的文本对齐方式。
5. `.post-img` 元素应设置为图片填满卡片的整个宽度，并有底部边框。
6. `.post-content` 元素应设置卡片内部的内边距。
7. `.post-title` 和 `.post-excerpt` 元素应有非默认的文本颜色，并在四周有外边距。
8. `.read-more` 元素应被设置为按钮样式，并具有：
    - 非默认的文本颜色。
    - 背景色。
    - 四周有外边距。
    - 四周有内边距。
    - 圆角。
    - `display` 属性设置为 `inline-block`。
    - 悬停时背景色发生变化的效果。

**注意：** 请确保在 HTML 中正确引入你的样式表并应用 CSS。

# --hints--

你应该有一个带有 `blog-post-card` 类的 `div`。

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

你应该有一个带有 `post-img` 类的 `img` 元素。确保图片有带文本的 `alt` 属性和非空的 `src` 属性。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

你应该有一个带有 `post-content` 类的 `div`。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

你应该有一个带有 `post-title` 类的 `h2` 元素，并有标题文本。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

你应该有一个带有 `post-excerpt` 类的 `p` 元素，并有摘要文本。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

你应该有一个带有 `read-more` 类的 `a` 元素。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

你的 `.read-more` 元素文本应为 `Read More`。

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent, "Read More");
```

你的 `.blog-post-card` 元素应有大于 0 的 `border-radius` 属性值。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

你的 `.blog-post-card` 元素应有白色背景。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

你应为 `.blog-post-card` 设置 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

你应为 `.blog-post-card` 设置 `text-align` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

你应为 `.post-content` 设置 `padding` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

你的 `.read-more` 元素应有悬停效果。

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

你应为 `.read-more` 设置 `color` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

你应为 `.read-more` 设置 `background-color` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

你应为 `.read-more` 设置 `margin` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

你应为 `.read-more` 设置 `display` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

你应为 `.read-more` 设置 `border-radius` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

你应为 `.read-more` 设置 `padding` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

你的 `.post-img` 元素应填满卡片宽度。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

你的 `.post-img` 元素应有 `border-bottom` 属性值。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

你的 `.post-title` 和 `.post-excerpt` 元素应有外边距和非默认文本颜色。

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// 是否使用了 CSS 选择器为元素设置样式：
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```

