---
id: 66def5467aee701733aaf8cc
title: 构建书签管理器应用
challengeType: 25
dashedName: build-a-bookmark-manager-app
demoType: onClick
---

# --description--

在本实验中，所有的 HTML 和 CSS 都已为你提供。你需要使用 JavaScript 完成书签管理器应用，使其能够在本地存储中保存书签集合并从中读取数据。

请完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要有一个名为 `getBookmarks` 的函数，该函数返回存储在本地存储中的 `bookmarks` 数组。如果本地存储中还没有设置 `bookmarks` 键，或者其内容不是有效的书签数组，则 `getBookmarks` 应返回一个空数组。
1. 本地存储中的 `bookmarks` 键应为书签对象数组。每个书签对象应包含三个键：`name`、`category` 和 `url`。
1. 你需要有一个名为 `displayOrCloseForm` 的函数，用于切换 `#main-section` 和 `#form-section` 的 `hidden` 类。
1. 当点击 `#add-bookmark-button` 时，应将 `.category-name` 的文本内容更新为 `#category-dropdown` 下拉菜单中选中的值，并调用 `displayOrCloseForm`，以显示表单区域并隐藏主区域。
1. 当点击 `#close-form-button` 时，应运行你的函数以隐藏表单区域并显示主区域。
1. 当点击 `#add-bookmark-button-form` 时，应通过在数组末尾添加一个书签对象来更新本地存储中的 `bookmarks` 键。该对象的 `name` 应为 `#name` 输入框的值，`category` 为分类下拉菜单选中的值，`url` 为 `#url` 输入框的值。
1. 一旦更新了 `bookmarks` 键，应将 `#name` 和 `#url` 的值重置为空字符串，然后运行你的函数以隐藏表单区域并显示主区域。
1. 你需要有一个名为 `displayOrHideCategory` 的函数，用于切换 `#main-section` 和 `#bookmark-list-section` 的 `hidden` 类。
1. 当点击 `#view-category-button` 时，应将 `.category-name` 的文本内容更新为 `#category-dropdown` 下拉菜单中选中的值，根据以下用户故事修改 `#category-list` 的 inner HTML，并调用 `displayOrHideCategory` 函数。
1. 如果本地存储中的书签没有属于该分类的，应将 `#category-list` 的 inner HTML 设置为一个文本为 `No Bookmarks Found` 的 `p` 元素。
1. 如果本地存储中有一个或多个书签属于所选分类，应为每个书签在 `#category-list` 的 inner HTML 中添加一个单选按钮，`id` 和 `value` 属性都设置为书签名，并且所有单选按钮应有相同的 `name` 属性。
1. 每个单选按钮都应有一个对应的 label，label 内包含一个锚点元素，锚点文本为书签名，`href` 属性为书签 URL。
1. 当点击 `#close-list-button` 时，应运行你的函数以隐藏 `#bookmark-list-section` 并显示主区域。
1. 当点击 `#delete-bookmark-button` 时，应删除本地存储中对应所选单选按钮和分类的书签，并更新显示的书签列表。

# --hints--

你需要有一个 `getBookmarks` 函数。

```js
assert.isFunction(getBookmarks);
```

你的 `getBookmarks` 函数应返回一个数组。

```js
assert.isArray(getBookmarks());
```

你的 `getBookmarks` 函数应返回本地存储中的 `bookmarks` 键。

```js
try {
  setLocalStorage()
  const testArr = getBookmarks();
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(testArr, expected);
} finally {
  resetLocalStorage();
}
```

当本地存储中未设置 `bookmarks` 键或其为一个空数组时，`getBookmarks` 应返回空数组。

```js
try {
  localStorage.setItem("bookmarks", JSON.stringify([]));
  assert.isEmpty(getBookmarks());
  localStorage.removeItem("bookmarks");
  assert.isEmpty(getBookmarks());
} finally {
  resetLocalStorage();
}
```

当本地存储中的 `bookmarks` 键不是有效的书签对象数组时，`getBookmarks` 应返回空数组。

```js
try {
  localStorage.setItem("bookmarks", 'invalid');
  const arrayFromInvalidValue = getBookmarks();
  assert.isArray(arrayFromInvalidValue);
  assert.isEmpty(arrayFromInvalidValue);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", notUrl: "example1.com"}]));
  const arrayFromInvalidArray = getBookmarks();
  assert.isArray(arrayFromInvalidArray)
  assert.isEmpty(arrayFromInvalidArray)
} finally {
  resetLocalStorage();
}
```

你需要有一个名为 `displayOrCloseForm` 的函数。

```js
assert.isFunction(displayOrCloseForm);
```

你的 `displayOrCloseForm` 函数应切换 `#main-section` 和 `#form-section` 的 `hidden` 类。

```js
const hidden = getHidden("form");
displayOrCloseForm();
const hiddenAfterCall = getHidden("form");
assert.isDefined(hidden, hiddenAfterCall);
assert.notEqual(hidden, hiddenAfterCall);
```

点击 `#add-bookmark-button` 时，应将 `.category-name` 的文本内容更新为 `#category-dropdown` 选中的值。

```js
addBookMarkButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText;
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

点击 `#add-bookmark-button` 时，应调用 `displayOrCloseForm`，显示表单区域并隐藏主区域。

```js
assert.strictEqual(getHidden("form"), "form");
addBookMarkButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "main");
```

点击 `#close-form-button` 时，应调用 `displayOrCloseForm`，隐藏表单区域并显示主区域。

```js
assert.strictEqual(getHidden("form"), "main");
closeFormButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "form");
```

点击 `#add-bookmark-button-form` 时，应通过在数组末尾添加对象来更新本地存储中的 `bookmarks` 键。添加的对象的 `name`、`category`、`url` 分别为 `#name`、分类下拉菜单、`#url` 的值。

```js
setLocalStorage();
try {
  bookmarkNameTest.value = "test example5";
  bookmarkURLTest.value = "test example5.com";
  const bookmarkObject = {
    name: bookmarkNameTest.value,
    category: categoryDropdownTest.value,
    url: bookmarkURLTest.value
  };
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}, bookmarkObject];
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.deepEqual(getBookmarks(), expected);
} finally {
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
  resetLocalStorage();
}
```

点击 `#add-bookmark-button-form` 时，应将 `#name` 和 `#url` 的值重置为空字符串。

```js
try {
  bookmarkNameTest.value = "test example2";
  bookmarkURLTest.value = "test example2.com";
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(bookmarkNameTest.value, "");
  assert.strictEqual(bookmarkURLTest.value, "");
} finally {
  resetLocalStorage()  
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
}
```

点击 `#add-bookmark-button-form` 时，应运行 `displayOrCloseForm`，隐藏表单区域并显示主区域。

```js
try {
  bookmarkNameTest.value = "test example3";
  bookmarkURLTest.value = "test example3.com";  
  addBookMarkButtonTest.dispatchEvent(new Event("click"));  
  assert.strictEqual(getHidden("form"), "main");
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(getHidden("form"), "form");
} finally {
  resetLocalStorage()  
}
```

你需要有一个名为 `displayOrHideCategory` 的函数。

```js
assert.isFunction(displayOrHideCategory);
```

你的 `displayOrHideCategory` 函数应切换 `#main-section` 和 `#bookmark-list-section` 的 `hidden` 类。

```js
const hidden = getHidden("bookmark list");
displayOrHideCategory();
const hiddenAfterCall = getHidden("bookmark list");
assert.isDefined(hidden, hiddenAfterCall)
assert.notEqual(hidden, hiddenAfterCall)
```

点击 `#view-category-button` 时，应将 `.category-name` 的文本内容更新为 `#category-dropdown` 选中的值。

```js
viewCategoryButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText;
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

点击 `#view-category-button` 时，如果本地存储中没有该分类的书签，应在 `#category-list` 的 inner HTML 中添加一个文本为 `No Bookmarks Found` 的 `p` 元素。

```js
try {
  setLocalStorage()
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const categoryList = document.getElementById("category-list").innerHTML;
  assert.strictEqual(categoryList, "<p>No Bookmarks Found</p>")
} finally {
  resetLocalStorage()
  clearCategoryList()
}
```

点击 `#view-category-button` 时，应通过为每个属于所选分类的书签添加单选按钮来修改 `#category-list` 元素的 inner HTML。每个单选按钮的 `id` 和 `value` 属性都应为书签名，并且所有单选按钮应有相同的 `name` 属性。

```js
try {
  setLocalStorage();
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(bookmarksDisplayed, 2);
  assert.strictEqual(bookmarksDisplayed[0].id, "example1");
  assert.strictEqual(bookmarksDisplayed[0].value, "example1");
  assert.strictEqual(bookmarksDisplayed[1].id, "example4");
  assert.strictEqual(bookmarksDisplayed[1].value, "example4");
  assert.isOk(bookmarksDisplayed[0].name)
  assert.strictEqual(bookmarksDisplayed[0].name, bookmarksDisplayed[1].name);
} finally {
  resetLocalStorage();
}
```

添加到 `#category-list` 的每个单选按钮都应有一个对应的 label，label 内包含一个锚点元素，锚点文本为书签名，`href` 属性为书签 URL。

```js
setLocalStorage();
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarkLabelsDisplayed, 2);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example1");
  assert.strictEqual(bookmarkLabelsDisplayed[1].htmlFor, "example4");
} finally {
  resetLocalStorage();
}
```

每个 `label` 元素应包含一个锚点元素，锚点文本为书签名，`href` 属性为书签 URL。

```js
setLocalStorage();
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkAnchorsDisplayed = document.querySelectorAll('#category-list label>a');
  assert.lengthOf(bookmarkAnchorsDisplayed, 2);
  assert.strictEqual(bookmarkAnchorsDisplayed[0].innerText, "example1");
  assert.isTrue(bookmarkAnchorsDisplayed[0].href.endsWith("example1.com"));
  assert.strictEqual(bookmarkAnchorsDisplayed[1].innerText, "example4");
  assert.isTrue(bookmarkAnchorsDisplayed[1].href.endsWith("example4.com"));
} finally {
  resetLocalStorage();
}
```

点击 `#close-list-button` 时，应隐藏 `#bookmark-list-section` 并显示主区域。

```js
viewCategoryButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "main");
closeListButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "bookmark list");
```

点击 `#close-list-button` 后再次打开任意分类时，`#category-list` 只应包含所选分类的数据，不应重复。

```js
setLocalStorage();
try {
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const categoryList = document.getElementById("category-list").innerHTML;
  assert.strictEqual(categoryList, "<p>No Bookmarks Found</p>");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "news";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const newsBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(newsBookmarks, 2);
  assert.strictEqual(newsBookmarks[0].id, "example1");
  assert.strictEqual(newsBookmarks[0].value, "example1");
  assert.strictEqual(newsBookmarks[1].id, "example4");
  assert.strictEqual(newsBookmarks[1].value, "example4");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "work";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const workBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(workBookmarks, 1);
  assert.strictEqual(workBookmarks[0].id, "example3");
  assert.strictEqual(workBookmarks[0].value, "example3");

  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const paragraphs = document.querySelectorAll("#category-list p");
  assert.lengthOf(paragraphs, 1);
  assert.strictEqual(paragraphs[0].innerText, "No Bookmarks Found");
} finally {
  resetLocalStorage();
  clearCategoryList();
}
```

点击 `#delete-bookmark-button` 时，应删除本地存储中对应所选单选按钮和分类的书签，并更新显示的书签列表。

```js
setLocalStorage();
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const firstRadioButton = document.querySelector('#category-list input[type="radio"]');
  firstRadioButton.checked = true;
  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));
  let bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  let bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarksDisplayed, 1);
  assert.strictEqual(bookmarksDisplayed[0].id, "example4");
  assert.strictEqual(bookmarksDisplayed[0].value, "example4");
  assert.lengthOf(bookmarkLabelsDisplayed, 1);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example4");
  const expected = [{name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(getBookmarks(), expected);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "duplicated-name", category: "news", url: "example1.com"}, {name: "duplicated-name", category: "entertainment", url: "example2.com"}]));

  categoryDropdownTest.value = "entertainment";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const radioForDuplicate = document.querySelector('#category-list input[type="radio"]');
  radioForDuplicate.checked = true;

  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));

  assert.deepEqual(getBookmarks(), [{name: "duplicated-name", category: "news", url: "example1.com"}])
} finally {
  resetLocalStorage();
}
```

# --seed--

## --after-user-code--

```js
const viewCategoryButtonTest = document.getElementById("view-category-button");
const closeListButtonTest = document.getElementById("close-list-button");
const closeFormButtonTest = document.getElementById("close-form-button");
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const deleteBookmarkButtonTest = document.getElementById("delete-bookmark-button");
const bookmarkNameTest = document.getElementById("name");
const categoryDropdownTest = document.getElementById("category-dropdown");
const bookmarkURLTest = document.getElementById("url");
const mainSection = document.querySelector("#main-section");
const formSection = document.querySelector("#form-section");
const bookmarkListSection = document.querySelector("#bookmark-list-section");
const getHidden = (notMain) => {
  if (notMain === "form") {
    if (mainSection.classList.contains("hidden") && !formSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && formSection.classList.contains("hidden")) {
      return notMain;
    }
  } else if (notMain === "bookmark list") {
    if (mainSection.classList.contains("hidden") && !bookmarkListSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && bookmarkListSection.classList.contains("hidden")) {
      return notMain;
    }
  }  
}

const originalLocalStorage = JSON.parse(localStorage.getItem("bookmarks")) || [];
const temp = originalLocalStorage;
function resetLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify(temp));
}

function setLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}]));
}

function clearCategoryList() {
  const categoryList = document.getElementById("category-list");
  categoryList.innerHTML = "";
}
```

## --seed-contents--

```js

```

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=0.0">
    <title>Bookmark Manager</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <section id="main-section">
            <h1>Bookmark Manager</h1>
            <div id="dropdown">
                <label for="category-dropdown">Select a category:</label>
                <select id="category-dropdown" name="options">
                    <option value="news" selected>News</option>
                    <option value="entertainment">Entertainment</option>
                    <option value="work">Work</option>
                    <option value="miscellaneous">Miscellaneous</option>
                </select>
            </div>
            <div id="buttons">
                <button type="button" id="view-category-button">View Category</button>
                <button type="button" id="add-bookmark-button">Add Bookmark</button>
            </div>
        </section>

        <section id="form-section" class="hidden">
            <form>
                <h2 class="category-name"></h2>
                <div>
                    <label for="name">Name:</label>
                    <input type="text" id="name">
                </div>
                <div>
                    <label for="url">URL:</label>
                    <input type="text" id="url">
                </div>
                <div>
                    <button type="button" id="close-form-button">Go Back</button>
                    <button type="button" id="add-bookmark-button-form">Add Bookmark</button>
                </div>
            </form>
        </section>

        <section id="bookmark-list-section" class="hidden">
            <h2 class="category-name"></h2>
            <div id="category-list">
            </div>
            <div>
                <button type="button" id="close-list-button">Go Back</button>
                <button type="button" id="delete-bookmark-button">Delete Bookmark</button>
            </div>
        </section>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-grey: #0a0a23;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

main {
  display: flex;
  justify-content: center;
}

body {
  background-color: var(--dark-grey);
}

.hidden {
  display: none;
}

section {
  display: flex;
  flex-direction: column;
  justify-content: center;
}

select,
input,
label {
  margin-left: 10px;
}

div {
  padding: 30px;
  display: flex;
  justify-content: center;
}

.close-form-button {
  background: none;
  border: none;
  cursor: pointer;
}

h1, h2 {
  margin-top: 20px;
  text-align: center;
}

#category-list {
  text-align: center;
  display: flex;
  flex-direction: column;
  justify-content: flex-start;
  background-color: var(--light-grey);
  align-self: center;
  width: 80%;
  margin-top: 15px;
  border-radius: 10px;
}

#category-list,
h1,
h2,
label {
  color: var(--light-grey);
}

#category-list p {
  color: var(--dark-grey);
}

button {
  cursor: pointer;
  padding: 5px;
  width: 100px;
  margin: 10px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

section {
  margin-top: 60px;
  border: 2px solid var(--golden-yellow);
  width: fit-content;
  border-radius: 10px;
}
```

