---
id: 6690e10ebe2181212abc9652
title: 设计一张名片
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成实验。

**用户需求:**

1. 你需要在 HTML 文件的 `head` 部分使用 `link` 标签引入样式表 `styles.css`。
2. 页面背景应设置为 `rosybrown`，整体字体为 `Arial`，备用字体为 `sans-serif`。
3. 所有的 `p` 元素上下 `margin` 应为 `5px`。
4. 页面上的所有链接默认不应有下划线。
5. 你需要有一个 `div`，其 `class` 属性值为 `business-card`，用于包含所有其他元素。
6. `.business-card` 选择器应设置元素宽度为 `300px`，背景色可自选，内边距为 `20px`，顶部外边距为 `100px`，文本居中对齐，字体大小为 `16px`，左右外边距为 `auto`。
7. 在 `.business-card` 元素内部，应有一个 `img` 元素，其 `class` 属性值为 `profile-image`。你可以将图片源设置为 `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`，并为其添加有意义的 `alt` 描述。
8. `.profile-image` 选择器应设置 `max-width` 属性为 `100%`。
9. 在 `.business-card` 元素内部，图片元素之后应有三个 `p` 元素，`class` 属性分别为 `full-name`、`designation` 和 `company`。
10. 第一个 `p` 元素应包含你的名字。
11. 第二个 `p` 元素应包含你的职位。
12. 第三个 `p` 元素应包含你的公司名称。
13. 在第三个 `p` 元素下方应有一个 `hr` 元素。
14. 在 `hr` 元素之后，应有两个 `p` 元素。第一个 `p` 元素文本为你的电子邮件地址，第二个 `p` 元素文本为你的电话号码。紧接着这两个 `p` 元素后，应有一个文本为 `Portfolio` 的 `a` 元素，链接到有效网址。并在该 `a` 元素后添加一个 `hr` 元素。
15. 你还需要有一个 `div` 元素，`class` 属性值为 `social-media`。在该元素内应有一个 `h2` 元素，文本为 `Connect with me`。
16. 在 `.social-media` 元素内应有三个 `a` 元素，文本分别为 `Twitter`、`LinkedIn` 和 `GitHub`，并分别链接到有效网站。

# --hints--

你的 `link` 元素应该在你的 `head` 元素内。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

你应该在 `body` 选择器中将页面背景颜色设置为 `rosybrown`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

你应该在 `body` 元素中将页面字体设置为 `Arial`，备用字体为 `sans-serif`。

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```


你的 `p` 元素应该有 `5px` 的顶部和底部 `margin`。

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```
页面上的所有链接默认不应有下划线。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

你应该有一个主 `div`，其 `class` 属性值为 `business-card`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

你的 `.business-card` 选择器应具有值为 `300px` 的 `width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

你的 `.business-card` 选择器应具有一个有效颜色值的 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

你的 `.business-card` 选择器应具有值为 `20px` 的 `padding` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

你的 `.business-card` 选择器应具有 `100px` 的顶部外边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

`.business-card` 元素内的文本应居中对齐。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

`.business-card` 元素内文本的字体大小应为 `16px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

`.business-card` 元素的左右外边距应设置为 `auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

在 `.business-card` 元素内部，应该有一个 `img` 元素，其 `class` 属性值为 `profile-image`。

```js
assert.exists(document.querySelector('img.profile-image'));
```

图片源应设置为有效图片。

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

图片的 `alt` 属性应设置为有意义的描述。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

你的 `.profile-image` 选择器应具有值为 `100%` 的 `max-width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

在 `.business-card` 元素内部，在 `img` 元素之后，应该有一个 `p` 元素，其 `class` 属性值为 `full-name`。

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

在 `.full-name` 元素之后，应该有一个 `p` 元素，其 `class` 属性值为 `designation`。

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

在 `.designation` 元素之后，应该有一个 `p` 元素，其 `class` 属性值为 `company`。

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

第一个 `p` 元素（类名为 `full-name` 的元素）应包含你的名字。

```js
assert.isAbove(document.querySelector(".full-name")?.innerText.length, 0);
```

第二个 `p` 元素（`class` 属性为 `designation` 的元素）应包含你的职称。

```js
assert.isAbove(document.querySelector(".designation")?.innerText.length, 0);
```

第三个 `p` 元素（`class` 属性值为 `company` 的元素）应包含你的公司名称。

```js
assert.isAbove(document.querySelector(".company")?.innerText.length, 0);
```

在第三个 `p` 元素（`class` 属性值为 `company` 的元素）之后，应该有一个 `hr` 元素。

```js
assert.exists(document.querySelector('p.company + hr'));
```

在第一个 `hr` 元素之后，应该有一个 `p` 元素，其文本为你的电子邮件地址。

```js
assert.isAbove(document.querySelectorAll('.business-card hr + p')[0].innerText.length, 0);
```

在电子邮件 `p` 元素之后，应该有另一个 `p` 元素，其文本为你的电话号码。

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

在电话号码 `p` 元素之后，应该有一个文本为 `Portfolio` 的 `a` 元素。

```js
assert.equal(document.querySelector('.business-card > a')?.textContent, 'Portfolio');
```

你的作品集 `a` 元素应具有有效的 `href`。

```js
assert.isAbove(document.querySelectorAll('.business-card a')[0].href.length, 0);
```

在包含你的作品集的 `a` 元素之后，应该有另一个 `hr` 元素作为分隔线。

```js
assert.exists(document.querySelector('a + hr'));
```

在第二个 `hr` 元素之后，应该有一个 `div` 元素，其 `class` 属性值为 `social-media`。

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

在 `.social-media` 元素内部，应该有一个文本为 `Connect with me` 的 `h2` 元素。

```js
assert.equal(document.querySelector('.social-media h2')?.textContent, 'Connect with me');
```

在 `.social-media` 元素内部应该有三个 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

第一个 `a` 元素的文本应为 `Twitter`。

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent, 'Twitter');
```

第一个 `a` 元素的 `href` 应指向一个有效链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

第二个 `a` 元素的文本应为 `LinkedIn`。

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent, 'LinkedIn');
```

第二个 `a` 元素的 `href` 应指向一个有效链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

第三个 `a` 元素的文本应为 `GitHub`。

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent, 'GitHub');
```

第三个 `a` 元素的 `href` 应指向一个有效链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>名片</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>数字名片</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="一朵盛开的鲜艳花朵" class="profile-image">
        <p class="full-name">你的名字</p>
        <p class="designation">软件开发者</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>邮箱：your-email@example.com</p>
        <p>电话：(123) 456-7890</p>
        <a href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```

