---
id: 668f08ea07b99b1f4a91acab
title: 构建一个食谱页面
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个 `!DOCTYPE html` 声明。
1. 你应该有一个 `html` 元素，并将 `lang` 属性设置为 `en`。
1. 你应该有一个包含 `title` 元素（内容为你的食谱名称）和一个 `charset` 属性为 `UTF-8` 的 `meta` 元素的 `head` 元素。
1. 你应该有一个 `body` 元素。
1. 你应该有一个包含食谱名称的 `h1` 元素。
1. 你应该在 `h1` 元素下方有一个介绍该食谱的 `p` 元素。
1. 你应该有一个文本为 `Ingredients` 的 `h2` 元素，用于配料部分。
1. 你应该在第一个 `h2` 元素下方有一个无序列表（`ul` 元素），其中包含至少四个配料的列表项（`li` 元素）。
1. 你应该有第二个文本为 `Instructions` 的 `h2` 元素，用于步骤部分。
1. 你应该在第二个 `h2` 元素下方有一个有序列表（`ol` 元素），其中包含至少四个步骤的列表项。
1. 你应该有一个 `img` 元素，其 `src` 属性设置为有效图片地址（你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`），并有描述图片的 `alt` 属性。

# --hints--

你的食谱页面应有 `!DOCTYPE html` 声明。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

你应该有一个 `lang` 设置为 `en` 的 `html` 元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

你应该在 `html` 元素内有一个 `head` 元素。

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

你应该在 `head` 元素内有 `title` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `title` 元素应包含你的食谱标题。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你应该在 `head` 元素内有 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `meta` 元素的 `charset` 属性应设置为 `UTF-8`。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你应该在 `html` 元素内有 `body` 元素。

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

你应该有一个包含食谱名称的 `h1` 元素。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你只能有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该在 `h1` 元素下方有一个 `p` 元素。

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

你的第一个 `p` 元素应描述你的食谱。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一个 `h2` 元素文本应为 `Ingredients`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

你应该在第一个 `h2` 元素下方有一个无序列表元素。

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

你的无序列表中应有至少四个配料的列表项。

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你的第二个 `h2` 元素文本应为 `Instructions`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

你应该在第二个 `h2` 元素下方有一个有序列表元素。

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

你的有序列表中应有至少四个步骤的列表项。

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你应该至少有一个 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

所有 `img` 元素应有有效的 `src` 属性和值。

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

所有 `img` 元素应有描述图片的 `alt` 属性。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>巧克力曲奇饼干食谱</title>
</head>

<body>
  <h1>巧克力曲奇饼干</h1>
  <p>欢迎来到迷你巧克力曲奇饼干的终极制作指南！这些一口大小的小饼干非常适合满足你的甜食欲望，又不会吃得太多。按照这个简单的食谱，制作外酥里软、人人都爱的美味迷你巧克力曲奇饼干吧。</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="烘焙原料：三个鸡蛋、一碗面粉、一杯牛奶和一个打蛋器摆放在木桌上。">
  <h2>Ingredients</h2>
  <ul>
    <li>1 杯中筋面粉</li>
    <li>1/2 茶匙小苏打</li>
    <li>1/4 杯无盐黄油（软化）</li>
    <li>1/4 杯白砂糖</li>
    <li>1/2 茶匙香草精</li>
    <li>1/2 杯迷你巧克力豆</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>将烤箱预热至 350°F（175°C），并在烤盘上铺上烘焙纸。</li>
    <li>在一个碗中混合面粉和小苏打。</li>
    <li>在另一个碗中将黄油、糖和香草精打发至顺滑。</li>
    <li>将干性材料逐渐加入湿性材料中，然后拌入迷你巧克力豆。</li>
    <li>用小勺将面团分成小份，放在烤盘上。</li>
    <li>烘烤 8-10 分钟，取出冷却后即可享用！</li>
  </ol>
</body>

</html>
```

