---
id: 66c051d13a6a20255a963695
title: 构建一个风格化的待办事项列表
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

在本实验中，你将练习为链接应用不同的样式，包括悬停、聚焦、点击和访问后的效果。

请完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个带有 `todo-list` 类的无序列表。
2. 在该无序列表中，应该有四个列表项。
3. 每个列表项中应包含：

  - 一个 `input` 元素，类型为 `checkbox`，并且 `id` 设置为唯一值。
  - 一个 `label` 元素，`for` 属性设置为对应的 `input` 元素的 `id`。
  - 一个带有 `sub-item` 类的无序列表。
  - 该无序列表中有一个列表项，里面有一个锚点元素。该锚点应有 `sub-item-link` 类、有效的 `href` 值，并且 `target` 属性设置为在新标签页打开链接。

4. 你的锚点元素默认不应有下划线。
5. 你应将链接的文本颜色设置为你选择的颜色。
6. 当链接被访问后，颜色应变为你选择的另一种颜色。
7. 当链接被悬停时，颜色应变为你选择的另一种颜色。
8. 当链接获得焦点时，应该有彩色轮廓线包围链接。
9. 当链接被点击时，颜色应变为你选择的另一种颜色。

**注意：** 请确保在 HTML 中引入你的样式表并应用 CSS。

# --hints--

你应该有一个带有 `todo-list` 类的无序列表。

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

你应该在无序列表中有四个列表项。

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

`ul.todo-list` 下的 `li` 应包含一个类型为 `checkbox` 的 `input` 元素。

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

`ul.todo-list` 下的 `li` 应包含一个 `label` 元素。

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

所有 `input` 元素都应有 `id` 属性。

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

所有 `label` 元素都应有 `for` 属性。

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

所有 `label` 元素都应有文本内容。

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

`input` 和 `label` 元素的 `id` 和 `for` 属性值应一一对应。

```js
const labels = document.querySelectorAll("label");

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

在 `label` 元素后应有一个带有 `sub-item` 类的无序列表。

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

`ul.sub-item` 下的 `li` 应包含一个带有 `sub-item-link` 类的锚点元素。

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

所有锚点元素都应有有效的 `href` 属性。

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

所有锚点元素都应有文本内容。

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

每个 `a` 元素的 `target` 属性值应为 `_blank`。

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

页面上的链接默认不应有下划线。

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
   ?.textDecoration === 'none' ||
   new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
   new __helpers.CSSHelp(document).getStyle('.sub-item-link')
    ?.textDecoration === 'none'
);
```

链接悬停时应变色。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
   new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
   new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

链接被点击时应变色。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
   new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
   new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

链接获得焦点时应有轮廓线。

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
   css.getStyle('a:focus')?.outline ||
   css.getStyle('a.sub-item-link:focus')?.outline ||
   css.getStyle('.sub-item-link:focus')?.border ||
   css.getStyle('a:focus')?.border ||
   css.getStyle('a.sub-item-link:focus')?.border
);
```

链接被访问后应变色。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
   new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
   new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
   <meta charset="utf-8">
   <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
   <meta charset="UTF-8" />
   <meta name="viewport" content="width=device-width, initial-scale=1.0" />
   <title>风格化待办事项列表</title>
   <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
   <div class="container">
    <h1>我的待办事项</h1>
    <ul class="todo-list">
      <li>
       <input type="checkbox" id="task1" />
       <label for="task1">探索游戏键盘</label>
       <ul class="sub-item">
        <li>
          <a href="#" target="_blank" class="sub-item-link">商店链接</a>
        </li>
       </ul>
      </li>

      <li>
       <input type="checkbox" id="task2" />
       <label for="task2">完成报告</label>
       <ul class="sub-item">
        <li>
          <a href="#" target="_blank" class="sub-item-link"
           >请求访问</a
          >
        </li>
       </ul>
      </li>

      <li>
       <input type="checkbox" id="task3" />
       <label for="task3">查看手机保修</label>
       <ul class="sub-item">
        <li>
          <a href="#" target="_blank" class="sub-item-link"
           >查看收据</a
          >
        </li>
       </ul>
      </li>

      <li>
       <input type="checkbox" id="task4" />
       <label for="task4">检查处理器规格</label>
       <ul class="sub-item">
        <li>
          <a href="#" target="_blank" class="sub-item-link"
           >查看型号</a
          >
        </li>
       </ul>
      </li>
    </ul>
   </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```

