---
id: 66c06fad3475cd92421b9ac2
title: 构建一个闰年计算器
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

闰年是能被 `4` 整除的年份，但能被 `100` 整除且不能被 `400` 整除的年份除外。例如，`2000` 是闰年，但 `1900` 不是。此外，闰年二月有一天额外的日期，即该月的第 29 天。

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 定义一个名为 `isLeapYear` 的函数，该函数接收一个数字作为参数。
2. 在函数外声明一个变量 `year`，用于存储你想要检查的年份。
3. 在函数内部，使用 `if/else` 语句或三元运算符判断该年份是否为闰年。
4. 判断闰年的条件如下：

  - 如果年份能被 `4` 整除，则为闰年。
  - 除非该年份还能被 `100` 整除，则不是闰年。
  - 除非该年份还能被 `400` 整除，则为闰年。

5. 如果年份是闰年，返回 `[year] is a leap year.`，否则返回 `[year] is not a leap year.`。你需要将 `[year]` 替换为 `isLeapYear` 函数中的参数。
6. 你需要调用 `isLeapYear` 函数，并将 `year` 作为参数传入，并将结果赋值给名为 `result` 的变量。
7. 你需要使用 `console.log()` 输出 `result` 变量。

# --hints--

你应该定义一个名为 `isLeapYear` 的函数。

```js
assert.isFunction(isLeapYear);
```

`isLeapYear` 函数应该接收一个数字作为参数。

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

你应该声明一个变量 `year` 并赋值，用于检查是否为闰年。

```js
assert.isDefined(year);
```

`year` 变量不应为空。

```js
assert.isNotNull(year);
```

当 `year` 变量的值为 `2024` 时，`result` 应为 `2024 is a leap year.`。

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

当 `year` 变量的值为 `2000` 时，`result` 应为 `2000 is a leap year.`。

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

当 `year` 变量的值为 `1900` 时，`result` 应为 `1900 is not a leap year.`。

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

你应该调用 `isLeapYear` 函数，并传入 `year` 作为参数。

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

你应该声明一个 `result` 变量。

```js
assert.isDefined(result);
```

你应该将调用 `isLeapYear` 函数的结果存储在名为 `result` 的变量中。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

你应该使用 `console.log()` 输出 `result`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
   return(year + " is a leap year.");
  } else {
   return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```

