---
id: 66bb6a9c2dd58b73cd759034
title: 构建成绩册应用
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个名为 `getAverage` 的函数，它接收一个成绩数组作为参数，并返回平均分。
1. 你应该有一个名为 `getGrade` 的函数，它接收一个学生分数作为参数，并根据分数返回一个表示字母等级的字符串。分数与对应的字母等级如下：

  | 分数范围   | 等级   |
  | ---------- | ------ |
  | `100`      | `"A+"` |
  | `90 - 99`  | `"A"`  |
  | `80 - 89`  | `"B"`  |
  | `70 - 79`  | `"C"`  |
  | `60 - 69`  | `"D"`  |
  | `0 - 59`   | `"F"`  |

1. 你应该有一个名为 `hasPassingGrade` 的函数，它接收一个分数作为参数，并根据分数是否及格返回 `true` 或 `false`。
1. `hasPassingGrade` 函数应使用 `getGrade` 函数获取字母等级，并据此判断是否及格。只要不是 `"F"` 就算及格。
1. 你应该有一个名为 `studentMsg` 的函数，它接收一个成绩数组和一个学生分数作为参数。该函数应返回如下格式的字符串：

  - 如果学生及格，返回：`"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`
  - 如果学生不及格，返回：`"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`

  其中 `average-goes-here` 替换为总成绩的平均分，`grade-goes-here` 替换为学生的等级。使用 `getAverage` 获取平均分，使用 `getGrade` 获取学生等级。

# --hints--

你应该有一个名为 `getAverage` 的函数。

```js
assert.isFunction(getAverage);
```

你的 `getAverage` 函数应返回一个数字。

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 应返回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 应返回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 应返回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

你的 `getAverage` 函数应返回测试分数的平均值。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

你应该有一个名为 `getGrade` 的函数。

```js
assert.isFunction(getGrade);
```

你的 `getGrade` 函数在分数为 `100` 时应返回 `"A+"`。

```js
assert.strictEqual(getGrade(100), "A+");
```

你的 `getGrade` 函数在分数为 `90` 到 `99` 时应返回 `"A"`。

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

你的 `getGrade` 函数在分数为 `80` 到 `89` 时应返回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

你的 `getGrade` 函数在分数为 `70` 到 `79` 时应返回 `"C"`。

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

你的 `getGrade` 函数在分数为 `60` 到 `69` 时应返回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

你的 `getGrade` 函数在分数为 `0` 到 `59` 时应返回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```

你应该有一个名为 `hasPassingGrade` 的函数。

```js
assert.isFunction(hasPassingGrade);
```

你的 `hasPassingGrade` 函数应返回布尔值。

```js
assert.isBoolean(hasPassingGrade(100));
```

你的 `hasPassingGrade` 函数在等级为 `"A"` 时应返回 `true`。

```js
assert.isTrue(hasPassingGrade(100));
```

你的 `hasPassingGrade` 函数在等级为 `"F"` 时应返回 `false`。

```js
assert.isFalse(hasPassingGrade(53));
```

你的 `hasPassingGrade` 函数对于所有及格等级应返回 `true`。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

你应该有一个名为 `studentMsg` 的函数。

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` 应返回：`"Class average: 71.7. Your grade: F. You failed the course."`。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` 应返回：`"Class average: 50.8. Your grade: A+. You passed the course."`。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

你的 `studentMsg` 函数应根据学生分数和班级平均分返回正确的信息。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
  }

  function getGrade(score) {
  if (score === 100) {
    return "A+";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
  }

  function hasPassingGrade(score) {
  return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
  }
```

