---
id: 66f53dc2c5bd6a11d6c3282f
title: 构建一个密码生成器应用
challengeType: 26
dashedName: lab-password-generator
---

# --description--

在本实验中，你将通过构建一个随机密码生成器来练习使用函数。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你应该创建一个名为 `generatePassword` 的函数，该函数带有一个参数，表示生成密码的长度。你可以自行命名该参数。
2. 你的函数应返回一个字符串，表示随机生成的密码。你应使用以下字符串和不同的 `Math` 方法，帮助你返回一个包含随机字符的新字符串：`ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`。
3. 你应定义一个名为 `password` 的变量，并将调用 `generatePassword` 函数的结果赋值给它，参数为所需密码长度的数字。
4. 你应有一个 `console.log`，输出消息 `"Generated password:"`，后跟 `password` 变量。

# --hints--

你应该有一个带参数的 `generatePassword` 函数。你可以自行命名该参数。

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

你的 `generatePassword` 函数应返回一个字符串。

```js
const result = generatePassword(5);
assert.isString(result);
```

你的 `generatePassword` 函数应返回一个长度正确的新字符串。

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

你的函数应返回一个包含有效字符的随机密码。

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

你的函数每次调用都应返回一个新的随机字符串。

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

你应该有一个 `password` 变量。

```js
assert.isDefined(password);
```

你的 `password` 变量应为字符串。

```js
assert.isString(password);
```

你应调用 `generatePassword` 函数，并用数字参数将返回的密码存储在 `password` 变量中。

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const)\s+password\s*=\s*generatePassword\(\d+\)\;/);
```

你应将生成的密码输出到控制台。

```js
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```

