---
id: 66b36358ed4f261d64840c24
title: 构建可用性表格
challengeType: 25
dashedName: build-an-availability-table
demoType: onClick
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该有一个包含至少三列和五行（包括表头）的表格。
1. 第一行的单元格应为表头，内容为一周中的天数。
1. 第一列的单元格应为表头，`class` 为 `time`，内容为时间值。
1. 所有数据单元格应有 `available-#` 的类名，其中 `#` 是 `0` 到 `5` 的数字，表示该时间段可用人数。
1. 在你的 `:root` 选择器中，定义六个 CSS 变量，命名为 `--color#`（`#` 为 `0` 到 `5`），并分别赋予颜色值。使用这些变量为对应的 `.available-#` 元素设置背景色。
1. 在你的 `:root` 选择器中，定义 CSS 变量 `--solid-border` 和 `--dashed-border`。用这些变量为你的数据单元格设置底部边框，根据行的不同交替使用实线和虚线边框。为行添加 `sharp` 或 `half` 类以实现样式。
1. 你应该有一个 `id` 为 `legend` 的 `div` 元素。它应包含一个文本为 `Availability` 的 `span` 元素，以及一个 `id` 为 `legend-gradient` 的 `div` 元素。
1. 你应该为 `#legend-gradient` 元素设置一个线性渐变背景，颜色从 `--color0` 过渡到 `--color5`。每个颜色值应有两个颜色断点（用百分比表示），以实现颜色之间的硬切换。

**注意：** 请确保在 HTML 中正确链接样式表以应用 CSS。

# --hints--

你应该有一个 `table` 元素。

```js
assert.isNotNull(document.querySelector('table'));
```

你应该只有一个 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

你的表格应至少有 `5` 行。

```js
assert.isAtLeast(document.querySelectorAll('tr')?.length, 5);
```

你的表格应至少有 `3` 列。

```js
const rows = [...document.querySelectorAll('tr')];
assert.isNotEmpty(rows);
rows.forEach((row) => {
    assert.isAtLeast(row.children.length, 3);
})
```

表格的第一行应包含至少三个 `th` 元素。

```js
assert.isAtLeast(document.querySelectorAll('tr:first-child>th').length, 3)
```

表格的第一行不应包含 `td` 元素。

```js
const firstRow = document.querySelector('tr');
assert.isNotNull(firstRow);
assert.isEmpty(firstRow.querySelectorAll('td'));
```

你应至少有两行带有 `sharp` 类。

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="sharp"]').length, 2);
```

你应至少有两行带有 `half` 类。

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="half"]').length, 2);
```

第一列的 `th` 元素应有 `time` 类且内容为时间值。

```js
const firstColumnData = [...document.querySelectorAll('tr th:first-child')];
assert.isNotEmpty(firstColumnData);
for (let cell of firstColumnData) {
    if (cell.innerText.match(/\d/)?.length) {
        assert.isTrue(cell.classList.contains('time'));
    }
}
```

你应至少有四个带有 `time` 类且内容为时间值的 `th` 元素。

```js
const timeClass = document.querySelectorAll('.time');
assert.isAtLeast(timeClass.length, 4);
for (let cell of timeClass) {
    assert.match(cell.innerText, /\d/)
}
```

所有 `td` 元素应有 `available-#` 类，其中 `#` 是 `0` 到 `5` 的数字。

```js
const cells = [...document.querySelectorAll('tr[class="sharp"]>td, tr[class="half"]>td')];
assert.isNotEmpty(cells);
for (let cell of cells) {
    const classString = [...cell.classList].join(' ');
    assert.lengthOf(classString.match(/(?<=\s|^)available-[0-5](?=\s|$)/gm), 1);
}
```

你应至少有八个 `.available-#` 元素，其中 `#` 是 `0` 到 `5` 之间的数字。

```js
assert.isAtLeast(document.querySelectorAll('td[class*="available-"]').length, 8);
```

你应在 `:root` 选择器中定义 `--color0` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color0'))
```

你应在 `:root` 选择器中定义 `--color1` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color1'))
```

你应在 `:root` 选择器中定义 `--color2` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color2'))
```

你应在 `:root` 选择器中定义 `--color3` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color3'));
```

你应在 `:root` 选择器中定义 `--color4` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color4'));
```

你应在 `:root` 选择器中定义 `--color5` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color5'));
```

你应使用 `--color0` 设置 `.available-0` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-0');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color0)');
```

你应使用 `--color1` 设置 `.available-1` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-1');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color1)');
```

你应使用 `--color2` 设置 `.available-2` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-2');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color2)');
```

你应使用 `--color3` 设置 `.available-3` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-3');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color3)');
```

你应使用 `--color4` 设置 `.available-4` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-4');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color4)');
```

你应使用 `--color5` 设置 `.available-5` 元素的背景色。

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-5');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color5)');
```

你应在 `:root` 选择器中定义 `--solid-border` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--solid-border'));
```

你应在 `:root` 选择器中定义 `--dashed-border` 变量。

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--dashed-border'));
```

你应选中 `.sharp` 元素下的 `td` 元素。

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
```

你应使用 `--solid-border` 设置 `.sharp` 元素下 `td` 的底部边框。

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--solid-border)');
```

你应选中 `.half` 元素下的 `td` 元素。

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
```

你应使用 `--dashed-border` 设置 `.half` 元素下 `td` 的底部边框。

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--dashed-border)');
```

你应有一个 `id` 为 `legend` 的 `div` 元素。

```js
assert.equal(document.querySelector('#legend')?.tagName, 'DIV');
```

你的 `#legend` 内应有一个文本为 `Availability` 的 `span` 元素。

```js
assert.equal(document.querySelector('#legend span')?.innerText, 'Availability');
```

你的 `#legend` 内应有一个 `id` 为 `legend-gradient` 的 `div` 元素。

```js
assert.equal(document.querySelector('#legend #legend-gradient')?.tagName, 'DIV');
```

你应将 `#legend-gradient` 的背景图像设置为线性渐变。

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient); 
assert.include(legendGradient.backgroundImage, "linear-gradient(")
```

你应为 `#legend-gradient` 使用两个颜色断点（百分比表示），使每种颜色之间的过渡为硬切换，并使用你的 `--color#` 变量。

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient);
const matched = legendGradient.backgroundImage.match(/var\(\s*--color[0-5]\s*\)\s+\d+%(\s+\d+%)?/g);
assert.lengthOf(matched, 6);
matched.forEach((arg, i) => {
  if (i !== 0 && i !== 5) assert.lengthOf(arg.match(/%/g), 2);
})
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <table>
            <tbody>
                <tr class="sharp">
                    <th rowspan="2" class="time">9 AM</th>
                    <th>Monday</th>
                    <th>Tuesday</th>
                    <th>Wednesday</th>
                    <th>Thursday</th>
                    <th>Friday</th>
                </tr>
                <tr class="half">
                    <td class="available-0"></td>
                    <td class="available-1"></td>
                    <td class="available-2"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">10 AM</th>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>

                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">11 AM</th>
                    <td class="available-0"></td>
                    <td class="available-2"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-1"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">12 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">1 PM</th>
                    <td class="available-2"></td>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">2 PM</th>
                    <td class="available-0"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">3 PM</th>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">4 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">5 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                </tr>
                <tr>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                </tr>
            </tbody>
        </table>
        <div id="legend">
            <span id="legend-title">Availability</span>
            <div id="legend-line"><span>0</span>
                <div id="legend-gradient"></div><span>5+</span>
            </div>
        </div>
    </main>
</body>

</html>
```

```css
:root {
    font-size: 16px;
    font-family: Arial, Helvetica, sans-serif;
    --color0: hsl(0, 0%, 100%);
    --color1: hsl(120, 95%, 90%);
    --color2: hsl(120, 95%, 80%);
    --color3: hsl(120, 95%, 65%);
    --color4: hsl(120, 95%, 40%);
    --color5: hsl(120, 95%, 25%);
    --solid-border: 0.1rem solid black;
    --dashed-border: 0.09rem dashed black;
}

* {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

body {
    background-color: hsl(150, 30%, 75%);
    width: 100vw;
    height: 100vh;
}

main {
    width: 100%;
    height: 100%;
    display: flex;
    flex-direction: column;
    align-items: center;
    justify-content: space-evenly;
}

table {
    border-collapse: collapse;
    table-layout: fixed;
    width: 32rem;
}

tr {
    height: 1.1rem;
}

tr :first-child {
    width: 4rem;
}

td {
    border-right: var(--solid-border);
    border-left: var(--solid-border);
}

.sharp td {
    border-bottom: var(--solid-border);
}

.sharp th:not([class*="time"]) {
    border-bottom: var(--solid-border);
}

.half td {
    border-bottom: var(--dashed-border);
}

tbody :last-child td {
    border: 0;
}

.time {
    text-align: right;
    border: 0;
    padding: 0.4rem;
}

.available-0 {
    background-color: var(--color0);
}

.available-1 {
    background-color: var(--color1);
}

.available-2 {
    background-color: var(--color2);
}

.available-3 {
    background-color: var(--color3);
}

.available-4 {
    background-color: var(--color4);
}

.available-5 {
    background-color: var(--color5);
}

#legend {
    display: flex;
    flex-direction: column;
    align-items: center;
    width: 18rem;
    height: 3.5rem;
    text-align: center;
}

#legend span {
    display: inline-block;
    width: 5rem;
    height: 2rem;
}

#legend-gradient {
    width: 100%;
    height: 60%;
    border: var(--solid-border);
    background-image: linear-gradient(90deg,
            var(--color0) 0% 17%,
            var(--color1) 17% 34%,
            var(--color2) 34% 50%,
            var(--color3) 50% 67%,
            var(--color4) 67% 83%,
            var(--color5) 83% 100%);
}

#legend-line {
    width: 60%;
    height: 100%;
    display: flex;
}
```

