---
id: 66a37f37ef5823a313de8c26
title: 构建月球轨道
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

在本实验中，你将使用 HTML 和 CSS 创建一个月球绕地球公转的简单动画。地球位于系统的中心，月球将围绕地球旋转。

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应通过将 `body` 元素的高度设置为视口高度的 `100%`，并设置为 flexbox 布局，将所有元素居中对齐。

2. 你应有一个 `class` 为 `space` 的 `div` 元素。

3. 在 `.space` 元素内部，应依次包含两个 `div` 元素，分别具有 `earth` 和 `orbit` 类。

4. 在 `.orbit` 元素内部，应有一个 `class` 为 `moon` 的 `div` 元素。

5. `class` 为 `space` 的 `div` 元素应居中显示，宽高均为 `200px`。

6. `class` 为 `space` 的 `div` 元素应使用 `relative` 定位。

7. `.earth` 元素应使用 `absolute` 定位。将 `.earth` 元素的中心定位在父元素的垂直（top）和水平（left）轴的中点。之后，通过将 `.earth` 元素向上和向左各移动自身宽高的一半，使其在父元素 `.space` 内居中。

8. `.earth` 元素的宽高应为 `100px`。

9. `.orbit` 元素的宽高应为 `200px`。

10. `.orbit` 元素应使用 `absolute` 定位。其右下角应通过 `transform` 属性向上和向左各移动 `-50%`，以位于 `.space` 元素的中心。

11. 月球的轨道路径应为一个圆形。

12. `.moon` 元素应使用 `absolute` 定位，宽高为 `30px`。`.moon` 元素应位于 `.orbit` 元素的顶部，并水平居中。

13. 你应通过将 `.moon` 元素向左移动自身宽度的一半，进一步调整其水平位置。

14. `.earth` 和 `.moon` 元素应有背景色，并设置 `border-radius` 为 `50%`，使其看起来像行星。

15. 你应定义一个名为 `orbit` 的 `@keyframes` 动画，使 `.orbit` 元素围绕其中心旋转 360 度。应将该动画应用于 `.orbit` 元素，持续时间为 `5` 秒，线性（linear）计时函数，无限循环。

**注意：** 请确保在 HTML 中正确链接样式表并应用 CSS。

# --hints--

`body` 元素的内容应通过将高度设置为视口高度的 `100%`，`display` 设置为 `flex`，`justify-content` 和 `align-items` 设置为 `center`，实现页面居中。

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

你应有一个主 `div`，其 `class` 为 `space`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

`.space` 元素的宽高应为 `200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

`.space` 元素应设置 `position` 属性为 `relative`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

你应在 `class` 为 `space` 的 `div` 内有一个 `class` 为 `earth` 的 `div`。

```js
assert.exists(document.querySelector('div.space div.earth'));
```

`.earth` 元素应为 `.space` 元素的第一个子元素。

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

`.earth` 元素的宽高应为 `100px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

`.earth` 元素应使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

`.earth` 元素的 `top` 和 `left` 应为 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

`.earth` 元素应通过设置 `transform` 属性为 `translate(-50%, -50%)`，向上和向左各移动 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

你应有一个 `class` 为 `orbit` 的 `div`，位于 `.space` 元素内。

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

`orbit` 类的 `div` 应紧跟在 `earth` 类的 `div` 之后。

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

`.orbit` 元素的宽高应为 `200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

`.orbit` 元素应使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

`.orbit` 元素应通过 `transform` 属性向上和向左各移动 `-50%`，以居中于 `.space` 元素。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

你应在 `.orbit` 元素内有一个 `class` 为 `moon` 的 `div`。

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

`.moon` 元素的宽高应为 `30px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

`.moon` 元素应使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

`.moon` 元素的上边缘应与其包含块的上边缘对齐。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

你应将 `.moon` 元素的左边缘移动到包含块的中心（距离左侧 `50%`）。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

`.moon` 元素应通过 `translateX(-50%)` 在 x 轴上向左移动 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.transform, 'translateX(-50%)');
```

你的 `.earth` 元素应有背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

你的 `.earth` 元素应设置 `border-radius` 为 `50%`，使其呈圆形。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

你的 `.moon` 元素应有背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

你的 `.moon` 元素应设置 `border-radius` 为 `50%`，使其呈圆形。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

你应有一个 `@keyframes` 规则。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

你的新 `@keyframes` 规则应命名为 `orbit`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

你的 `@keyframes orbit` 规则应有 `0%` 选择器。

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

你的 `@keyframes orbit` 规则应有 `100%` 选择器。

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

你的 `0%` 选择器应设置 `transform` 属性为 `rotate(0deg) translate(-50%, -50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

你的 `100%` 选择器应在 `0%` 选择器之后。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

你的 `100%` 选择器应设置 `transform` 属性为 `rotate(360deg) translate(-50%, -50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

你的 `.orbit` 选择器应设置 `animation` 属性为 `orbit 5s linear infinite`。

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

