---
id: 66d6a7a3e1aa411e94bf2346
title: 搭建一个房屋绘画
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

在本实验中，你将使用 HTML 创建一个房屋的结构。然后，你将使用 CSS 定位来排列房屋的各个元素，如窗户和门。

完成以下用户故事并通过所有测试即可完成本实验。

**用户需求：**

1. 你需要有一个 `id` 为 `house` 的 `div` 元素。
1. 你的 `#house` 应设置 `position` 为 `relative`，以便其子元素可以相对于它进行定位。
1. 你的 `#house` 应设置宽度为 `500px`，高度为 `400px`。
1. 你的 `#house` 应设置背景色和边框。
1. 你需要在 `#house` 内部包含五个 `div` 元素，分别具有以下 `id`：`chimney`、`roof`、`window-1`、`window-2` 和 `door`。
1. `#house` 的所有直接子元素都应设置 `position` 为 `absolute`。
1. `#roof`、`#chimney`、`#window-1`、`#window-2` 和 `#door` 都应设置宽度、高度、边框和背景色。
1. 你的 `#roof` 应设置 `top` 值为 `0`。
1. 你的 `#door` 应放置在房屋的底部。
1. 你的窗户应放置在 `#roof` 下方，并且至少高于 `#door` 高度的三分之一。
1. 你的窗户和门都应设置 `left` 或 `right`，使它们位于房屋边界内。
1. 你的 `#chimney` 应设置 `top` 值，使其位于 `#house` 的顶部。
1. 你的 `#chimney` 应设置 `z-index`，使其位于房屋后面。

**注意：** 请确保在 HTML 中正确链接你的样式表，以应用 CSS。

# --hints--

你需要有一个 `id` 为 `house` 的 `div` 元素。

```js
assert.exists(document.querySelector("div#house"));
```

你需要选择 `#house` 并将其 `position` 设置为 `relative`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

你需要选择 `#house` 并设置其 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

你需要选择 `#house` 并设置其 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

你需要选择 `#house` 并将其宽度设置为 `500px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

你需要选择 `#house` 并将其高度设置为 `400px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

你需要在 `#house` 内部有一个 `id` 为 `chimney` 的 `div`。

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

`#chimney` 应设置 `position` 属性为 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

你需要选择 `#chimney` 并设置其 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

你需要选择 `#chimney` 并设置其 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

你需要选择 `#chimney` 并设置其 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

你需要选择 `#chimney` 并设置其 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

你需要在 `#house` 内部有一个 `id` 为 `roof` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

`#roof` 应设置 `position` 属性为 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

你需要选择 `#roof` 并设置其 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

你需要选择 `#roof` 并设置其 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

你需要选择 `#roof` 并设置其 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

你需要选择 `#roof` 并设置其 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

你需要在 `#house` 内部有一个 `id` 为 `window-1` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1` 应设置 `position` 属性为 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

你需要选择 `#window-1` 并设置其 `width` 属性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

你需要选择 `#window-1` 并设置其 `height` 属性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

你需要选择 `#window-1` 并设置其 `border` 属性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

你需要选择 `#window-1` 并设置其 `background-color` 属性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

你需要在 `#house` 内部有一个 `id` 为 `window-2` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2` 应设置 `position` 属性为 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

你需要选择 `#window-2` 并设置其 `width` 属性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

你需要选择 `#window-2` 并设置其 `height` 属性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

你需要选择 `#window-2` 并设置其 `border` 属性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

你需要选择 `#window-2` 并设置其 `background-color` 属性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

你需要在 `#house` 内部有一个 `id` 为 `door` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door` 应设置 `position` 属性为 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

你需要选择 `#door` 并设置其 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

你需要选择 `#door` 并设置其 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

你需要选择 `#door` 并设置其 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

你需要选择 `#door` 并设置其 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

你的 `#roof` 应设置 `top` 属性为 `0`。

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

你的 `#door` 应放置在房屋底部。

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

你应为 `#door` 设置 `left` 或 `right`，并将其放置在房屋边界内。

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isAbove(doorLeft, 0);
assert.isAbove(doorRight, 0);
```

你的 `#chimney` 应设置 `top` 值，使其位于 `#house` 顶部。

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

你的 `#chimney` 应设置 `z-index`，使其位于房屋后面。

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
  assert.isBelow(Number(chimneyZIndex), 0)
} else {
  assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

你的窗户应放置在 `#roof` 下方，并且至少高于 `#door` 高度的三分之一。

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

你应为窗户设置 `left` 或 `right`，并将其放置在房屋边界内。

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isAbove(w1Left, 0);
assert.isAbove(w1Right, 0);
assert.isAbove(w2Left, 0);
assert.isAbove(w2Right, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>房屋绘画</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh-CN">
<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>房屋绘画</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <div id="house">
    <div id="chimney">
      <div id="smoke"></div>
    </div>
    <div id="roof"></div>
    <div id="window-1"></div>
    <div id="window-2"></div>
    <div id="door">
      <div id="door-knob"></div>
    </div>
    <div id="welcome">欢迎光临</div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
  }

  body {
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background: linear-gradient(
    #b3e6ff,
    #b3e6ff 60%,
    green 60%
  );
  }

  #house {
  flex: 0 0 auto;
  width: 500px;
  height: 400px;
  position: relative;
  background-color: #ff9980;
  border: 6px solid #b35900;
  }

  #chimney {
  position: absolute;
  width: 90px;
  height: 100px;
  top: -100px;
  left: 65%;
  background: repeating-linear-gradient(
    #e6e6e6,
    #e6e6e6 10%,
    black 10%,
    black 12%
  );
  border: 2px solid black;
  z-index: -1;
  }

  #smoke {
  width: 40px;
  height: 40px;
  background-color: rgba(0,0,0,0.3);
  position: absolute;
  top: -20px;
  left: 20px;
  border-radius: 50%;
  box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
  border: none;
  }

  #roof {
  border: 2px;
  width: 100%;
  height: 110px;
  background-color: green;
  position: absolute;
  top: 0;
  background: repeating-linear-gradient(
    45deg,
    #b35900,
    #b35900 2%,
    transparent 2%,
    transparent 5%
  ), repeating-linear-gradient(
    -45deg,
    #b35900,
    #b35900 2%,
    #ff9980 2%,
    #ff9980 5%
  )
  }

  #window-1, #window-2 {
  width: 100px;
  height: 100px;
  background: linear-gradient(
    to right,
    #ffffb3,
    #ffffb3 48%,
    #b35900 48%,
    #b35900 52%,
    #ffffb3 52%,
    #ffffb3 100%
  );
  position: absolute;
  top: 40%;
  border: 6px solid #b35900;
  }

  #window-1 {
  left: 7.5%;
  }

  #window-2 {
  right: 7.5%;
  }

  #door {
  width: 130px;
  height: 180px;
  position: absolute;
  bottom: 0%;
  left: 185px;
  background-color: #e6e6e6;
  border: 6px solid #b35900;
  transform: translateY(6px);
  }

  #door-knob {
  width: 20px;
  height: 20px;
  background-color: #b35900;
  border-radius: 50%;
  position: absolute;
  top: 50%;
  right: 5%;
  }

  #welcome {
  font-family: sans-serif;
  font-weight: bold;
  width: 130px;
  height: 40px;
  position: absolute;
  bottom: -47px;
  left: 178px;
  background-color: #85e085;
  border: 2px solid black;
  transform: skewX(-20deg);
  display: flex;
  align-items: center;
  justify-content: center;
  }
```

