---
id: 66fe4f33a2cc9b33f4d5cd9b
title: 构建一个排列生成器
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

在本实验中，你将构建一个排列生成器，它接收一个字符串并返回该字符串所有可能的字符排列。例如，字符串 `cat` 的所有可能排列为：`cat`、`cta`、`act`、`atc`、`tac` 和 `tca`。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你需要创建一个名为 `permuteString` 的函数。
2. `permuteString` 函数应接收三个参数——一个字符串、一个前缀值和一个用于存储及返回结果的空数组。前缀值用于累加字符以形成一个排列。
3. 在函数内部，应检查传入字符串的长度是否为 `0`。如果是，则将当前前缀推入结果数组并返回结果。
4. 遍历输入字符串的每个字符，在每次迭代中移除当前字符，并用更新后的参数递归调用 `permuteString` 函数，以生成剩余字符的排列。
5. 你应返回最终的结果数组。
6. 你应确保排列结果唯一，去除重复项。

# --hints--

你应该有一个名为 `permuteString` 的函数。

```js
assert.isFunction(permuteString);
```

你应该在 `permuteString` 函数中使用递归。

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` 应返回 `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`。

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` 应返回 `[ "fcc", "cfc", "ccf" ]`。

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` 应返回 `[ "p" ]`。

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` 应返回 `[""]`。

```js
let emptyArr= permuteString("");

// 1，因为空字符串本身作为一个排列被推入
assert.lengthOf(emptyArr, 1);

```

`permuteString("101")` 应返回 `[ "101", "011", "110" ]`。

```js
assert.sameMembers(permuteString("101"), [ "101", "011", "110"]);
```

`permuteString` 应返回正确的结果。

```js
// 防止硬编码的测试
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // 首先对字符串排序以分组重复字符
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // 跳过重复字符
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```

