---
id: 66d9af3897e7d75a895b72c2
title: 构建第 n 个斐波那契数生成器
challengeType: 14
dashedName: lab-nth-fibonacci-number-generator
---

# --description--

在本实验中，你将使用动态规划的方法来计算并返回斐波那契数列中的第 `n` 个数字，其中每个数字都是其前两个数字之和。

请完成以下用户需求，并通过所有测试以完成本实验。

**用户需求：**

1. 你应在全局作用域中定义一个名为 `sequence` 的数组。它应包含数字 `0` 和 `1`。
1. 你应创建一个名为 `fibonacci` 的函数。
1. `fibonacci` 函数应接收一个参数，该参数为正整数。
1. 如果输入不是正整数，函数应返回 `"Invalid input"`。
1. 调用 `fibonacci(n)` 时，应使用动态规划方法计算并返回斐波那契数列中的第 `n` 个数字。
1. 每次计算出的第 `n` 个斐波那契数应存储在 `sequence` 数组的索引 `n - 1` 处。
1. 如果输入的数字请求的是已计算过的数字，则不应进行新的计算，应直接从 `sequence` 数组中获取该数字。

# --hints--

你应有一个 `sequence` 数组。

```js
assert.isArray(sequence);
```

`sequence` 数组在定义时应包含 `0` 和 `1`。

```js
assert.match(__helpers.removeJSComments(code), /(let|const)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]\s*;?/)
```

你应创建一个 `fibonacci` 函数。

```js
assert.isFunction(fibonacci);
```

`fibonacci(3.5)` 应返回 `"Invalid input"`。

```js
assert.strictEqual(fibonacci(3.5), "Invalid input")
```

`fibonacci(0)` 应返回 `"Invalid input"`。

```js
assert.strictEqual(fibonacci(0), "Invalid input")
```

`fibonacci(-1)` 应返回 `"Invalid input"`。

```js
assert.strictEqual(fibonacci(-1), "Invalid input")
```

`fibonacci("three")` 应返回 `"Invalid input"`。

```js
assert.strictEqual(fibonacci("three"), "Invalid input")
```

`fibonacci(1)` 应返回 `0`。

```js
assert.strictEqual(fibonacci(1), 0)
```

`fibonacci(32)` 应返回 `1346269`。

```js
assert.strictEqual(fibonacci(32), 1346269)
```

`fibonacci(5)` 应返回 `3`。

```js
assert.strictEqual(fibonacci(5), 3);
```

`sequence` 应包含到目前为止调用的最大数的所有数字。

```js
sequence.length = 2
fibonacci(2)
assert.lengthOf(sequence, 2)
fibonacci(5)
assert.lengthOf(sequence, 5)
fibonacci(3)
assert.lengthOf(sequence, 5)
fibonacci(10)
assert.lengthOf(sequence, 10)
assert.deepEqual(sequence, [0, 1, 1, 2, 3, 5, 8, 13, 21, 34])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const sequence = [0, 1];

const fibonacci = (n) => {
    if (typeof n !== "number" || n < 1 || n % 1 !== 0) {
        console.log("Invalid input")
        return "Invalid input"
    }
    if (n <= sequence.length) {
        return sequence[n - 1];
    }
    for (let i = sequence.length; i < n; i++) {
        sequence.push(sequence[i - 1] + sequence[i - 2]);
    }
    return sequence[n - 1]
}
```

