---
id: 66d97c3be4030e8f41a92083
title: 实现快速排序算法
challengeType: 14
dashedName: build-the-quicksort-algorithm
---

# --description--

快速排序是一种使用分治法对数组进行排序的算法。

其实现步骤如下：

- 在输入数组中选择一个基准值（可以使用数组的第一个或最后一个元素）。
- 将输入数组划分为三个子数组，分别包含大于、小于和等于基准值的元素。
- 递归调用 `quicksort` 对子数组进行排序。

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该定义一个名为 `quicksort` 的函数来实现快速排序算法。

2. `quicksort` 函数应接收一个整数数组作为输入，并返回一个按从小到大排序后的新数组。

# --hints--

`quicksort` 应该是一个函数。

```js
assert.isFunction(quicksort);
```

`quicksort` 应该返回一个已排序（从小到大）的数组。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert(
  isSorted(
    quicksort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quicksort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` 应返回一个仅顺序改变的数组。

```js
assert.sameMembers(
  quicksort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` 不应使用内置的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  quicksort([0, 1]);
  Array.prototype.sort = temp;
  return !sortUsed;
}

assert(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const quicksort = (array) => {
    if (array.length === 0) {
        return [];
    }
    const pivot = array[0];
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let i of array) {
        if (i < pivot) {
            lesser.push(i);
        } else if (i > pivot) {
            greater.push(i);
        } else {
            equal.push(i);
        }
    }
    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
}
```

