---
id: 587d825d367417b2b2512c96
title: 实现深度优先搜索算法
challengeType: 14
dashedName: implement-the-depth-first-search-algorithm
---

# --description--

<dfn>深度优先搜索</dfn>（DFS）是一种用于遍历树或图等数据结构的算法。从根节点开始，搜索会沿着一条路径尽可能深入。当到达路径的尽头时，搜索会回溯到上一个有未访问边的节点，并继续沿着未访问的边进行搜索。

**目标：** 完成以下用户需求，并通过所有测试以完成实验。

**用户需求：**

1. 你需要编写一个实现<dfn>深度优先搜索</dfn>算法的函数 `dfs()`。
1. 该函数应能记录已访问的节点。
1. 该函数应使用栈（即后进先出数组）来确保总是优先访问最近添加节点的邻居。
1. 该函数应接收一个无向图的邻接矩阵 `graph` 和一个节点标签 `root` 作为参数。节点标签是介于 `0` 和 `n - 1` 之间的数字，其中 `n` 是图中节点的总数。
1. 该函数应输出一个包含所有从 `root` 可达节点的数组。

# --hints--

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`，起始节点为 `1` 时，应返回包含 `0`、`1`、`2` 和 `3` 的数组。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`，起始节点为 `3` 时，应返回包含 `3`、`2`、`1` 和 `0` 的数组。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`，起始节点为 `1` 时，应返回一个包含四个元素的数组。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`，起始节点为 `3` 时，应返回包含 `3` 的数组。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`，起始节点为 `3` 时，应返回一个包含一个元素的数组。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`，起始节点为 `3` 时，应返回包含 `2` 和 `3` 的数组。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`，起始节点为 `3` 时，应返回一个包含两个元素的数组。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`，起始节点为 `0` 时，应返回包含 `0` 和 `1` 的数组。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

输入图 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`，起始节点为 `0` 时，应返回一个包含两个元素的数组。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
    const stack = [];
    let tempV;
    const visited = [];
    let tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (let i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}

const exDFSGraph = [
    [0, 1, 0, 0],
    [1, 0, 1, 0],
    [0, 1, 0, 1],
    [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 0));
```

