---
id: 673b02b03134b04637bf7055
title: 构建可复用的页脚
challengeType: 25
dashedName: build-a-reusable-footer
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你需要导出一个 `Footer` 组件。
2. 你的 `Footer` 组件应返回一个包含所有其他元素的 `footer` 元素。
3. 你的 `footer` 元素应包含：

- 至少三个无序列表，每个列表至少有两个列表项。
- 至少一个包含版权符号 `©` 的段落元素。
- 至少三个 `href` 值为 `#` 且内容为你选择的图标或文本的链接。

# --hints--

你需要导出一个 `Footer` 组件。

```js
assert.isFunction(window.index.Footer);
```

你的 `Footer` 组件应返回一个 `footer` 元素。

```js
const footer = document.getElementById('root').firstElementChild;
assert.equal(footer.tagName, 'FOOTER');
```

你的 `Footer` 组件只能包含一个 `footer` 元素。

```js
assert.lengthOf(document.querySelectorAll('footer'), 1);
```

你的 `footer` 元素不能有任何兄弟元素。

```js
assert.lengthOf(document.getElementById('root').children, 1);
```

你的 `footer` 元素应包含至少三个无序列表。

```js
assert.isAtLeast(document.querySelectorAll('ul').length, 3);
```

每个无序列表应至少有两个列表项。

```js
const uls = [...document.querySelectorAll('ul')];
uls.forEach(ul => {
  assert.isAtLeast(ul.querySelectorAll('li').length, 2);
  }
);
```

你的 `footer` 应至少有一个段落元素。

```js
assert.isAtLeast(document.querySelectorAll('p').length, 1);
```

你应在某个段落中包含版权符号（`©`）。

```js
const texts = [...document.querySelectorAll('p')].map(p => p.textContent);
const hasCopyright = texts.some(text => text.includes('©'));
assert.isTrue(hasCopyright);
```

你的 `footer` 应至少有三个 `href` 值为 `#` 的链接。

```js
const links = document.querySelectorAll('a[href="#"]');
assert.isAtLeast(links.length, 3);
```

你的链接内容不能为空。

```js
const links = [...document.querySelectorAll('a[href="#"]')];
links.forEach(link => {
  assert.isNotEmpty(link.textContent);
  }
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
  <meta charset="UTF-8" />
  <title>可复用页脚组件</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { Footer } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
  </script>
  </body>
</html>
```

```css

```

```jsx
export const Footer = () => {};
```

# --solutions--

```html
<!doctype html>
<html>
  <head>
  <meta charset="UTF-8" />
  <title>可复用页脚组件</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { Footer } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
  </script>
  </body>
</html>
```

```css
.footer {
  background-color: #e3e1eb;
  padding: 20px;
  text-align: center;
  border-top: 1px solid #e5e7eb;
  font-family: Arial, sans-serif;
}

.footer-section {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  justify-content: center;
  margin-bottom: 20px;
  align-items: center;
}

.footer ul {
  list-style: none;
  padding: 0;
  margin: 0;
}

.footer li {
  margin-bottom: 10px;
}

.footer a {
  text-decoration: none;
  color: #1f2937;
}

.footer a:hover {
  color: #6366f1;
}

.footer p {
  color: #6b7280;
  margin-top: 20px;
}

.footer-icons {
  margin-top: 20px;
  display: flex;
  justify-content: center;
}

.footer-icons a {
  margin: 0 10px;
  font-size: 24px;
  color: #6b7280;
  text-decoration: none;
}

.footer-icons a:hover {
  color: #6366f1;
}
```

```jsx
export const Footer = () => (
  <footer className='footer'>
  <div className='footer-section'>
    <ul>
    <li>
      <a href='#'>健身仪表盘</a>
    </li>
    <li>
      <a href='#'>服务</a>
    </li>
    </ul>
    <ul>
    <li>
      <a href='#'>观看视频</a>
    </li>
    <li>
      <a href='#'>Discord</a>
    </li>
    </ul>
    <ul>
    <li>
      <a href='#'>隐私政策</a>
    </li>
    <li>
      <a href='#'>条款与条件</a>
    </li>
    </ul>
  </div>
  <p>© 2024 健身仪表盘。保留所有权利。</p>
  <div className='footer-icons'>
    <a href='#'>&#x1F3AE;</a>
    <a href='#'>&#x1F426;</a>
    <a href='#'>&#x1F4BB;</a>
    <a href='#'>&#x1F3C0;</a>
  </div>
  </footer>
);
```

