---
id: 67279fe50237291f80eed8b8
title: 构建客户投诉表单
challengeType: 25
dashedName: build-a-customer-complaint-form
demoType: onClick
---

# --description--

在本实验中，已为你提供所有 HTML 和 CSS。你需要使用 JavaScript 对投诉表单进行校验。

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 当表单提交时，你需要确保：
    - `#full-name` 不为空。
    - `#email` 格式为有效的电子邮件地址。
    - `#order-no` 是以 `2024` 开头的十位数字序列。
    - `#product-code` 符合 `XX##-X###-XX#` 的模式，其中 `X` 表示大小写字母，`#` 表示数字。
    - `#quantity` 是正整数。
    - `#complaints-group` 至少有一个复选框被选中。
    - 如果勾选了 `Other` 复选框，`#complaint-description` 至少包含二十个字符。
    - `#solutions-group` 至少有一个单选按钮被选中。
    - 如果选择了 `Other` 单选按钮，`#solution-description` 至少包含二十个字符。
1. 你需要有一个名为 `validateForm` 的函数，返回一个包含以下键的对象：`full-name`、`email`、`order-no`、`product-code`、`quantity`、`complaints-group`、`complaint-description`、`solutions-group` 和 `solution-description`。每个键的值为 `true` 表示对应字段填写正确，否则为 `false`。
1. 你需要有一个名为 `isValid` 的函数，接收 `validateForm` 返回的对象作为参数，如果所有字段都填写正确则返回 `true`，否则返回 `false`。
1. 当表单字段触发 change 事件且值有效时，应将其边框颜色设置为 `green`。对于复选框和单选按钮组，应设置父级 `fieldset` 的边框颜色。
1. 当表单字段触发 change 事件且值无效时，应将其边框颜色设置为 `red`。对于复选框和单选按钮组，应设置父级 `fieldset` 的边框颜色。
1. 当尝试提交表单时，应调用 `isValid` 对表单进行校验。
1. 当尝试提交表单且存在无效字段时，应将每个无效字段的输入框、文本域或 fieldset（针对复选框和单选按钮组）的边框颜色设置为 `red` 以高亮显示。

# --hints--

你需要有一个名为 `validateForm` 的函数。

```js
assert.isFunction(validateForm);
```

`validateForm` 应返回一个对象。

```js
assert.isObject(validateForm());
```

当 `#full-name` 为空时，`validateForm()["full-name"]` 应为 `false`，否则为 `true`。

```js
const field = document.getElementById("full-name");
field.value = "";
assert.isFalse(validateForm()["full-name"]);
field.value = "testing";
assert.isTrue(validateForm()["full-name"]);
```

当 `#full-name` 触发 `change` 事件且值有效时，应将其边框颜色设置为 `green`，否则为 `red`。

```js
const field = document.getElementById("full-name");
field.value = "testing"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
field.value = ""
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#email` 格式为有效邮箱时，`validateForm()["email"]` 应为 `true`，否则为 `false`。

```js
const field = document.getElementById("email");
field.value = "example@domain.com"
assert.isTrue(validateForm()["email"]);
field.value = "testing"
assert.isFalse(validateForm()["email"]);
```

当 `#email` 触发 `change` 事件且值有效时，应将其边框颜色设置为 `green`，否则为 `red`。

```js
const field = document.getElementById("email");
field.value = "example@domain.com"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
field.value = ""
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#order-no` 格式有效时，`validateForm()["order-no"]` 应为 `true`，否则为 `false`。

```js
const field = document.getElementById("order-no");
field.value = "2024001122"
assert.isTrue(validateForm()["order-no"]);
field.value = "testing"
assert.isFalse(validateForm()["order-no"]);
```

当 `#order-no` 触发 `change` 事件且值有效时，应将其边框颜色设置为 `green`，否则为 `red`。

```js
const field = document.getElementById("order-no");
field.value = "2024001122"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
field.value = "20240011"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#product-code` 格式有效时，`validateForm()["product-code"]` 应为 `true`，否则为 `false`。

```js
const field = document.getElementById("product-code");
field.value = "Xa22-X123-Xb4";
assert.isTrue(validateForm()["product-code"]);
field.value = "testing"
assert.isFalse(validateForm()["product-code"]);
```

当 `#product-code` 触发 `change` 事件且值有效时，应将其边框颜色设置为 `green`，否则为 `red`。

```js
const field = document.getElementById("product-code");
field.value = "Xa22-X123-Xb4"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
field.value = "Xa22-123-Xb4"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#quantity` 有效时，`validateForm()["quantity"]` 应为 `true`，否则为 `false`。

```js
const field = document.getElementById("quantity");
field.value = "5";
assert.isTrue(validateForm()["quantity"]);
field.value = "0";
assert.isFalse(validateForm()["quantity"]);
```

当 `#quantity` 触发 `change` 事件且值有效时，应将其边框颜色设置为 `green`，否则为 `red`。

```js
const field = document.getElementById("quantity");
field.value = "2"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
field.value = "0"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#complaints-group` 至少有一个复选框被选中时，`validateForm()["complaints-group"]` 应为 `true`。

```js
document.getElementById("damaged-product").checked = true
assert.isTrue(validateForm()["complaints-group"]);
```

当 `#complaints-group` 所有复选框都未选中时，`validateForm()["complaints-group"]` 应为 `false`。

```js
document.getElementById("damaged-product").checked = false;
document.getElementById("nonconforming-product").checked = false;
document.getElementById("delayed-dispatch").checked = false;
document.getElementById("other-complaint").checked = false;
assert.isFalse(validateForm()["complaints-group"]);
```

当 `#complaints-group` 至少有一个复选框被选中时，应将 `#complaints-group` 的边框颜色设置为 `green`。

```js
document.getElementById("nonconforming-product").checked = false;
document.getElementById("delayed-dispatch").checked = false;
document.getElementById("other-complaint").checked = false;
const damagedProduct = document.getElementById("damaged-product");
damagedProduct.checked = true;
damagedProduct.dispatchEvent(new Event("change", { bubbles: true }));
const fieldset = document.getElementById("complaints-group");
assert.equal(fieldset.style.borderColor, "green");
```

当 `#complaints-group` 所有复选框都未选中时，应将 `#complaints-group` 的边框颜色设置为 `red`。

```js
document.getElementById("damaged-product").checked = false;
document.getElementById("delayed-dispatch").checked = false;
document.getElementById("other-complaint").checked = false;
const nonConfProduct = document.getElementById("nonconforming-product");
nonConfProduct.checked = false;
nonConfProduct.dispatchEvent(new Event("change", { bubbles: true }));
const fieldset = document.getElementById("complaints-group");
assert.equal(fieldset.style.borderColor, "red");
```

当勾选了 `#other-complaint` 且 `#complaint-description` 至少包含二十个字符时，`validateForm()["complaint-description"]` 应为 `true`。

```js
document.getElementById("other-complaint").checked = true
document.getElementById("complaint-description").value = "A sentence with at least twenty characters";
assert.isTrue(validateForm()["complaint-description"]);
```

当勾选了 `#other-complaint` 且 `#complaint-description` 少于二十个字符时，`validateForm()["complaint-description"]` 应为 `false`。

```js
document.getElementById("other-complaint").checked = true
document.getElementById("complaint-description").value = "Less than 20 chars";
assert.isFalse(validateForm()["complaint-description"]);
```

当 `#complaint-description` 的值变为有效时，应将其边框颜色设置为 `green`。

```js
const field = document.getElementById("complaint-description");
field.value = "A sentence with at least twenty characters"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
```

当 `#complaint-description` 的值变为无效时，应将其边框颜色设置为 `red`。

```js
const field = document.getElementById("complaint-description");
field.value = "Not enough chars"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

当 `#solutions-group` 至少有一个单选按钮被选中时，`validateForm()["solutions-group"]` 应为 `true`。

```js
document.getElementById("refund").checked = true
assert.isTrue(validateForm()["solutions-group"]);
```

当 `#solutions-group` 所有单选按钮都未选中时，`validateForm()["solutions-group"]` 应为 `false`。

```js
document.getElementById("refund").checked = false;
document.getElementById("exchange").checked = false;
document.getElementById("other-solution").checked = false;
assert.isFalse(validateForm()["solutions-group"]);
```

当 `#solutions-group` 至少有一个单选按钮被选中时，应将 `#solutions-group` 的边框颜色设置为 `green`。

```js
document.getElementById("exchange").checked = false;
document.getElementById("other-solution").checked = false;
const refundRadioBtn = document.getElementById("refund");
refundRadioBtn.checked = true;
refundRadioBtn.dispatchEvent(new Event("change", { bubbles: true }));
const fieldset = document.getElementById("solutions-group");
assert.equal(fieldset.style.borderColor, "green");
```

当 `#solutions-group` 所有单选按钮都未选中时，应将 `#solutions-group` 的边框颜色设置为 `red`。

```js
document.getElementById("refund").checked = false;
document.getElementById("other-solution").checked = false;
const exchangeRadioBtn = document.getElementById("exchange");
exchangeRadioBtn.checked = false;
exchangeRadioBtn.dispatchEvent(new Event("change", { bubbles: true }));
const fieldset = document.getElementById("solutions-group");
assert.equal(fieldset.style.borderColor, "red");
```

当勾选了 `#other-solution` 且 `#solution-description` 至少包含二十个字符时，`validateForm()["solution-description"]` 应为 `true`。

```js
document.getElementById("other-solution").checked = true
document.getElementById("solution-description").value = "A sentence with at least twenty characters";
assert.isTrue(validateForm()["solution-description"]);
```

当勾选了 `#other-solution` 且 `#solution-description` 少于二十个字符时，`validateForm()["solution-description"]` 应为 `false`。

```js
document.getElementById("other-solution").checked = true
document.getElementById("solution-description").value = "Less than 20 chars";
assert.isFalse(validateForm()["solution-description"]);
```

当 `#solution-description` 的值变为有效时，应将其边框颜色设置为 `green`。

```js
const field = document.getElementById("solution-description");
field.value = "A sentence with at least twenty characters"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "green");
```

当 `#solution-description` 的值变为无效时，应将其边框颜色设置为 `red`。

```js
const field = document.getElementById("solution-description");
field.value = "Not enough chars"
field.dispatchEvent(new Event("change", { bubbles: true }));
assert.equal(field.style.borderColor, "red");
```

你需要有一个名为 `isValid` 的函数。

```js
assert.isFunction(isValid);
```

你的 `isValid` 函数应接收 `validateForm()` 作为参数，当所有表单字段都填写正确时返回 `true`。

```js
document.getElementById("full-name").value = "testing"
document.getElementById("email").value = "example@domain.com"
document.getElementById("order-no").value = "2024001122"
document.getElementById("product-code").value = "Xa22-X123-Xb4"
document.getElementById("quantity").value = "5"
document.getElementById("damaged-product").checked = true
document.getElementById("other-complaint").checked = false
document.getElementById("refund").checked = true
document.getElementById("other-solution").checked = false
assert.isTrue(isValid(validateForm()));

document.getElementById("other-complaint").checked = true
document.getElementById("complaint-description").value = "A sentence with at least twenty characters";
document.getElementById("other-solution").checked = true
document.getElementById("solution-description").value = "A sentence with at least twenty characters";
assert.isTrue(isValid(validateForm()));
```

你的 `isValid` 函数应接收 `validateForm()` 作为参数，当有任一表单字段未正确填写时返回 `false`。

```js
const name = document.getElementById("full-name");
const email = document.getElementById("email");
const orderNo = document.getElementById("order-no");
const prCode = document.getElementById("product-code");
const quantity = document.getElementById("quantity");
const damaged = document.getElementById("damaged-product");
const nonConforming = document.getElementById("nonconforming-product");
const delayed = document.getElementById("delayed-dispatch");
const otherC = document.getElementById("other-complaint");
const complaintDescr = document.getElementById("complaint-description");
const refund = document.getElementById("refund");
const exchange = document.getElementById("exchange");
const otherS = document.getElementById("other-solution");
const solutionDescr = document.getElementById("solution-description");

// full-name 无效
name.value = "";
email.value = "example@domain.com";
orderNo.value = "2024001122";
prCode.value = "Xa22-X123-Xb4";
quantity.value = "5";
damaged.checked = true;
refund.checked = true;
assert.isFalse(isValid(validateForm()));

// email 无效
name.value = "testing";
email.value = "";
assert.isFalse(isValid(validateForm()));

// order-no 无效
email.value = "example@domain.com";
orderNo.value = "202400";
assert.isFalse(isValid(validateForm()));

// product-code 无效
orderNo.value = "2024001122";
prCode.value = "1a22-Xa23-Xb4";
assert.isFalse(isValid(validateForm()));

// quantity 无效
prCode.value = "Xa22-X123-Xb4";
quantity.value = 0;
assert.isFalse(isValid(validateForm()));

// complaints-group 无效
quantity.value = 5;
damaged.checked = false;
nonConforming.checked = false;
delayed.checked = false;
otherC.checked = false;
assert.isFalse(isValid(validateForm()));

// complaint-description 无效
otherC.checked = true;
complaintDescr.value = "not enough chars";
assert.isFalse(isValid(validateForm()));


// solutions-group 无效
complaintDescr.value = "A sentence with at least twenty characters";
refund.checked = false;
exchange.checked = false;
otherS.checked = false;
assert.isFalse(isValid(validateForm()));

// solutions-description 无效
otherS.checked = true;
solutionDescr.value = "not enough chars";
assert.isFalse(isValid(validateForm()));
```

当尝试提交表单时，应调用 `isValid`。

```js
let flag = false;
const temp = isValid;
try {
    isValid = (obj) => {flag = true};
    document.getElementById("form").dispatchEvent(new Event("submit"));
    assert.isTrue(flag);
} finally {
    isValid = temp;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
    <title>客户投诉表单</title>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>投诉表单</h1>
    <form id="form">
        <fieldset id="personal-info">
            <div>
                <label for="full-name">姓名：</label>
                <input type="text" id="full-name" name="full-name" placeholder="张三">
            </div>

            <div>
                <label for="email">电子邮箱：</label>
                <input type="email" id="email" name="email" placeholder="example@domain.com">
            </div>
        </fieldset>
        <hr>
        <fieldset id="product-info">
            <div>
                <label for="order-no">订单号：</label>
                <input type="text" id="order-no" name="order-no" placeholder="2024######">
            </div>
            <div>
                <label for="product-code">产品编码：</label>
                <input type="text" id="product-code" name="product-code" placeholder="XX##-X###-XX#">
            </div>
            <div>
                <label for="quantity">数量：</label>
                <input type="number" id="quantity" name="quantity" min="1">
            </div>
        </fieldset>

        <fieldset id="complaints-group">
            <legend>投诉原因：</legend>
            <div>
                <input type="checkbox" id="damaged-product" name="complaint" value="damaged-product">
                <label for="damaged-product">产品损坏</label>
            </div>

            <div>
                <input type="checkbox" id="nonconforming-product" name="complaint" value="nonconforming-product">
                <label for="nonconforming-product">产品不符</label>
            </div>

            <div>
                <input type="checkbox" id="delayed-dispatch" name="complaint" value="delayed-dispatch">
                <label for="delayed-dispatch">发货延迟</label>
            </div>

            <div>
                <input type="checkbox" id="other-complaint" name="complaint" value="other">
                <label for="other-complaint">其他</label>
            </div>
        </fieldset>

        <div id="complaint-description-container">
            <legend>投诉原因描述</legend>
            <textarea placeholder="请用不少于20个字符描述您的投诉原因"
                name="complaint-textarea" id="complaint-description"></textarea>
        </div>

        <fieldset id="solutions-group">
            <legend>期望解决方案</legend>
            <input type="radio" name="solutions" id="refund" value="refund">
            <label for="refund">退款</label>

            <input type="radio" name="solutions" id="exchange" value="exchange">
            <label for="exchange">换货</label>

            <input type="radio" name="solutions" id="other-solution" value="other">
            <label for="other-solution">其他</label>
        </fieldset>

        <div id="solution-description-container">
            <legend>期望解决方案描述</legend>
            <textarea placeholder="请用不少于20个字符描述您期望的解决方案"
                name="solution-textarea" id="solution-description"></textarea>
        </div>
        <div id="btn-container">
            <button type="submit" id="submit-btn">提交</button>
            <span id="message-box"></span>
        </div>

    </form>

    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

h1 {
    text-align: center;
}

#form {
    max-width: 70%;
    margin: auto;
    border-radius: 10px;
    box-shadow: rgba(0, 0, 0, 0.35) 0px 5px 15px;
    padding: 10px;
}

input {
    border-color: rgb(118, 118, 118);
}

#personal-info input, #product-info input {
    width: 100%;
    margin-bottom: 10px;
}


fieldset {
    margin-bottom: 10px;
    border-radius: 5px;
    border-color: rgb(118, 118, 118);
}

textarea {
    width: 100%;
    border-color: rgb(118, 118, 118);
}

#btn-container {
    display: flex;
    justify-content: space-between;
    align-items: center;
}

#submit-btn, #clear-btn {
    margin: 10px 15px 0;
}
```

```js

```

## --solutions--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
    <title>客户投诉表单</title>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>投诉表单</h1>
    <form id="form">
        <fieldset id="personal-info">
            <div>
                <label for="full-name">姓名：</label>
                <input type="text" id="full-name" name="full-name" placeholder="张三">
            </div>

            <div>
                <label for="email">电子邮箱：</label>
                <input type="email" id="email" name="email" placeholder="example@domain.com">
            </div>
        </fieldset>
        <hr>
        <fieldset id="product-info">
            <div>
                <label for="order-no">订单号：</label>
                <input type="text" id="order-no" name="order-no" placeholder="2024######">
            </div>
            <div>
                <label for="product-code">产品编码：</label>
                <input type="text" id="product-code" name="product-code" placeholder="XX##-X###-XX#">
            </div>
            <div>
                <label for="quantity">数量：</label>
                <input type="number" id="quantity" name="quantity" min="1">
            </div>
        </fieldset>

        <fieldset id="complaints-group">
            <legend>投诉原因：</legend>
            <div>
                <input type="checkbox" id="damaged-product" name="complaint" value="damaged-product">
                <label for="damaged-product">产品损坏</label>
            </div>

            <div>
                <input type="checkbox" id="nonconforming-product" name="complaint" value="nonconforming-product">
                <label for="nonconforming-product">产品不符</label>
            </div>

            <div>
                <input type="checkbox" id="delayed-dispatch" name="complaint" value="delayed-dispatch">
                <label for="delayed-dispatch">发货延迟</label>
            </div>

            <div>
                <input type="checkbox" id="other-complaint" name="complaint" value="other">
                <label for="other-complaint">其他</label>
            </div>
        </fieldset>

        <div id="complaint-description-container">
            <legend>投诉原因描述</legend>
            <textarea placeholder="请用不少于20个字符描述您的投诉原因"
                name="complaint-textarea" id="complaint-description"></textarea>
        </div>

        <fieldset id="solutions-group">
            <legend>期望解决方案</legend>
            <input type="radio" name="solutions" id="refund" value="refund">
            <label for="refund">退款</label>

            <input type="radio" name="solutions" id="exchange" value="exchange">
            <label for="exchange">换货</label>

            <input type="radio" name="solutions" id="other-solution" value="other">
            <label for="other-solution">其他</label>
        </fieldset>

        <div id="solution-description-container">
            <legend>期望解决方案描述</legend>
            <textarea placeholder="请用不少于20个字符描述您期望的解决方案"
                name="solution-textarea" id="solution-description"></textarea>
        </div>
        <div id="btn-container">
            <button type="submit" id="submit-btn">提交</button>
            <span id="message-box"></span>
            <button type="button" id="clear-btn">清空</button>
        </div>

    </form>

    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

h1 {
    text-align: center;
}

#form {
    max-width: 70%;
    margin: auto;
    border-radius: 10px;
    box-shadow: rgba(0, 0, 0, 0.35) 0px 5px 15px;
    padding: 10px;
}

input {
    border-color: rgb(118, 118, 118);
}

#personal-info input, #product-info input {
    width: 100%;
    margin-bottom: 10px;
}


fieldset {
    margin-bottom: 10px;
    border-radius: 5px;
    border-color: rgb(118, 118, 118);
}

textarea {
    width: 100%;
    border-color: rgb(118, 118, 118);
}

#btn-container {
    display: flex;
    justify-content: space-between;
    align-items: center;
}

#submit-btn, #clear-btn {
    margin: 10px 15px 0;
}
```

```js
const form = document.getElementById("form");
const fullName = document.getElementById("full-name");
const emailAddress = document.getElementById("email");
const orderNo = document.getElementById("order-no");
const productCode = document.getElementById("product-code");
const quantity = document.getElementById("quantity");
const complaints = document.querySelectorAll('input[name="complaint"]');
const complaintDescription = document.getElementById("complaint-description");
const complaintTextAreaDiv = document.getElementById("complaint-description-container");
const solutions = document.querySelectorAll('input[name="solutions"]');
const solutionDescription = document.getElementById("solution-description");
const solutionTextAreaDiv = document.getElementById("solution-description-container");
const submitBtn = document.getElementById("submit-btn");
const messageBox = document.getElementById("message-box");
const clearBtn = document.getElementById("clear-btn");

const hideTextareas = () => {
    complaintTextAreaDiv.style.display = "none";
    solutionTextAreaDiv.style.display = "none";
}

hideTextareas();

const formFields = {
    "full-name": fullName,
    "email": emailAddress,
    "order-no": orderNo,
    "product-code": productCode,
    "quantity": quantity,
    "complaints-group": Array.from(complaints),
    "complaint-description": complaintDescription,
    "solutions-group": Array.from(solutions),
    "solution-description": solutionDescription
}

const clearForm = () => {
    Object.entries(formFields).forEach(entry => {
        const [key, val] = entry;
        if (Array.isArray(val)) {
            val.forEach(i => i.checked = false);
        } else {
            val.value = "";
        }
        hideTextareas();
        messageBox.innerText = "";
        document.getElementById(key).style.borderColor = "rgb(118, 118, 118)";
    })
}

const validateForm = () => {
    const complaintsArr = Array.from(complaints).map(i => i.checked);
    const solutionsArr = Array.from(solutions).map(i => i.checked);
    const validationObject = {
        "full-name": Boolean(fullName.value.trim()),
        "email": /^\S+@\S+\.\S+$/.test(emailAddress.value.trim()),
        "order-no": /^2024\d{6}$/.test(orderNo.value.trim()),
        "product-code": /[A-Z]{2}\d{2}-[A-Z]\d{3}-[A-Z]{2}\d/i.test(productCode.value.trim()),
        "quantity": Number(quantity.value.trim()) > 0,
        "complaints-group": complaintsArr.some(i => i),
        "complaint-description": null,
        "solutions-group": solutionsArr.some(i => i),
        "solution-description": null
    }

    if (complaintsArr[3]) {
        complaintTextAreaDiv.style.display = "block";
        const complaintsDescriptionVal = complaintDescription.value.trim();
        if (complaintsDescriptionVal.length < 20) {
            validationObject["complaint-description"] = false;
        } else {
            validationObject["complaint-description"] = true;
        }
    } else {
        complaintTextAreaDiv.style.display = "none";
    }
    if (solutionsArr[2]) {
        solutionTextAreaDiv.style.display = "block";
        const solutionDescriptionVal = solutionDescription.value.trim();
        if (solutionDescriptionVal.length < 20) {
            validationObject["solution-description"] = false;
        } else {
            validationObject["solution-description"] = true;
        }
    } else {
        solutionTextAreaDiv.style.display = "none";
    }
    if (validationObject["complaint-description"] === null) delete validationObject["complaint-description"];
    if (validationObject["solution-description"] === null) delete validationObject["solution-description"];

    return validationObject;
}

clearBtn.addEventListener("click", clearForm);

Object.keys(formFields).forEach(key => {
    document.getElementById(key).addEventListener("change", (e) => {
        const validationObject = validateForm();
        if (Object.values(validationObject).every(val => val)) messageBox.innerText = "";
        if (!validationObject[key]) {
            document.getElementById(key).style.borderColor = "red";
        } else {
            document.getElementById(key).style.borderColor = "green";
        }
    });
});

const isValid = (validationObj) => Object.values(validationObj).every(i => i);

form.addEventListener("submit", (e) => {
    e.preventDefault();
    const validationObject = validateForm();
    if (!isValid(validationObject)) {
        messageBox.innerText = "请正确填写所有必填项后再提交。";
        Object.keys(validationObject).forEach(key => {
            if (!validationObject[key]) {
                document.getElementById(key).style.borderColor = "red";
            } else {
                document.getElementById(key).style.borderColor = "green";
            }
        });
    } else {
        messageBox.innerText = "";
        alert("表单提交成功。");
        clearForm();
    }
});
```

