---
id: 66ed41f912d0bb1dc62da5dd
title: 构建一个 JavaScript 趣味问答机器人
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你需要在控制台输出 `"Hello! I'm your coding fun fact guide!"` 作为对用户的问候信息。
2. 你需要创建三个变量：`botName`、`botLocation` 和 `favoriteLanguage`，分别存储机器人的名字、来自哪里以及它最喜欢的编程语言。
3. 你需要使用字符串拼接，将 `"My name is (botName) and I live on (botLocation)."` 输出到控制台。
4. 你需要使用字符串拼接，将 `"My favorite programming language is (favoriteLanguage)."` 输出到控制台。
5. 你需要使用 `let` 创建一个变量 `codingFact`，并赋值为一个关于你机器人最喜欢编程语言的有趣事实，且该字符串需通过拼接包含 `favoriteLanguage` 变量。
6. 你需要将 `codingFact` 输出到控制台。
7. 你需要重新为 `codingFact` 赋值为另一个关于机器人最喜欢语言的新事实，并再次使用 `favoriteLanguage` 变量。
8. 你需要再次将 `codingFact` 输出到控制台。
9. 你需要再次为 `codingFact` 赋值为另一个关于机器人最喜欢语言的新事实，并再次使用 `favoriteLanguage` 变量。
10. 你需要第三次将 `codingFact` 输出到控制台。
11. 你需要使用字符串拼接，将 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` 作为机器人的告别语输出到控制台。

# --hints--

你需要在控制台输出 `"Hello! I'm your coding fun fact guide!"`。

```js
assert.equal(output[0], "Hello! I'm your coding fun fact guide!")
```

你需要声明一个 `botName` 变量。请检查拼写和大小写。

```js
assert.exists(botName);
```

你的 `botName` 变量应该是字符串类型。

```js
assert.isString(botName);
```

你需要声明一个 `botLocation` 变量。请检查拼写和大小写。

```js
assert.exists(botLocation);
```

你的 `botLocation` 变量应该是字符串类型。

```js
assert.isString(botLocation);
```

你需要声明一个 `favoriteLanguage` 变量。请检查拼写和大小写。

```js
assert.exists(favoriteLanguage);
```

你的 `favoriteLanguage` 变量应该是字符串类型。

```js
assert.isString(favoriteLanguage);
```

你需要使用拼接将 `"My name is (botName) and I live on (botLocation)."` 输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[1], `My name is ${botName} and I live on ${botLocation}.`)
assert.match(codeWithoutComments, /is ("|')\s*\+\s*botName\s*\+\s*("|') and I live on \2\s*\+\s*botLocation\s*\+\s*('|")\./)
```

你需要使用拼接将 `"My favorite programming language is (favoriteLanguage)."` 输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[2], `My favorite programming language is ${favoriteLanguage}.`)
assert.match(codeWithoutComments, /language is ('|")\s*\+\s*favoriteLanguage\s*\+\s*('|")\./);
```

你需要使用 `let` 声明一个新变量 `codingFact`。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

你需要给 `codingFact` 赋值，并通过拼接包含 `favoriteLanguage`。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|')?.+?\1?\s*\+\s*|favoriteLanguage\s*\+\s*(("|')?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
```

你需要将 `codingFact` 输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(output[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

你需要为 `codingFact` 赋予一个新值（同样包含 `favoriteLanguage`），并输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(output[4], favoriteLanguage);
assert.notEqual(output[4], output[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

你需要第三次为 `codingFact` 赋值（同样包含 `favoriteLanguage`），并输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(output[5], favoriteLanguage);
assert.notEqual(output[5], output[4]);
assert.equal(output[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.exists(third);
assert.isNotEmpty(codingFact); 
```

你需要使用拼接将 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` 输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
assert.match(codeWithoutComments, /\. Goodbye! - ("|')\s*\+\s*botName\s*\+\s*('|") from \2\s*\+\s*botLocation\s*\+\s*("|')\./)
```

# --seed--

## --before-user-code--

```js
const temp = console.log
const output = []
console.log = function (...args) {
temp(...args)
output.push(...args)
}
```

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```

