---
id: 61fd67a656743144844941cb
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕阅读器根据文档流朗读 HTML 元素。我们最终希望资产负债表有一个“Balance Sheet”的标题和一个“AcmeWidgetCorp”的副标题。但是，如果由屏幕阅读器朗读，这个顺序就没有意义了。

为你现有的 `span` 设置 `class` 属性为 `flex`，并在其中添加两个 `span` 元素。第一个 `span` 的文本为 `AcmeWidgetCorp`。第二个 `span` 的文本为 `Balance Sheet`。你将使用 CSS 来反转页面上文本的顺序，但 HTML 顺序对于屏幕阅读器来说更有意义。

# --hints--

你现有的 `span` 元素应该将 `class` 属性设置为 `flex`。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

你现有的 `span` 元素中应该有两个新的 `span` 元素。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

你的新 `span` 元素不应该有 `class` 属性。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

你的第一个新 `span` 元素应该有文本 `AcmeWidgetCorp`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent,
  'AcmeWidgetCorp'
);
```

你的第二个新 `span` 元素应该有文本 `Balance Sheet`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent,
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```

