---
id: 61fd75ea7f663457612dba02
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

在每个新的 `th` 元素中，嵌套一个 `span` 元素，其 `class` 设置为 `sr-only year`。按顺序给它们以下文本：`2019`、`2020` 和 `2021`。

给你的第三个 `th` 元素设置 `class` 属性为 `current`。

将 `td` 元素留空。此元素仅用于确保你的表格具有四列布局并将标题与正确的列相关联。

# --hints--

每个 `th` 元素都应该有一个 `span` 元素。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

每个新的 `span` 元素都应该将 `class` 属性设置为 `sr-only year`。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

你的第一个 `span` 元素应该有文本 `2019`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent,
  '2019'
);
```

你的第二个 `span` 元素应该有文本 `2020`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent,
  '2020'
);
```

你的第三个 `span` 元素应该有文本 `2021`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent,
  '2021'
);
```

你的第三个 `th` 元素应该将 `class` 设置为 `current`。

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

你的 `td` 元素应该是空的。

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent, '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```

