```markdown
---
id: 61fd7b3fcaa5406257abc5d1
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

现在来处理你的第二个表格。首先，给它一个 `caption` 元素，内容设置为 `负债`。然后添加 `thead` 和 `tbody`。

# --hints--

你的第二个 `table` 元素应该有一个 `caption` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('table')?.[1]?.children?.[0]?.localName,
  'caption'
);
```

你的 `caption` 元素的文本应为 `负债`。

```js
assert.strictEqual(
  document.querySelectorAll('table')?.[1]?.querySelector('caption')
    ?.textContent,
  '负债'
);
```

你的第二个 `table` 元素应该有一个 `thead` 元素。

```js
assert.isNotNull(
  document.querySelectorAll('table')?.[1]?.querySelector('thead')
);
```

你的第二个 `table` 元素应该有一个 `tbody` 元素。

```js
assert.isNotNull(
  document.querySelectorAll('table')?.[1]?.querySelector('tbody')
);
```

你的 `thead` 元素应该紧跟在 `caption` 元素之后。

```js
const table = document.querySelectorAll('table')?.[1];
assert.strictEqual(
  table?.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

你的 `tbody` 元素应该紧跟在 `thead` 元素之后。

```js
const table = document.querySelectorAll('table')?.[1];
assert.strictEqual(
  table?.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>资产负债表</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>资产负债表</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              资产
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>现金 <span class="description">这是我们目前持有的现金。</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>支票账户 <span class="description">我们的主要交易账户。</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>储蓄账户 <span class="description">为应急情况预留的资金。</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>总计 <span class="sr-only">资产</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
