---
id: 67fccc2463253b213a000142
title: 构建多媒体播放器
challengeType: 14
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

在之前的讲座视频中，你已经学习了如何使用 `audio` 和 `video` 元素。在本实验中，你将构建一个多媒体播放器，展示一个 `audio` 音轨和一个带有字幕的 `video`。

对于 `audio` 元素，你需要包含一个 `source` 元素，用于指定所用的媒体资源。

示例：

```html
<audio controls aria-label="此处填写描述性标签">
  <source src="音频文件链接" type="audio/mpeg">
</audio>
```

`source` 元素同样可以用于 `video` 元素，如下所示：

```html
<video controls width="600" aria-label="此处填写描述性标签">
  <source src="mp4视频链接" type="video/mp4">
  <!-- 其余代码放在这里 -->  
</video>
```

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个 `h1` 元素作为页面主标题。
2. 你应该有三个 `section` 元素。
3. 在第一个 `section` 元素内，应该有一个 `h2` 元素，作为正在播放的歌曲标题。
4. 在 `h2` 元素下方，应该有一个带有 `controls` 属性和 `aria-label` 属性的 `audio` 元素。
5. 在 `audio` 元素内，应该有一个带有指向音频文件的 `src` 属性和 `type` 属性的 `source` 元素。你可以使用此音频链接：`https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 在第二个 `section` 元素内，应该有一个 `h2` 元素，作为正在播放的视频标题。
7. 在 `h2` 元素下方，应该有一个带有 `controls`、`width` 属性和 `aria-label` 属性的 `video` 元素。
8. 在 `video` 元素内，应该有一个带有指向视频文件的 `src` 属性和 `type` 属性的 `source` 元素。你可以使用此视频链接：`https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4`
9. 在 `source` 元素下方，应该有一个带有指向字幕文件的 `src` 属性、`kind` 属性、`srclang` 属性和 `label` 属性的 `track` 元素。
10. 在第三个 `section` 元素内，应该有一个 `h2` 元素作为该部分的标题，例如“文字稿”。
11. 在 `h2` 元素下方，应该有一个包含视频文字稿的 `p` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.isNotNull(document.querySelector("h1"));
```

你的 `h1` 应该包含代表页面主标题的文本。

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

你应该有三个 `section` 元素。

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

第一个 `section` 应包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

你的 `h2` 应该包含代表正在播放歌曲标题的文本。

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

第一个 `section` 应包含一个 `audio` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

你的 `audio` 元素应有 `controls` 属性。

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

你的 `audio` 元素应有一个描述音频的 `aria-label` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

你的 `audio` 元素应包含一个 `source` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

你的 `source` 元素应有指向音频文件的 `src` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

你的 `source` 元素应有指定音频类型的 `type` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

第二个 `section` 应包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

你的 `h2` 应该包含代表正在播放视频标题的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第二个 `section` 应包含一个 `video` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

你的 `video` 元素应有 `controls` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

你的 `video` 元素应有指定视频宽度的 `width` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

`video` 元素应包含一个 `source` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

`source` 元素应有指向视频文件的 `src` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

`source` 元素应有指定视频类型的 `type` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

你的 `video` 元素应包含一个 `track` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

你的 `track` 元素应有指向字幕文件的 `src` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

你的 `track` 元素应有描述文件类型的 `kind` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

你的 `track` 元素应有描述字幕语言的 `srclang` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

你的 `track` 元素应有描述字幕语言的 `label` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

第三个 `section` 应包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

你的 `h2` 应该包含代表该部分标题的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第三个 `section` 应包含一个用于文字稿的 `p` 元素。

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>多媒体播放器</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>多媒体播放器</title>
</head>
<body>
  <h1>多媒体播放器</h1>
  <section>
    <h2>正在播放：Sailing Away</h2>
    <audio controls aria-label="Sailing Away 音频播放器">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>什么是 map 方法，它是如何工作的？</h2>
    <video controls width="600" aria-label="什么是 map 方法，它是如何工作的？">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4" type="video/mp4">

      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>文字稿</h2>
    <p>什么是 map 方法，它是如何工作的？map 方法是 JavaScript 中一个功能强大且广泛使用的数组方法。它用于通过对原数组的每个元素应用给定函数来创建一个新数组。该方法不会修改原数组，而是返回一个包含每个元素应用函数结果的新数组。
    </p>
  </section>

</body>
</html> 
```
