---
id: 66f1805aac05f00a7028fbcc
title: 构建哈希表类
challengeType: 14
dashedName: build-a-hash-table-class
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你需要定义一个名为 `HashTable` 的类，并在其 `constructor` 方法中初始化一个名为 `collection` 的属性为一个空对象。`collection` 对象应根据键的哈希值存储键值对。
1. `HashTable` 类应包含四个实例方法，分别为 `hash`、`add`、`remove` 和 `lookup`。
1. `hash` 方法应接受一个字符串参数，并返回该字符串每个字符的 UTF-16 码单元之和作为哈希值。请使用 `charCodeAt` 方法实现。
1. `add` 方法应接受一个键值对作为参数，对键进行哈希。哈希值应被用来在哈希表（`collection` 对象）中存储该键值对。当多个不同的键产生相同的哈希值时，每个键值对都应存储在同一个哈希值下。
1. `remove` 方法应接受一个键作为参数，并从哈希表中移除对应的键值对。如果这是该哈希值下的最后一个键值对，也应移除该哈希键。
1. `lookup` 方法应接受一个键作为参数，并返回哈希表中对应的值，如果不存在则返回 `null`。

# --hints--

你应该有一个 `HashTable` 类。

```js
assert.isFunction(HashTable);
assert.isObject(new HashTable());
```

你的 `HashTable` 类应有一个在构造函数中初始化为空对象的 `collection` 属性。

```js
const table = new HashTable();
assert.isObject(table.collection);
assert.isEmpty(table.collection);
```

你的 `HashTable` 类应有一个 `hash` 方法。

```js
const table = new HashTable();
assert.isFunction(table.hash);
```

`new HashTable().hash("freeCodeCamp")` 应返回 `1182`。

```js
const table = new HashTable();
assert.strictEqual(table.hash("freeCodeCamp"), 1182);
```

`hash` 方法应接受一个字符串参数，并返回该字符串每个字符的 UTF-16 码单元之和作为哈希值，需使用 `charCodeAt` 方法。

```js
const table = new HashTable();
assert.strictEqual(table.hash("Hey you"), 675);
```

你的 `HashTable` 类应有一个 `add` 方法。

```js
const table = new HashTable();
assert.isFunction(table.add);
```

调用 `hashTable.add("FCC", "freeCodeCamp")` 后，`hashTable.collection` 应变为 `{ 204: { FCC: 'freeCodeCamp' } }`。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
assert.deepEqual(table.collection, { 204: { FCC: 'freeCodeCamp' } });
```

`add` 方法应接受一个键值对作为参数，对键进行哈希，并将键值对存储在 `collection` 对象的哈希键下。

```js
const table = new HashTable();
table.add("test", "this is a test");
assert.deepEqual(table.collection, { 448: { test: 'this is a test' } });
```

依次调用 `hashTable.add("FCC", "freeCodeCamp")` 和 `hashTable.add("CFC", "chloroflourocarbon")` 后，`hashTable.collection` 应变为 `{ 204: { FCC: 'freeCodeCamp', CFC: 'chloroflourocarbon' } }`。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.add("CFC", "chloroflourocarbon");
assert.deepEqual(table.collection, { 204: { FCC: 'freeCodeCamp', CFC: 'chloroflourocarbon' } });
```

当多个不同的键产生相同的哈希值时，`add` 方法应将每个键值对存储在同一个哈希值下。

```js
const table = new HashTable();
table.add("pan", "frying");
table.add("nap", "sleeping");
assert.deepEqual(table.collection, { 319: { pan: 'frying', nap: 'sleeping' } });
```

你的 `HashTable` 类应有一个 `remove` 方法。

```js
const table = new HashTable();
assert.isFunction(table.remove);
```

当 `hashtable.collection` 等于 `{ 204: { FCC: 'freeCodeCamp' } }` 时，`hashTable.remove("FCC")` 应同时移除 `FCC` 键和 `204` 哈希键。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.remove("FCC");
assert.isEmpty(table.collection);
```

当 `hashtable.collection` 等于 `{ 204: { FCC: 'freeCodeCamp', CFC: 'chloroflourocarbon' } }` 时，`hashTable.remove("FCC")` 只应移除 `204` 哈希键下的 `FCC` 键。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.add("CFC", "chloroflourocarbon");
table.remove("FCC");
assert.deepEqual(table.collection, { 204: { CFC: 'chloroflourocarbon' } });
```

`remove` 方法应接受一个键作为输入，并在哈希键下为空时移除该哈希键。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.remove("FCC");
assert.isEmpty(table.collection);
```

`remove` 方法应只移除哈希键对象下正确的键值对。

```js
const table = new HashTable();
table.add("pan", "frying");
table.add("nap", "sleeping");
table.remove("pan");
assert.deepEqual(table.collection, { 319: { nap: 'sleeping' } });
```

你的 `HashTable` 类应有一个 `lookup` 方法。

```js
const table = new HashTable();
assert.isFunction(table.lookup);
```

`hashTable.lookup("FCC")` 应返回存储在 `hashtable.collection` 中 `FCC` 键的值。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.add("CFC", "chloroflourocarbon");
assert.strictEqual(table.lookup("FCC"), "freeCodeCamp");
```

`hashTable.lookup("FCC")` 若未找到 `FCC` 键，应返回 `null`。

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.add("CFC", "chloroflourocarbon");
assert.isNull(table.lookup("x"));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class HashTable {
    constructor() {
        this.collection = {};
    }

    hash(string) {
        let hashed = 0;
        for (let i = 0; i < string.length; i++) {
            hashed += string.charCodeAt(i);
        }
        return hashed
    }

    add(key, val) {
        const theHash = this.hash(key);
        if (!this.collection.hasOwnProperty(theHash)) {
            this.collection[theHash] = {};
        }
        this.collection[theHash][key] = val;
    }

    remove(key) {
        const theHash = this.hash(key);
        const hashedObj = this.collection[theHash];
        if (hashedObj.hasOwnProperty(key)) {
            delete hashedObj[key];
        }
        if (!Object.keys(hashedObj).length) {
            delete this.collection[theHash];
        }
    }

    lookup(key) {
        const theHash = this.hash(key);
        if (this.collection.hasOwnProperty(theHash)) {
            return this.collection[theHash][key];
        }
        return null
    }
}
```

