---
id: 67172b43f84bcd2dec238a3d
title: 构建一本图书整理器
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个名为 `books` 的对象数组，数组中的每个对象都应包含一个字符串类型的 `title`，另一个字符串类型的 `authorName`，以及一个数字类型的 `releaseYear`。

2. 你的 `books` 数组中至少应有三个对象。

3. 你应该有一个名为 `sortByYear` 的回调函数，该函数接受两个图书对象作为参数用于排序。

4. 如果第一个图书的 `releaseYear` 小于第二个图书的 `releaseYear`，`sortByYear` 函数应返回 `-1`。

5. 如果第一个图书的 `releaseYear` 大于第二个图书的 `releaseYear`，`sortByYear` 函数应返回 `1`。

6. 如果两个图书的 `releaseYear` 相等，`sortByYear` 函数应返回 `0`。

7. 你应该从 `books` 数组中过滤出出版年份早于某一年（如 1950 年）的图书，并将过滤后的数组保存到一个名为 `filteredBooks` 的新数组中。

8. 你应该按照 `releaseYear` 升序对 `filteredBooks` 数组中的图书进行排序。你在之前的视频中学到，`sort()` 方法会原地排序数组，这意味着 `filteredBooks` 数组会被改变。

# --hints--

你的代码中应该有一个 `sortByYear` 函数。

```js
assert.isFunction(sortByYear);
```

你的 `sortByYear` 函数应接受两个参数。

```js
assert.lengthOf(sortByYear, 2);
```

你的 `sortByYear` 函数应在 `book1` 的 `releaseYear` 小于 `book2` 时返回 `-1`，大于时返回 `1`，其他情况返回 `0`。

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

你的代码中应该有一个 `books` 数组。

```js
assert.isArray(books)
```

你的 `books` 数组中应至少有三个图书对象。

```js
assert.isAtLeast(books.length, 3)
```

你的 `books` 数组应包含每个对象都具有字符串类型的 `title`、字符串类型的 `authorName` 和数字类型的 `releaseYear`。

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

你的代码中应该有一个 `filteredBooks` 数组。

```js
assert.isArray(filteredBooks)
```

`filteredBooks` 数组应从 `books` 数组中过滤出部分图书且不为空。

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

你应该在 `filteredBooks` 数组上调用 `sort` 高阶函数，并传入 `sortByYear` 回调函数。

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "了不起的盖茨比",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "追忆似水年华",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "尤利西斯",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "百年孤独",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "麦田里的守望者",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("1950 年前出版的图书（按出版年份排序）");
filteredBooks.forEach((book) => {
  console.log(`${book.title} 作者：${book.authorName}（${book.releaseYear}）`);
});
```

