---
id: 673c91f0b934834bc4a3ecc2
title: 构建一个 fCC 论坛排行榜
challengeType: 25
dashedName: build-an-fcc-forum-leaderboard
demoType: onClick
---

# --description--

在本实验中，你将构建一个 freeCodeCamp 论坛排行榜，显示来自 [freeCodeCamp 论坛](https://forum.freecodecamp.org/) 的最新话题、用户和回复。HTML、CSS 和部分 JS 已为你提供。你可以自由探索它们。

完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个名为 `timeAgo` 的函数，它接收一个 ISO 8601 格式的时间戳作为参数。
1. `timeAgo` 函数应计算传入时间与当前时间的差值，并返回：
  - 如果相差小于 60 分钟，返回 `xm ago`（`x` 代表分钟数）。
  - 如果相差小于 24 小时，返回 `xh ago`（`x` 代表小时数）。
  - 否则返回 `xd ago`（`x` 代表天数）。
1. 你应该有一个名为 `viewCount` 的函数，它接收一个帖子的浏览数作为参数。
1. 如果传入的浏览数大于等于 1000，`viewCount` 应返回一个字符串，将浏览数除以 1000 向下取整并加上字母 `k`。否则返回原始浏览数。
1. 你应该有一个名为 `forumCategory` 的函数，它接收所选分类的 id 作为参数。
1. `forumCategory` 函数应验证所选分类 id 是否为 `allCategories` 对象的属性，并返回一个包含以下内容的锚点元素字符串：
  - 选中分类的 `category` 作为文本。
  - 类名为 `category` 加上选中分类的 `className`。
  - `href` 的值为 `<forumCategoryUrl>/<className>/<id>`，其中 `<className>` 是选中分类的 `className`，`id` 是传入的参数。
1. 如果 `allCategories` 对象没有该分类 id，应将 `category` 标记为 `General`，`className` 标记为 `general`。
1. 你应该有一个名为 `avatars` 的函数，它接收两个数组，分别代表 posters 和 users。
1. `avatars` 函数应返回一个字符串，由每个在用户数组中找到的 poster 生成的 `img` 元素拼接而成。_提示：_ 你可以通过比较 poster 的 `user_id` 和 user 的 `id` 来找到用户。
1. `avatars` 函数应通过访问 `avatar_template` 属性并将 `{size}` 替换为 `30` 来设置每个头像的大小。
1. 每个图片元素应有 alt 文本，值为 poster 的 `name` 属性。
1. 每个图片元素的 src 应为 poster 的 `avatar_template` 属性。如果 `avatar_template` 是相对路径，应设置为 `<avatarUrl>/<avatar_template>`。
1. 你应该有一个名为 `showLatestPosts` 的函数，接收一个参数。
1. `showLatestPosts` 应从传入对象中提取 `users` 和 `topic_list` 属性，并处理 `topic_list` 下 `topics` 数组中的以下属性：
  - `id`：帖子的 id
  - `title`：帖子的标题
  - `views`：帖子的浏览数
  - `posts_count`：该话题的回复数
  - `slug`：帖子的 slug
  - `posters`：该话题的发帖人
  - `category_id`：帖子的分类 id
  - `bumped_at`：ISO 8601 格式的时间戳
1. `showLatestPosts` 应将 `#posts-container` 的 inner HTML 设置为由每个 `topics` 项拼接而成的 `tr` 元素字符串。
1. 每个 `tr` 元素应包含五个 `td` 元素：
  - 一个 `td`，包含两个锚点元素，第一个 class 为 `post-title`，`href` 为 `<forumTopicUrl><slug>/<id>`，文本为 `<title>`，第二个通过调用 `forumCategory` 并传入 `category_id` 获得。
  - 一个 `td`，包含一个 class 为 `avatar-container` 的 `div`，其中包含调用 `avatars` 并传入 `posters` 和 `users` 得到的图片。
  - 一个 `td`，显示该帖子的回复数。_提示：_ 使用 `posts_count - 1`。
  - 一个 `td`，显示该帖子的浏览数。
  - 一个 `td`，显示自上次活动以来的时间。
1. 你应该有一个名为 `fetchData` 的异步函数。
1. `fetchData` 函数应从 `forumLatest` 请求数据，并将解析后的 JSON 传递给 `showLatestPosts`。
1. 如果获取数据时出错，`fetchData` 应使用 `console.log` 将错误输出到控制台。

# --hints--

你应该有一个名为 `timeAgo` 的函数，接收一个参数。

```js
assert.isFunction(timeAgo);
assert.lengthOf(timeAgo, 1);
```

当传入的时间与当前时间相差 50 分钟时，`timeAgo` 应返回 `50m ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 50).toISOString();
};
const expected = '50m ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 60 分钟时，`timeAgo` 应返回 `1h ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60).toISOString();
};
const expected = '1h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 115 分钟时，`timeAgo` 应返回 `1h ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * (60 * 115)).toISOString();
};
const expected = '1h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 15 小时时，`timeAgo` 应返回 `15h ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 15).toISOString();
};
const expected = '15h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 24 小时时，`timeAgo` 应返回 `1d ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 24).toISOString();
};
const expected = '1d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 46 小时时，`timeAgo` 应返回 `1d ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 46).toISOString();
};
const expected = '1d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

当时间差为 3 天时，`timeAgo` 应返回 `3d ago`。

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 24 * 3).toISOString();
};
const expected = '3d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

你应该有一个名为 `viewCount` 的函数，接收一个参数。

```js
assert.isFunction(viewCount);
assert.lengthOf(viewCount, 1);
```

`viewCount(597)` 应返回 `597`。

```js
assert.strictEqual(597, viewCount(597));
```

`viewCount(1000)` 应返回 `1k`。

```js
assert.equal('1k', viewCount(1000));
```

`viewCount(2730)` 应返回 `2k`。

```js
assert.equal('2k', viewCount(2730));
```

你应该有一个名为 `forumCategory` 的函数，接收一个参数。

```js
assert.isFunction(forumCategory);
assert.lengthOf(forumCategory, 1);
```

`forumCategory(299)` 应返回包含文本 `Career Advice` 的锚点元素字符串。

```js
let actual = forumCategory(299);
assert.match(actual, /^<\s*a.+?>\s*Career Advice\s*<\/a>$/);
// 防止硬编码

actual = forumCategory(409);
assert.match(actual, /^<\s*a.+?>\s*Project Feedback\s*<\/a>$/);
```

`forumCategory(299)` 应返回包含 `href="https://forum.freecodecamp.org/c/career/299"` 的锚点元素字符串。

```js
let actual = forumCategory(299);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/career\/299\1/
);

// 防止硬编码
actual = forumCategory(409);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/feedback\/409\1/
);
```

`forumCategory(299)` 应返回包含 `class="category career"` 的锚点元素字符串。

```js
let actual = forumCategory(299);
assert.match(actual, /class=("|')category\s+career\1/);

// 防止硬编码
actual = forumCategory(409);
assert.match(actual, /class=("|')category\s+feedback\1/);
```

`forumCategory(200)` 应返回包含文本 `General` 的锚点元素字符串。

```js
const actual = forumCategory(200);
assert.match(actual, /^<\s*a.+?>\s*General\s*<\/a>$/);
```

`forumCategory(200)` 应返回包含 `href="https://forum.freecodecamp.org/c/general/200"` 的锚点元素字符串。

```js
let actual = forumCategory(200);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/general\/200/
);

actual = forumCategory(220);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/general\/220/
);
```

`forumCategory(200)` 应返回包含 `class="category general"` 的锚点元素字符串。

```js
const actual = forumCategory(200);
assert.match(actual, /class=("|')category\s+general\1/);
```

你应该有一个名为 `avatars` 的函数，接收两个参数。

```js
assert.isFunction(avatars);
assert.lengthOf(avatars, 2);
```

`avatars` 函数应返回一个字符串，由每个在用户数组中找到的 poster 生成的 `img` 元素拼接而成。

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
   id: 6,
   name: 'Quincy Larson',
   username: 'QuincyLarson'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
   id: 285941,
   name: 'Jessica Wilkins',
   username: 'jwilkins.oboe'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
   id: 170865,
   name: 'Ilenia',
   username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
const matches = actual.match(/<\s*img\s+.+?>/g);
assert.lengthOf(matches, 3);
```

`avatars` 返回的字符串中每个 `img` 元素应有 alt 文本，值为 poster 的 `name` 属性。

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
   id: 6,
   name: 'Quincy Larson',
   username: 'QuincyLarson'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
   id: 285941,
   name: 'Jessica Wilkins',
   username: 'jwilkins.oboe'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
   id: 170865,
   name: 'Ilenia',
   username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
const matches = actual.match(/<\s*img\s+.+?>/g);

assert.match(matches[0], /alt="Quincy Larson"/);
assert.match(matches[1], /alt="Jessica Wilkins"/);
assert.match(matches[2], /alt="Ilenia"/);
```

`avatars` 函数应通过访问 `avatar_template` 属性并将 `{size}` 替换为 `30` 来设置每个头像的大小。

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
   id: 6,
   name: 'Quincy Larson',
   username: 'QuincyLarson'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
   id: 285941,
   name: 'Jessica Wilkins',
   username: 'jwilkins.oboe'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
   id: 170865,
   name: 'Ilenia',
   username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
assert.notMatch(actual, /\{size\}/);
assert.lengthOf(actual.match(/\/30\//g), 3);
```

`avatars` 返回的字符串中每个 `img` 元素的 src 应为 poster 的 `avatar_template` 属性。如果 `avatar_template` 是相对路径，应设置为 `<avatarUrl>/<avatar_template>`。

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
   id: 6,
   name: 'Quincy Larson',
   username: 'QuincyLarson'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
   id: 285941,
   name: 'Jessica Wilkins',
   username: 'jwilkins.oboe'
  },
  {
   avatar_template:
    '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
   id: 170865,
   name: 'Ilenia',
   username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);

const matches = actual.match(/<\s*img\s+.+?>/g);

assert.match(
  matches[0],
  /src="https:\/\/sea1\.discourse-cdn\.com\/freecodecamp\/user_avatar\/forum\.freecodecamp\.org\/quincylarson\/30\/212400_2\.png"/
);
assert.match(
  matches[1],
  /src="https:\/\/sea1\.discourse-cdn\.com\/freecodecamp\/user_avatar\/forum\.freecodecamp\.org\/jwilkins\.oboe\/30\/179497_2\.png"/
);
assert.match(
  matches[2],
  /src="https:\/\/sea1\.discourse-cdn\.com\/freecodecamp\/user_avatar\/forum\.freecodecamp\.org\/ilenia\/30\/270648_2\.png"/
);
```

你应该有一个名为 `showLatestPosts` 的函数，接收一个参数。

```js
assert.isFunction(showLatestPosts);
assert.lengthOf(showLatestPosts, 1);
```

你应该有一个名为 `fetchData` 的函数。

```js
assert.isFunction(fetchData);
```

你的 `fetchData` 函数应从 `https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json` 请求数据。

```js
const testArr = [];
const temp = fetch;
try {
  fetch = source => {
   testArr.push(source);
   return temp(source);
  };
  fetchData();
  assert.deepEqual(testArr, [
   'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json'
  ]);
} finally {
  fetch = temp;
}
```

你的 `fetchData` 函数应调用 `showLatestPosts`。

```js
const testArr = [];
const temp = showLatestPosts;
async () => {
  try {
   showLatestPosts = data => {
    testArr.push(data);
    return temp(data);
   };
   await fetchData();
   assert.isNotEmpty(testArr);
  } catch (err) {
   throw new Error(err);
  } finally {
   fetch = temp;
  }
};
```

如果出错，你的 `fetchData` 函数应使用 `console.log` 输出错误。

```js
const testArr = [];
const temp1 = fetch;
const temp2 = console.log;
async () => {
  try {
   console.log = obj => {
    testArr.push(obj.toString());
   };
   fetch = source => {
    throw new Error('This is a test error');
   };
   await fetchData();
   assert.deepEqual(testArr, ['Error: This is a test error']);
  } finally {
   fetch = temp1;
   console.log = temp2;
  }
};
```

`showLatestPosts` 应将 `#posts-container` 的 inner HTML 设置为由每个 `topics` 项拼接而成的 `tr` 元素字符串。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
assert.lengthOf(pContainer.querySelectorAll('tr'), 2);
```

`showLatestPosts` 返回的每个表格行 `tr` 元素应包含 5 个 `td` 元素。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
assert.lengthOf(pContainer.querySelectorAll('tr:first-child>td'), 5);
assert.lengthOf(pContainer.querySelectorAll('tr:last-child>td'), 5);
```

每个表格行的第一个 `td` 元素应包含两个锚点，第一个 class 为 `post-title`，`href` 为 `<forumTopicUrl><slug>/<id>`，文本为 `<title>`，第二个通过调用 `forumCategory` 并传入 `category_id` 获得。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
const anchors1 = pContainer.querySelectorAll('tr:first-child>td>a');
assert.lengthOf(anchors1, 2);

const anchors2 = pContainer.querySelectorAll('tr:last-child>td>a');
assert.lengthOf(anchors2, 2);

assert.equal(anchors1[0].classList[0], 'post-title');
assert.equal(
  anchors1[0].href,
  'https://forum.freecodecamp.org/t/the-freecodecamp-podcast-is-back-now-with-video/684569'
);
assert.equal(
  anchors1[0].innerText.trim(),
  'The freeCodeCamp Podcast is back – now with video'
);

assert.equal(anchors1[1].classList[0], 'category');
assert.equal(anchors1[1].classList[1], 'general');
assert.equal(anchors1[1].href, 'https://forum.freecodecamp.org/c/general/1');

assert.equal(anchors2[0].classList[0], 'post-title');
assert.equal(
  anchors2[0].href,
  'https://forum.freecodecamp.org/t/problem-with-making-changes-to-styles-js/686149'
);
assert.equal(
  anchors2[0].innerText.trim(),
  'Problem with making changes to styles. (JS)'
);

assert.equal(anchors2[1].classList[0], 'category');
assert.equal(anchors2[1].classList[1], 'javascript');
assert.equal(
  anchors2[1].href,
  'https://forum.freecodecamp.org/c/javascript/421'
);
```

每个表格行的第二个 `td` 元素应包含调用 `avatars` 并传入 `posters` 和 `users` 得到的图片，嵌套在 class 为 `avatar-container` 的 `div` 元素中。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

const div1 = pContainer.querySelector('tr:first-child>td:nth-child(2)>div');
assert.equal(div1.classList[0], 'avatar-container');

const div2 = pContainer.querySelector('tr:last-child>td:nth-child(2)>div');
assert.equal(div2.classList[0], 'avatar-container');

const imgs1 = div1.querySelectorAll('img');
assert.lengthOf(imgs1, 3);
assert.equal(
  imgs1[0].src,
  'https://sea1.discourse-cdn.com/freecodecamp/user_avatar/forum.freecodecamp.org/quincylarson/30/212400_2.png'
);
assert.equal(imgs1[0].alt, 'Quincy Larson');
assert.equal(
  imgs1[1].src,
  'https://sea1.discourse-cdn.com/freecodecamp/user_avatar/forum.freecodecamp.org/ilenia/30/270648_2.png'
);
assert.equal(imgs1[1].alt, 'Ilenia');
assert.equal(
  imgs1[2].src,
  'https://sea1.discourse-cdn.com/freecodecamp/user_avatar/forum.freecodecamp.org/jwilkins.oboe/30/179497_2.png'
);
assert.equal(imgs1[2].alt, 'Jessica Wilkins');

const imgs2 = div2.querySelectorAll('img');
assert.lengthOf(imgs2, 1);
assert.equal(
  imgs2[0].src,
  'https://sea1.discourse-cdn.com/freecodecamp/user_avatar/forum.freecodecamp.org/ilenia/30/270648_2.png'
);
assert.equal(imgs2[0].alt, 'Ilenia');
```

每个表格行的第三个 `td` 元素应显示帖子的回复数。_提示：_ 使用 `posts_count - 1`。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(3)').innerText,
  '7'
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(3)').innerText,
  '0'
);
```

每个表格行的第四个 `td` 元素应显示帖子的浏览数。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(4)').innerText,
  '542'
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(4)').innerText,
  '9'
);
```

每个表格行的第五个 `td` 元素应显示自上次活动以来的时间，使用 `timeAgo` 函数生成。

```js
const data = {
  users: [
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/quincylarson/{size}/212400_2.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/jwilkins.oboe/{size}/179497_2.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
   },
   {
    avatar_template:
      '/user_avatar/forum.freecodecamp.org/ilenia/{size}/270648_2.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
   }
  ],
  topic_list: {
   topics: [
    {
      bumped_at: '2024-04-15T16:01:26.403Z',
      category_id: 1,
      id: 684569,
      posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
      posts_count: 8,
      slug: 'the-freecodecamp-podcast-is-back-now-with-video',
      title: 'The freeCodeCamp Podcast is back – now with video',
      views: 542
    },
    {
      bumped_at: '2024-04-19T13:52:03.523Z',
      category_id: 421,
      id: 686149,
      posters: [{ user_id: 170865 }],
      posts_count: 1,
      slug: 'problem-with-making-changes-to-styles-js',
      title: 'Problem with making changes to styles. (JS)',
      views: 9
    }
   ]
  }
};
const calcTime = time => {
  const currentTime = new Date();
  const lastPost = new Date(time);
  const timeDifference = currentTime - lastPost;
  const msPerMinute = 1000 * 60;
  const minutesAgo = Math.floor(timeDifference / msPerMinute);
  const hoursAgo = Math.floor(minutesAgo / 60);
  const daysAgo = Math.floor(hoursAgo / 24);
  return `${daysAgo}d ago`;
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(5)').innerText,
  calcTime('2024-04-15T16:01:26.403Z')
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(5)').innerText,
  calcTime('2024-04-19T13:52:03.523Z')
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
   <meta charset="UTF-8" />
   <meta http-equiv="X-UA-Compatible" content="IE=edge" />
   <meta name="viewport" content="width=device-width, initial-scale=1.0" />
   <title>fCC Forum Leaderboard</title>
   <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
   <header>
    <nav>
      <img
       class="fcc-logo"
       src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
       alt="freeCodeCamp logo"
      />
    </nav>
    <h1 class="title">Latest Topics</h1>
   </header>
   <main>
    <div class="table-wrapper">
      <table>
       <thead>
        <tr>
          <th id="topics">Topics</th>
          <th id="avatars">Avatars</th>
          <th id="replies">Replies</th>
          <th id="views">Views</th>
          <th id="activity">Activity</th>
        </tr>
       </thead>
       <tbody id="posts-container"></tbody>
      </table>
    </div>
   </main>
   <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.fcc-logo {
  width: 210px;
  display: block;
  margin: auto;
}

.title {
  margin: 25px 0;
  text-align: center;
  color: var(--light-grey);
}

.table-wrapper {
  padding: 0 25px;
  overflow-x: auto;
}

table {
  width: 100%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
  overflow-x: scroll;
}

#topics {
  text-align: start;
  width: 60%;
}

th {
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.3rem;
}

td:not(:first-child) {
  text-align: center;
}

td {
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0;
}

.post-title {
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category {
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin: 10px 0 10px;
}

.career {
  background-color: var(--salmon-color);
}

.feedback,
.html-css {
  background-color: var(--light-blue);
}

.support {
  background-color: var(--light-orange);
}

.general {
  background-color: var(--light-green);
}

.javascript {
  background-color: var(--golden-yellow);
}

.backend {
  background-color: var(--gold);
}

.python {
  background-color: var(--green);
}

.motivation {
  background-color: var(--peach);
}

.avatar-container {
  display: flex;
  justify-content: center;
  gap: 10px;
  flex-wrap: wrap;
}

.avatar-container img {
  width: 30px;
  height: 30px;
}

@media (max-width: 750px) {
  .table-wrapper {
   padding: 0 15px;
  }

  table {
   width: 700px;
  }

  th {
   font-size: 1.2rem;
  }

  .post-title {
   font-size: 1.1rem;
  }
}
```

```js
const forumLatest =
  'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json';
const forumTopicUrl = 'https://forum.freecodecamp.org/t/';
const forumCategoryUrl = 'https://forum.freecodecamp.org/c/';
const avatarUrl = 'https://sea1.discourse-cdn.com/freecodecamp';

const allCategories = {
  299: { category: 'Career Advice', className: 'career' },
  409: { category: 'Project Feedback', className: 'feedback' },
  417: { category: 'freeCodeCamp Support', className: 'support' },
  421: { category: 'JavaScript', className: 'javascript' },
  423: { category: 'HTML - CSS', className: 'html-css' },
  424: { category: 'Python', className: 'python' },
  432: { category: 'You Can Do This!', className: 'motivation' },
  560: { category: 'Backend Development', className: 'backend' }
};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
   <meta charset="UTF-8" />
   <meta http-equiv="X-UA-Compatible" content="IE=edge" />
   <meta name="viewport" content="width=device-width, initial-scale=1.0" />
   <title>fCC Forum Leaderboard</title>
   <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
   <header>
    <nav>
      <img
       class="fcc-logo"
       src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
       alt="freeCodeCamp logo"
      />
    </nav>
    <h1 class="title">Latest Topics</h1>
   </header>
   <main>
    <div class="table-wrapper">
      <table>
       <thead>
        <tr>
          <th id="topics">Topics</th>
          <th id="avatars">Avatars</th>
          <th id="replies">Replies</th>
          <th id="views">Views</th>
          <th id="activity">Activity</th>
        </tr>
       </thead>
       <tbody id="posts-container"></tbody>
      </table>
    </div>
   </main>
   <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.fcc-logo {
  width: 210px;
  display: block;
  margin: auto;
}

.title {
  margin: 25px 0;
  text-align: center;
  color: var(--light-grey);
}

.table-wrapper {
  padding: 0 25px;
  overflow-x: auto;
}

table {
  width: 100%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
  overflow-x: scroll;
}

#topics {
  text-align: start;
  width: 60%;
}

th {
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.3rem;
}

td:not(:first-child) {
  text-align: center;
}

td {
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0;
}

.post-title {
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category {
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin: 10px 0 10px;
}

.career {
  background-color: var(--salmon-color);
}

.feedback,
.html-css {
  background-color: var(--light-blue);
}

.support {
  background-color: var(--light-orange);
}

.general {
  background-color: var(--light-green);
}

.javascript {
  background-color: var(--golden-yellow);
}

.backend {
  background-color: var(--gold);
}

.python {
  background-color: var(--green);
}

.motivation {
  background-color: var(--peach);
}

.avatar-container {
  display: flex;
  justify-content: center;
  gap: 10px;
  flex-wrap: wrap;
}

.avatar-container img {
  width: 30px;
  height: 30px;
}

@media (max-width: 750px) {
  .table-wrapper {
   padding: 0 15px;
  }

  table {
   width: 700px;
  }

  th {
   font-size: 1.2rem;
  }

  .post-title {
   font-size: 1.1rem;
  }
}
```

```js
const forumLatest =
  'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json';
const forumTopicUrl = 'https://forum.freecodecamp.org/t/';
const forumCategoryUrl = 'https://forum.freecodecamp.org/c/';
const avatarUrl = 'https://sea1.discourse-cdn.com/freecodecamp';

const allCategories = {
  299: { category: 'Career Advice', className: 'career' },
  409: { category: 'Project Feedback', className: 'feedback' },
  417: { category: 'freeCodeCamp Support', className: 'support' },
  421: { category: 'JavaScript', className: 'javascript' },
  423: { category: 'HTML - CSS', className: 'html-css' },
  424: { category: 'Python', className: 'python' },
  432: { category: 'You Can Do This!', className: 'motivation' },
  560: { category: 'Backend Development', className: 'backend' }
};
```

