---
id: 66bba6fff611169359d9d36a
title: 构建一个机密邮件页面
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要有一个 `main` 元素，`id` 为 `email`。
1. 你的 `#email` 元素应有 `50px` 的内边距（padding）、`50px` 的上外边距（margin-top）、`500px` 的宽度（width），以及 `2px` 的边框（border）。
1. 你的 `#email` 元素的总宽度（包括内边距和边框）应为 `500px`。
1. 在你的 `#email` 元素内应有两个 `div` 元素，分别具有 `id` 为 `confidential` 和 `top-secret`。
1. 你的 `#confidential` 和 `#top-secret` 元素应设置为 `inline-block` 显示。
1. 你的 `#confidential` 和 `#top-secret` 元素应有内边距、左外边距和边框。
1. `#confidential` 元素的文本应为 `CONFIDENTIAL`。
1. `#top-secret` 元素的文本应为 `TOP SECRET`。
1. 你的 `#confidential` 和 `#top-secret` 元素应通过 CSS transform 进行旋转。
1. 你的 `#email` 元素内应至少有三个段落（paragraph）元素。
1. 在这些段落元素中，至少应有三个带有 `blurred` 类的 `span` 元素。
1. 你应有一个 `.blurred` 选择器，通过 CSS filter 将元素模糊 3px。

# --hints--

你需要有一个 `main` 元素，`id` 为 `email`。

```js
assert.exists(document.querySelector('main#email'));
```

你应有一个 `#email` 选择器，将其元素的 `padding` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

你应有一个 `#email` 选择器，将其元素的 `margin-top` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

你应有一个 `#email` 选择器，将其元素的 `width` 设置为 `500px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

你的 `#email` 元素应有 `2px` 的边框。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

你的 `#email` 元素应设置 `box-sizing` 为 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

你应在 `#email` 元素内有一个 `id` 为 `confidential` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

你应在 `#email` 元素内有一个 `id` 为 `top-secret` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

你的 `#confidential` 元素应设置 `display` 为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

你的 `#top-secret` 元素应设置 `display` 为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

你的 `#confidential` 元素应有四个方向的内边距。

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#top-secret` 元素应有四个方向的内边距。

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#confidential` 元素应有左外边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

你的 `#top-secret` 元素应有左外边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

你的 `#confidential` 元素应有边框。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

你的 `#top-secret` 元素应有边框。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

你的 `#confidential` 元素文本应为 `CONFIDENTIAL`。请勿遗漏或拼写错误。

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

你的 `#top-secret` 元素文本应为 `TOP SECRET`。请勿遗漏或拼写错误。

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

你的 `#confidential` 元素应通过 CSS `transform` 进行旋转。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

你的 `#top-secret` 元素应通过 CSS `transform` 进行旋转。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

你的 `#email` 元素内应至少有三个段落元素。

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

在这些段落中，至少应有三个带有 `blurred` 类的 `span` 元素。

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

你的 `.blurred` 元素内容不能为空。

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

你应有一个 `.blurred` 选择器，将其元素的 `filter` 设置为 `blur(3px)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>机密邮件</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="zh">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>绝密棉花糖任务</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>亲爱的特工 <span class="blurred">S'more,</span></p>
    <p>我们遇到了紧急情况。<span class="blurred">超级棉花糖</span> 的秘密配方已被泄露。这个配方让我们的棉花糖变得最蓬松、最美味。</p>
    <p>我们怀疑 <span class="blurred">泡芙教授</span> 是幕后黑手。他已将配方带到他的秘密实验室。你的任务是潜入实验室，在为时已晚之前夺回配方。</p>
    <p>请务必保密实验室的位置。任何信息泄露都可能危及整个行动。</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```

