---
id: 66a8e9ca371965b6ef5f2989
title: 构建一个调查表单
challengeType: 14
dashedName: build-a-survey-form
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该有一个 `h1` 元素作为页面标题，并且其 `id` 为 `title`。
1. 你应该有一个简短说明，使用 `p` 元素，并且其 `id` 为 `description`。
1. 你应该有一个 `form` 元素，并且其 `id` 为 `survey-form`。
1. 在表单元素内，你应该有一个必填的 `input` 字段用于输入姓名，其 `id` 为 `name`，`type` 为 `text`。
1. 在表单元素内，你应该有一个必填的 `input` 字段用于输入邮箱，其 `id` 为 `email`。
1. 如果你输入的邮箱格式不正确，应该会看到 HTML5 验证错误。
1. 在表单元素内，你应该有一个用于输入数字的 `input` 字段，其 `id` 为 `number`。
1. 数字输入框不应接受非数字内容，可以通过阻止输入或显示 HTML5 验证错误（取决于浏览器）。
1. 如果输入的数字超出 `min` 和 `max` 属性定义的范围，应该会看到 HTML5 验证错误。
1. 对于姓名、邮箱和数字输入字段，表单中应有相应的 `label` 元素，分别描述每个字段的用途，`id` 分别为：`name-label`、`email-label` 和 `number-label`。
1. 对于姓名、邮箱和数字输入字段，应有占位符文本，提供描述或说明。
1. 在表单元素内，你应该有一个 `select` 下拉菜单，`id` 为 `dropdown`，且至少有两个可选项。
1. 在表单元素内，你可以从一组至少两个的单选按钮中选择一个，这些单选按钮通过 `name` 属性分组。
1. 在表单元素内，你可以从一系列复选框中选择多个，每个复选框都必须有 `value` 属性。
1. 在表单元素内，你应该有一个用于附加评论的 `textarea`。
1. 在表单元素内，你应该有一个 `id` 为 `submit` 的按钮，用于提交所有输入内容。

# --hints--

你应该有一个 `id` 为 `title` 的 `h1` 元素。

```js
const el = document.getElementById('title');
assert(!!el && el.tagName === 'H1');
```

你的 `#title` 不应为空。

```js
const el = document.getElementById('title');
assert(!!el && el.innerText.length > 0);
```

你应该有一个 `id` 为 `description` 的 `p` 元素。

```js
const el = document.getElementById('description');
assert(!!el && el.tagName === 'P');
```

你的 `#description` 不应为空。

```js
const el = document.getElementById('description');
assert(!!el && el.innerText.length > 0);
```

你应该有一个 `id` 为 `survey-form` 的 `form` 元素。

```js
const el = document.getElementById('survey-form');
assert(!!el && el.tagName === 'FORM');
```

你应该有一个 `id` 为 `name` 的 `input` 元素。

```js
const el = document.getElementById('name');
assert(!!el && el.tagName === 'INPUT');
```

你的 `#name` 应该是 `text` 类型。

```js
const el = document.getElementById('name');
assert(!!el && el.type === 'text');
```

你的 `#name` 应该是必填项。

```js
const el = document.getElementById('name');
assert(!!el && el.required);
```

你的 `#name` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #name');
assert(!!el);
```

你应该有一个 `id` 为 `email` 的 `input` 元素。

```js
const el = document.getElementById('email');
assert(!!el && el.tagName === 'INPUT');
```

你的 `#email` 应该是 `email` 类型。

```js
const el = document.getElementById('email');
assert(!!el && el.type === 'email');
```

你的 `#email` 应该是必填项。

```js
const el = document.getElementById('email');
assert(!!el && el.required);
```

你的 `#email` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #email');
assert(!!el);
```

你应该有一个 `id` 为 `number` 的 `input` 元素。

```js
const el = document.getElementById('number');
assert(!!el && el.tagName === 'INPUT');
```

你的 `#number` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #number');
assert(!!el);
```

你的 `#number` 应该是 `number` 类型。

```js
const el = document.getElementById('number');
assert(!!el && el.type === 'number');
```

你的 `#number` 应该有一个数值的 `min` 属性。

```js
const el = document.getElementById('number');
assert(!!el && el.min && isFinite(el.min));
```

你的 `#number` 应该有一个数值的 `max` 属性。

```js
const el = document.getElementById('number');
assert(!!el && el.max && isFinite(el.max));
```

你应该有一个 `id` 为 `name-label` 的 `label` 元素。

```js
const el = document.getElementById('name-label');
assert(!!el && el.tagName === 'LABEL');
```

你应该有一个 `id` 为 `email-label` 的 `label` 元素。

```js
const el = document.getElementById('email-label');
assert(!!el && el.tagName === 'LABEL');
```

你应该有一个 `id` 为 `number-label` 的 `label` 元素。

```js
const el = document.getElementById('number-label');
assert(!!el && el.tagName === 'LABEL');
```

你的 `#name-label` 应包含描述输入的文本。

```js
const el = document.getElementById('name-label');
assert(!!el && el.innerText.length > 0);
```

你的 `#email-label` 应包含描述输入的文本。

```js
const el = document.getElementById('email-label');
assert(!!el && el.innerText.length > 0);
```

你的 `#number-label` 应包含描述输入的文本。

```js
const el = document.getElementById('number-label');
assert(!!el && el.innerText.length > 0);
```

你的 `#name-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #name-label');
assert(!!el);
```

你的 `#email-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #email-label');
assert(!!el);
```

你的 `#number-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #number-label');
assert(!!el);
```

你的 `#name` 应该有 `placeholder` 属性和值。

```js
const el = document.getElementById('name');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

你的 `#email` 应该有 `placeholder` 属性和值。

```js
const el = document.getElementById('email');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

你的 `#number` 应该有 `placeholder` 属性和值。

```js
const el = document.getElementById('number');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

你应该有一个 `id` 为 `dropdown` 的 `select` 字段。

```js
const el = document.getElementById('dropdown');
assert(!!el && el.tagName === 'SELECT');
```

你的 `#dropdown` 应该至少有两个可选（未禁用）的 `option` 元素。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert(els.length >= 2);
```

你的 `#dropdown` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #dropdown');
assert(!!el);
```

你应该有至少两个 `type` 为 `radio` 的 `input` 元素（单选按钮）。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert(els.length >= 2);
```

你应该有至少两个单选按钮是 `#survey-form` 的子元素。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert(els.length >= 2);
```

所有单选按钮都应有 `value` 属性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][value=""], input[type="radio"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

所有单选按钮都应有 `name` 属性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][name=""], input[type="radio"]:not([name])');
assert(els1.length > 0 && els2.length === 0);
```

每组单选按钮应至少有 2 个按钮。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  })
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
})

assert(groupKeys.length > 0);
```

你应该有至少两个 `type` 为 `checkbox` 的 `input` 元素（复选框），且为 `#survey-form` 的子元素。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert(els.length >= 2);
```

`#survey-form` 内所有复选框都应有 `value` 属性和值。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll('#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

你应该有至少一个 `textarea` 元素，且为 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form textarea');
assert(!!el);
```

你应该有一个 `id` 为 `submit` 的 `input` 或 `button` 元素。

```js
const el = document.getElementById('submit');
assert(!!el && (el.tagName === 'INPUT' || el.tagName === 'BUTTON'));
```

你的 `#submit` 应该是 `submit` 类型。

```js
const el = document.getElementById('submit');
assert(!!el && el.type === 'submit');
```

你的 `#submit` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #submit');
assert(!!el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">

<head>
    <meta charset="UTF-8" />
    <title>调查表单</title>
</head>

<body>

</body>

</html> 
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp 调查表单</h1>
        <p id="description" class="description text-center">
          感谢您抽出宝贵时间帮助我们改进平台
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">姓名</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="请输入您的姓名"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">邮箱</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="请输入您的邮箱"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >年龄<span class="clue">(可选)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="年龄"
          />
        </div>
        <div class="form-group">
          <p>以下哪项最能描述您当前的身份？</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>请选择当前身份</option>
            <option value="student">学生</option>
            <option value="job">全职工作</option>
            <option value="learner">全职学习者</option>
            <option value="preferNo">不愿透露</option>
            <option value="other">其他</option>
          </select>
        </div>

        <div class="form-group">
          <p>您会推荐 freeCodeCamp 给朋友吗？</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />一定会</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />可能会</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />不确定</label
          >
        </div>

        <div class="form-group">
          <p>您最喜欢 freeCodeCamp 的哪个功能？</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>请选择一个选项</option>
            <option value="challenges">挑战</option>
            <option value="projects">项目</option>
            <option value="community">社区</option>
            <option value="openSource">开源</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            您希望看到哪些方面的改进？
            <span class="clue">(可多选)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />前端项目</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />后端项目</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />数据可视化</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />挑战</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />开源社区</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter 帮助房间</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />视频</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />城市聚会</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />论坛</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />更多课程</label
          >
        </div>

        <div class="form-group">
          <p>还有其他建议或意见吗？</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="请在此输入您的建议或意见……"
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            提交
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```

