---
id: 66c07238b01053abaf812065
title: 构建一个阶乘计算器
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

阶乘是一个整数与它以下所有整数的乘积。例如，`5` 的阶乘是 `5 * 4 * 3 * 2 * 1 = 120`。在本实验中，你将创建一个阶乘计算器，从用户处获取一个数字并计算该数字的阶乘。

完成以下用户故事并通过所有测试以完成本实验。

**用户需求：**

1. 你应该声明一个变量 `num` 并为其赋一个你选择的数字。赋值的数字应在 `1` 到 `20` 之间（包含）。
1. 创建一个名为 `factorialCalculator` 的函数，该函数接收一个数字作为参数并返回该数字的阶乘。
1. 在函数内部，声明一个 `result` 变量并赋值为 `1`。使用循环，从 `1` 循环到输入的数字（包含），每次将当前数字与 `result` 相乘，并将结果赋值给 `result` 变量。你可以选择使用 `for` 循环、`while` 循环或 `do...while` 循环。
1. 你应该用 `num` 作为参数调用 `factorialCalculator` 函数，并将结果赋值给变量 `factorial`。
1. 你应该将最终输出以 `Factorial of [num] is [factorial]` 的格式存储，并赋值给变量 `resultMsg`。
1. 你应该将 `resultMsg` 的值输出到控制台。

# --hints--

你应该有一个 `num` 变量。

```js
assert.isDefined(num);
```

你应该为 `num` 变量赋值

```js
assert.isNotNull(num);
```

`num` 的值应在 `1` 到 `20` 之间。

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

`num` 的值应为数字。

```js
assert.isNumber(num);
```

你应该有一个名为 `factorialCalculator` 的函数。

```js
assert.isFunction(factorialCalculator);
```

`factorialCalculator` 函数应接收一个数字作为参数。

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

`factorialCalculator(5)` 应返回 `120`。

```js
assert.strictEqual(factorialCalculator(5), 120);
```

`factorialCalculator(7)` 应返回 `5040`。

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

你应该用 `num` 变量作为参数调用 `factorialCalculator` 函数。

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s?$/m);
```

你应该定义一个 `factorial` 变量，并将 `factorialCalculator` 函数的结果赋值给它。

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

你的 `factorialCalculator` 应产生正确的结果。

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

你的 `resultMsg` 应包含 `Factorial of [num] is [factorial]` 格式的字符串。

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

你应该将 `resultMsg` 的值输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```

