---
id: 67d936de7055982b02baf186
title: 构建一个活动 RSVP 表单
challengeType: 25
dashedName: build-an-event-rsvp
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事，并通过所有测试以完成本实验。

**用户故事：**

1. 你需要创建一个表单，包含姓名、电子邮件、参与人数、饮食偏好，以及是否携带额外嘉宾的选项。

2. 你需要有一个文本输入框，用于输入你的姓名（必填）。

3. 你需要有一个电子邮件输入框，用于输入你的电子邮件地址（必填）。表单应验证格式，确保为有效邮箱。

4. 你需要有一个数字输入框，用于输入参与人数（必填），且人数不得小于 1。

5. 你需要有一个文本输入框，用于输入你的饮食偏好，该项为可选。

6. 你需要有一个复选框，用于指示是否携带额外嘉宾，可以勾选或取消勾选。

7. 你需要有一个按钮用于提交表单，表单提交后页面不应重新加载。

8. 提交后，你应在表单下方看到一条确认消息，并显示所填写的详细信息（姓名、邮箱、参与人数、饮食偏好和是否携带额外嘉宾）。

以下是提交表单后的示例消息：

```markdown
RSVP 已提交！
姓名: John Doe
邮箱: example@example.com
参与人数: 2
饮食偏好: 无
携带额外嘉宾: 是
```

# --hints--

你应有一个 `form` 元素包含所有表单内容。

```js
const forms = document.querySelectorAll('form');
assert.equal(forms.length, 1);
```

你应有两个 `type="text"` 的 `input` 元素。

```js
const inputs = document.querySelectorAll('form input[type="text"]');
assert.equal(inputs.length, 2);
```

你应有一个 `type="email"` 的 `input` 元素。

```js
try {
  const inputs = document.querySelectorAll('form input[type="email"]');
  assert.equal(inputs.length, 1);
} catch (e) {
  console.log(e);
  throw e;
}
```

你应有一个 `type="number"` 的 `input` 元素。

```js
try {
  const inputs = document.querySelectorAll('form input[type="number"]');
  assert.equal(inputs.length, 1);
} catch (e) {
  console.log(e);
  throw e;
}
```

你应有一个 `type="checkbox"` 的 `input` 元素。

```js
try {
  const inputs = document.querySelectorAll('form input[type="checkbox"]');
  assert.equal(inputs.length, 1);
} catch (e) {
  console.log(e);
  throw e;
}
```

其中一个 `input[type="text"]` 元素应为必填项。

```js
try {
  const input = document.querySelectorAll('form input[type="text"][required]');
  assert.equal(input.length, 1);
} catch (e) {
  console.log(e);
  throw e;
}
```

`input[type="email"]` 元素应为必填项。

```js
try {
  const input = document.querySelector('form input[type="email"][required]');
  assert.exists(input);
} catch (e) {
  console.log(e);
  throw e;
}
```

`input[type="number"]` 元素应为必填项。

```js
try {
  const input = document.querySelector('form input[type="number"][required]');
  assert.exists(input);
} catch (e) {
  console.log(e);
  throw e;
}
```

更改 `input[type="text"]` 元素的值应更新组件状态。

```js
async () => {
  try {
    const abuseState = __helpers.spyOn(React, 'useState');
    const script = [...document.querySelectorAll('script')]
      .find(s => s.dataset.src === 'index.jsx')
      .innerText.replace('_React.useState', 'abuseState');

    const exports = {};
    const a = eval(script);

    const s = await __helpers.prepTestComponent(exports.EventRSVPForm);
    const inp = s.querySelector('input[type="text"]');

    await React.act(async () => {
      inp.value = 'John Doe';
      const ev = new InputEvent('change', { bubbles: true, cancelable: false });
      inp[Object.keys(inp).find(k => k.startsWith('__reactProps'))].onChange({
        ...ev,
        target: inp
      });
    });

    // 对于 abuseState.returns 中的所有状态，将有两倍的调用——一次在更改前，一次在更改后。
    const alteredStates = abuseState.returns;
    const initialStates = alteredStates.splice(
      0,
      abuseState.returns.length / 2
    );

    const stateChanged = initialStates.some((s, i) => {
      try {
        assert.deepEqual(s[0], alteredStates[i][0]);
        return false;
      } catch (e) {
        return true;
      }
    });

    abuseState.restore();
    assert.isTrue(stateChanged);
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

更改 `input[type="email"]` 元素的值应更新组件状态。

```js
async () => {
  try {
    const abuseState = __helpers.spyOn(React, 'useState');
    const script = [...document.querySelectorAll('script')]
      .find(s => s.dataset.src === 'index.jsx')
      .innerText.replace('_React.useState', 'abuseState');

    const exports = {};
    const a = eval(script);

    const s = await __helpers.prepTestComponent(exports.EventRSVPForm);
    const inp = s.querySelector('input[type="email"]');

    await React.act(async () => {
      inp.value = 'fcc@freecodecamp.org';
      const ev = new InputEvent('change', { bubbles: true, cancelable: false });
      inp[Object.keys(inp).find(k => k.startsWith('__reactProps'))].onChange({
        ...ev,
        target: inp
      });
    });

    // 对于 abuseState.returns 中的所有状态，将有两倍的调用——一次在更改前，一次在更改后。
    const alteredStates = abuseState.returns;
    const initialStates = alteredStates.splice(
      0,
      abuseState.returns.length / 2
    );

    const stateChanged = initialStates.some((s, i) => {
      try {
        assert.deepEqual(s[0], alteredStates[i][0]);
        return false;
      } catch (e) {
        return true;
      }
    });

    abuseState.restore();
    assert.isTrue(stateChanged);
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

更改 `input[type="number"]` 元素的值应更新组件状态。

```js
async () => {
  try {
    const abuseState = __helpers.spyOn(React, 'useState');
    const script = [...document.querySelectorAll('script')]
      .find(s => s.dataset.src === 'index.jsx')
      .innerText.replace('_React.useState', 'abuseState');

    const exports = {};
    const a = eval(script);

    const s = await __helpers.prepTestComponent(exports.EventRSVPForm);
    const inp = s.querySelector('input[type="number"]');

    await React.act(async () => {
      inp.value = 2;
      const ev = new InputEvent('change', { bubbles: true, cancelable: false });
      inp[Object.keys(inp).find(k => k.startsWith('__reactProps'))].onChange({
        ...ev,
        target: inp
      });
    });

    // 对于 abuseState.returns 中的所有状态，将有两倍的调用——一次在更改前，一次在更改后。
    const alteredStates = abuseState.returns;
    const initialStates = alteredStates.splice(
      0,
      abuseState.returns.length / 2
    );

    const stateChanged = initialStates.some((s, i) => {
      try {
        assert.deepEqual(s[0], alteredStates[i][0]);
        return false;
      } catch (e) {
        return true;
      }
    });

    abuseState.restore();
    assert.isTrue(stateChanged);
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

更改 `input[type="checkbox"]` 元素的值应更新组件状态。

```js
async () => {
  try {
    const abuseState = __helpers.spyOn(React, 'useState');
    const script = [...document.querySelectorAll('script')]
      .find(s => s.dataset.src === 'index.jsx')
      .innerText.replace('_React.useState', 'abuseState');

    const exports = {};
    const a = eval(script);

    const s = await __helpers.prepTestComponent(exports.EventRSVPForm);
    const inp = s.querySelector('input[type="checkbox"]');

    await React.act(async () => {
      inp.checked = true;
      const ev = new InputEvent('change', { bubbles: true, cancelable: false });
      inp[Object.keys(inp).find(k => k.startsWith('__reactProps'))].onChange({
        ...ev,
        target: inp
      });
    });

    // 对于 abuseState.returns 中的所有状态，将有两倍的调用——一次在更改前，一次在更改后。
    const alteredStates = abuseState.returns;
    const initialStates = alteredStates.splice(
      0,
      abuseState.returns.length / 2
    );

    const stateChanged = initialStates.some((s, i) => {
      try {
        assert.deepEqual(s[0], alteredStates[i][0]);
        return false;
      } catch (e) {
        return true;
      }
    });

    abuseState.restore();
    assert.isTrue(stateChanged);
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交表单不应导致页面重新加载。

```js
async () => {
  try {
    // 理想情况下，应监听 window.onload。但测试环境禁用了提交导致的页面刷新。
    // 因此，测试需要手动检查 .preventDefault 是否被调用。
    const script = [...document.querySelectorAll('script')].find(
      s => s.dataset.src === 'index.jsx'
    ).innerText;

    const exports = {};
    const a = eval(script);

    const s = await __helpers.prepTestComponent(exports.EventRSVPForm);
    const f = s.querySelector('form');

    await React.act(async () => {
      let c = 0;
      const mockEv = {
        ...new SubmitEvent('submit'),
        preventDefault: () => {
          c++;
        }
      };
      f[Object.keys(f).find(k => k.startsWith('__reactProps'))].onSubmit(
        mockEv
      );
      assert.equal(c, 1);
    });
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交后，页面上应有元素显示姓名输入框的状态。

```js
async () => {
  try {
    const inp = document.querySelector(`input[type="text"]`);
    assert.exists(inp);
    const setter = Object.getOwnPropertyDescriptor(
      window.HTMLInputElement.prototype,
      'value'
    ).set;
    await React.act(async () => {
      // 为了提交表单，所有必填项都赋值。
      await adjustAllRequired();
      setter.call(inp, 'John Doe');
      inp.dispatchEvent(new Event('input', { bubbles: true }));

      submitForm();
    });

    const nonInputText = getInnerTextExcept('input,script');
    assert.include(nonInputText, 'John Doe');
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交后，页面上应有元素显示邮箱输入框的状态。

```js
async () => {
  try {
    const inp = document.querySelector(`input[type="email"]`);
    assert.exists(inp);
    const setter = Object.getOwnPropertyDescriptor(
      window.HTMLInputElement.prototype,
      'value'
    ).set;
    await React.act(async () => {
      // 为了提交表单，所有必填项都赋值。
      await adjustAllRequired();
      setter.call(inp, 'fcc@freecodecamp.org');
      inp.dispatchEvent(new Event('input', { bubbles: true }));

      submitForm();
    });

    const nonInputText = getInnerTextExcept('input,script');
    assert.include(nonInputText, 'fcc@freecodecamp.org');
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交后，页面上应有元素显示参与人数输入框的状态。

```js
async () => {
  try {
    const inp = document.querySelector(`input[type="number"]`);
    assert.exists(inp);
    const setter = Object.getOwnPropertyDescriptor(
      window.HTMLInputElement.prototype,
      'value'
    ).set;
    await React.act(async () => {
      // 为了提交表单，所有必填项都赋值。
      await adjustAllRequired();
      setter.call(inp, 2);
      inp.dispatchEvent(new Event('input', { bubbles: true }));

      submitForm();
    });

    const nonInputText = getInnerTextExcept('input,script');
    assert.include(nonInputText, '2');
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交后，页面上应有元素显示饮食偏好输入框的状态。

```js
async () => {
  try {
    const inp = document.querySelector(`input[type="text"]:not(:required)`);
    assert.exists(inp);
    const setter = Object.getOwnPropertyDescriptor(
      window.HTMLInputElement.prototype,
      'value'
    ).set;
    await React.act(async () => {
      // 为了提交表单，所有必填项都赋值。
      await adjustAllRequired();
      setter.call(inp, 'diet');
      inp.dispatchEvent(new Event('input', { bubbles: true }));

      submitForm();
    });

    const nonInputText = getInnerTextExcept('input,script');
    assert.include(nonInputText, 'diet');
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

提交后，页面上应有元素显示额外嘉宾复选框的状态。

```js
async () => {
  try {
    let inp = document.querySelector(`input[type="checkbox"]`);
    assert.exists(inp);
    await React.act(async () => {
      // 为了提交表单，所有必填项都赋值。
      await adjustAllRequired();
      const setter = Object.getOwnPropertyDescriptor(
        window.HTMLInputElement.prototype,
        'checked'
      ).set;
      setter.call(inp, false);
      inp.dispatchEvent(new Event('click', { bubbles: true }));

      submitForm();
    });

    // 先测试未勾选，再测试勾选，比较两次提交的内容是否不同。
    const stateWithFalse = getInnerTextExcept('input,script');

    inp = document.querySelector(`input[type="checkbox"]`);
    await React.act(async () => {
      const setter = Object.getOwnPropertyDescriptor(
        window.HTMLInputElement.prototype,
        'checked'
      ).set;
      setter.call(inp, true);
      inp.dispatchEvent(new Event('click', { bubbles: true }));

      submitForm();
    });

    const stateWithTrue = getInnerTextExcept('input,script');

    assert.notEqual(stateWithFalse, stateWithTrue);
  } catch (e) {
    console.log(e);
    throw e;
  }
};
```

# --before-all--

```js
// 提交按钮可以是表单内唯一的按钮，也可以是 `button[type="submit"]`。
function submitForm() {
  const maybeSubmit = document.querySelectorAll("form button");
  if (maybeSubmit.length === 1) {
    return maybeSubmit[0].click();
  }

  const submitBtn = document.querySelector("button[type='submit']");
  if (submitBtn) {
    return submitBtn.click();
  }
}

async function adjustAllRequired() {
  const inps = document.querySelectorAll(
    'input[required], select[required], textarea[required]'
  );
  inps.forEach(inp => {
    let setter = Object.getOwnPropertyDescriptor(
      window.HTMLInputElement.prototype,
      'value'
    ).set;
    switch (inp.type) {
      case 'text':
        setter.call(inp, 'required text');
        inp.dispatchEvent(new Event('change', { bubbles: true }));
        break;
      case 'email':
        setter.call(inp, 'required-email@freecodecamp.org');
        inp.dispatchEvent(new Event('change', { bubbles: true }));
        break;
      case 'number':
        setter.call(inp, 1);
        inp.dispatchEvent(new Event('change', { bubbles: true }));
        break;
      case 'checkbox':
        setter = Object.getOwnPropertyDescriptor(
          window.HTMLInputElement.prototype,
          'checked'
        ).set;
        setter.call(inp, true);
        inp.dispatchEvent(new Event('click', { bubbles: true }));
        break;
    }
    inp.dispatchEvent(new Event('change', { bubbles: true }));
  });
}

function getInnerTextExcept(removingSelector) {
  const body = document.body.cloneNode(true);

  const squareElements = body.querySelectorAll(removingSelector);
  squareElements.forEach(element => {
    element.parentNode.removeChild(element);
  });

  return body.innerText;
}
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>活动 RSVP</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { EventRSVPForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <EventRSVPForm />
      );
    </script>
  </body>
</html>
```

```css

```

```jsx
const { useState } = React;

export function EventRSVPForm() {}
```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>活动 RSVP</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { EventRSVPForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <EventRSVPForm />
      );
    </script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  margin: 0;
  font-family: 'Helvetica Neue', Arial, sans-serif;
  background-color: #e0e0e0;
}

.form-container {
  background-color: #ffffff;
  padding: 40px;
  border-radius: 12px;
  box-shadow: 0 8px 16px rgba(0, 0, 0, 0.2);
  width: 100%;
  max-width: 450px;
  text-align: center;
}

h2 {
  margin-bottom: 20px;
  font-size: 28px;
  color: #333;
  font-weight: bold;
}

form {
  display: flex;
  flex-direction: column;
}

.form-group {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 20px;
}

.form-group label {
  flex: 1;
  text-align: left;
  margin-right: 15px;
  font-size: 15px;
  color: #555;
}

.form-group input[type='text'],
.form-group input[type='email'],
.form-group input[type='number'] {
  flex: 2;
  padding: 10px 12px;
  border: 1px solid #ccc;
  border-radius: 8px;
  width: 100%;
  font-size: 15px;
  transition: border-color 0.3s ease;
}

.form-group input:focus {
  border-color: #4caf4f4f;
}

.form-group input[type='checkbox'] {
  margin-left: 10px;
}

button {
  background-color: #5957e4;
  color: white;
  border: none;
  padding: 12px 25px;
  border-radius: 8px;
  cursor: pointer;
  font-size: 16px;
  font-weight: bold;
  margin-top: 15px;
  transition: background-color 0.3s ease;
}

button:hover {
  background-color: #413fb3;
}

.submitted-message {
  margin-top: 30px;
  text-align: left;
}

.submitted-message h3 {
  color: #5957e4;
  margin-bottom: 10px;
  font-size: 22px;
}

@media (max-width: 500px) {
  .form-container {
    width: 90%;
    padding: 20px;
  }

  .form-group {
    flex-direction: column;
    align-items: flex-start;
  }

  .form-group label {
    text-align: left;
    margin-bottom: 5px;
  }

  .form-group input[type='text'],
  .form-group input[type='email'],
  .form-group input[type='number'] {
    width: 100%;
  }
}
```

```jsx
const { useState } = React;

export function EventRSVPForm() {
  const [formData, setFormData] = useState({
    name: '',
    email: '',
    attendees: '',
    dietaryPreferences: '',
    bringingOthers: false
  });

  const [submitted, setSubmitted] = useState(false);

  function handleChange(event) {
    const { name, value, type, checked } = event.target;
    setFormData({
      ...formData,
      [name]: type === 'checkbox' ? checked : value
    });
  }

  function handleSubmit(event) {
    event.preventDefault();
    setSubmitted(true);
  }

  return (
    <div className='form-container'>
      <h2>活动 RSVP 表单</h2>
      <form onSubmit={handleSubmit}>
        <div className='form-group'>
          <label>
            姓名：
            <input
              type='text'
              name='name'
              value={formData.name}
              onChange={handleChange}
              placeholder='请输入姓名'
              required
            />
          </label>
        </div>
        <div className='form-group'>
          <label>
            邮箱：
            <input
              type='email'
              name='email'
              value={formData.email}
              onChange={handleChange}
              placeholder='请输入邮箱'
              required
            />
          </label>
        </div>
        <div className='form-group'>
          <label>
            参与人数：
            <input
              type='number'
              name='attendees'
              value={formData.attendees || ''}
              onChange={handleChange}
              min='1'
              placeholder='参与人数'
              required
            />
          </label>
        </div>
        <div className='form-group'>
          <label>
            饮食偏好：
            <input
              type='text'
              name='dietaryPreferences'
              value={formData.dietaryPreferences}
              onChange={handleChange}
              placeholder='饮食偏好（可选）'
            />
          </label>
        </div>
        <div className='form-group'>
          <label>
            是否携带额外嘉宾？
            <input
              type='checkbox'
              name='bringingOthers'
              checked={formData.bringingOthers}
              onChange={handleChange}
            />
          </label>
        </div>
        <button type='submit'>提交 RSVP</button>
      </form>
      {submitted && (
        <div className='submitted-message'>
          <h3>RSVP 已提交！</h3>
          <p>
            <strong>姓名:</strong> {formData.name}
          </p>
          <p>
            <strong>邮箱:</strong> {formData.email}
          </p>
          <p>
            <strong>参与人数:</strong> {formData.attendees}
          </p>
          <p>
            <strong>饮食偏好:</strong>{' '}
            {formData.dietaryPreferences || '无'}
          </p>
          <p>
            <strong>携带额外嘉宾:</strong>{' '}
            {formData.bringingOthers ? '是' : '否'}
          </p>
        </div>
      )}
    </div>
  );
}
```

