---
id: 6762ec275cef87635acc4fe3
title: 构建一个鼓机
challengeType: 25
dashedName: build-drum-machine
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成本实验。

1. 你应该有一个 `id` 为 `drum-machine` 的 `div` 元素，包含所有其他元素。

2. 在 `#drum-machine` 元素内部，你应该有另一个 `id` 为 `pad-bank` 的 `div` 元素。

3. 在 `#drum-machine` 元素内部，你应该有一个 `id` 为 `display` 的 `p` 元素。

4. 在你的 `#pad-bank` 元素内，应该有九个可点击的鼓垫元素，每个元素都有 `drum-pad` 类、一个描述其音频片段的唯一 `id`，以及与下列键盘按键之一对应的文本内容：`Q`、`W`、`E`、`A`、`S`、`D`、`Z`、`X`、`C`。鼓垫必须按此顺序排列。

5. 每个 `.drum-pad` 内应有一个 `audio` 元素，具有 `clip` 类、指向音频片段的 `src` 属性，以及与其父 `.drum-pad` 元素文本内容对应的 `id`（如 `id="Q"`、`id="W"`、`id="E"` 等）。

6. 当你点击 `.drum-pad` 元素时，其子 audio 元素中的音频片段应被触发播放。

7. 当你按下与每个 `.drum-pad` 关联的触发键时，其子 audio 元素中的音频片段应被触发播放（例如按下 `Q` 键应触发包含 `Q` 的鼓垫，按下 `W` 键应触发包含 `W` 的鼓垫，依此类推）。

8. 当 `.drum-pad` 被触发时，你应在 `#display` 元素的文本内容中显示描述相关音频片段的字符串（每个字符串必须唯一）。

你可以使用以下音频样本，音频地址为 `https://cdn.freecodecamp.org/curriculum/drum/<fileName>`，其中 `<fileName>` 如下表所示：

| 鼓名称         | `fileName`         |
| -------------- | ----------------- |
| Heater 1       | `Heater-1.mp3`    |
| Heater 2       | `Heater-2.mp3`    |
| Heater 3       | `Heater-3.mp3`    |
| Heater 4       | `Heater-4_1.mp3`  |
| Clap           | `Heater-6.mp3`    |
| Open-HH        | `Dsc_Oh.mp3`      |
| Kick-n'-Hat    | `Kick_n_Hat.mp3`  |
| Kick           | `RP4_KICK_1.mp3`  |
| Closed-HH      | `Cev_H2.mp3`      |

# --hints--

你应该有一个 `id` 为 `drum-machine` 的 `div` 元素，包含所有其他元素。

```js
assert.isNotNull(document.querySelector('div#drum-machine'));
```

在 `#drum-machine` 元素内部，你应该有另一个 `id` 为 `pad-bank` 的 `div` 元素。

```js
assert.isNotNull(document.querySelector('#drum-machine div#pad-bank'));
```

在 `#drum-machine` 元素内部，你应该有一个 `id` 为 `display` 的 `p` 元素。

```js
assert.isNotNull(document.querySelector('#drum-machine p#display'));
```

在你的 `#pad-bank` 元素内，应该有九个带有 `drum-pad` 类的可点击鼓垫元素。

```js
assert.lengthOf(document.querySelectorAll('#pad-bank .drum-pad'), 9);
```

每个 `.drum-pad` 的 `innerText` 应按顺序为：`Q`、`W`、`E`、`A`、`S`、`D`、`Z`、`X`、`C`。

```js
const pads = document.querySelectorAll('.drum-pad');
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
for (let i = 0; i < 9; i++) {
  assert.equal(pads[i].innerText.trim(), letters[i]);
}
```

每个 `.drum-pad` 应有一个 `audio` 元素，具有 `clip` 类、指向音频片段的 `src` 属性，以及与其父 `.drum-pad` 元素文本内容对应的 `id`（如 `id="Q"`、`id="W"`、`id="E"` 等）。

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  assert.equal(audios[i].id, letters[i]);
  assert.isTrue(audios[i].classList.contains('clip'));
  assert.isTrue(audios[i].hasAttribute('src'));
}
```

当你点击 `.drum-pad` 元素时，其子 audio 元素中的音频片段应被触发播放。

```js
const pads = document.querySelectorAll('.drum-pad');
assert.isNotEmpty(pads);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  assert.isFalse(audio.paused);
  audio.pause();
});
```

当你按下 Q、W、E、A、S、D、Z、X、C 这些键盘按键时，对应的 `audio` 元素应播放相应的声音。

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  const audio = audios[i];
  document.dispatchEvent(
    new KeyboardEvent('keydown', { key: letters[i].toLowerCase() })
  );
  assert.isFalse(audio.paused);
  audio.pause();
}
```

当 `.drum-pad` 被触发时，你应在 `#display` 元素的文本内容中显示描述相关音频片段的字符串（每个字符串必须唯一）。

```js
const collection = new Set();
const disp = document.querySelector('#display');
const pads = document.querySelectorAll('.drum-pad');
assert.lengthOf(pads, 9);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  collection.add(disp.innerText.trim());
  audio.pause();
});
assert.lengthOf(collection, 9);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="zh">
  <head>
    <meta charset="utf-8" />
    <title>鼓机</title>
  </head>
  <body></body>
</html>
```

```css

```

```js

```

# --solutions--

```html
<!doctype html>
<html lang="zh">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>鼓机</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="inner-container" id="drum-machine">
      <div id="pad-bank"></div>

      <div class="controls-container">
        <div class="control power-control">
          <p>电源：<span id="power-status">开</span></p>
          <div class="select">
            <div class="inner"></div>
          </div>
        </div>

        <!-- 鼓垫名称和音量显示 -->
        <p id="display"></p>

        <div class="volume-slider">
          <input type="range" min="0" max="1" step="0.01" value="0.3" />
        </div>

        <div class="control bank-control">
          <p>音色组：<span id="current-bank">Heater Kit</span></p>
          <div class="select">
            <div class="inner"></div>
          </div>
        </div>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
/* 通用重置 */
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Arial', sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
  background-color: #282c34;
  color: white;
}

.inner-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  background: #3b3b5f;
  border-radius: 8px;
  padding: 20px;
  width: 300px;
}

/* 鼓垫样式 */
#pad-bank {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 10px;
  margin-bottom: 15px;
}

.drum-pad {
  display: flex;
  align-items: center;
  justify-content: center;
  width: 70px;
  height: 70px;
  background: #444;
  border-radius: 5px;
  font-size: 1.2em;
  cursor: pointer;
}

.drum-pad.active {
  background: #ffaa33;
}

.controls-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  width: 100%;
}

.control {
  display: flex;
  justify-content: space-between;
  width: 100%;
  margin: 10px 0;
}

.select {
  width: 50px;
  height: 20px;
  background: #555;
  border-radius: 10px;
  cursor: pointer;
  position: relative;
}

.select .inner {
  position: absolute;
  width: 18px;
  height: 18px;
  background: #ffaa33;
  border-radius: 50%;
  top: 1px;
  left: calc(100% - 18px);
}

#display {
  margin: 15px 0;
  text-align: center;
  font-size: 1em;
  padding: 20px;
  background: #444;
  border-radius: 5px;
  width: 100%;
}

.volume-slider input[type='range'] {
  width: 100%;
  -webkit-appearance: none;
  appearance: none;
  height: 4px;
  background: #ffaa33;
  border-radius: 2px;
  margin-bottom: 10px;
}

.volume-slider input[type='range']::-webkit-slider-thumb {
  -webkit-appearance: none;
  width: 12px;
  height: 12px;
  background: #fff;
  border-radius: 50%;
  cursor: pointer;
}
```

```js
const projectName = 'drum-machine';

const bankOne = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Heater-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Heater-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Heater-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Heater-4',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-4_1.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Clap',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-6.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dsc_Oh.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: "Kick-n'-Hat",
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Kick_n_Hat.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/RP4_KICK_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Cev_H2.mp3'
  }
];

const bankTwo = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Chord-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Chord-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Chord-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Shaker',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Give_us_a_light.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dry_Ohh.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Bld_H1.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: 'Punchy-Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/punchy_kick_1.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Side-Stick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/side_stick_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Snare',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Brk_Snr.mp3'
  }
];

const state = {
  power: true,
  display: ' ',
  currentPadBank: bankOne,
  currentPadBankId: 'Heater Kit',
  sliderVal: 0.3
};

const display = document.getElementById('display');

const setDisplay = text => {
  state.display = text;
  display.innerText = text;
};

const adjustVolume = e => {
  if (state.power) {
    state.sliderVal = e.target.value;
    setDisplay(`音量: ${Math.round(e.target.value * 100)}`);
    setTimeout(() => setDisplay(' '), 1000);
    document.querySelectorAll('.clip').forEach(sound => {
      sound.volume = state.sliderVal;
    });
  }
};

const activatePad = pad => {
  pad.classList.add('active');
  setTimeout(() => pad.classList.remove('active'), 100);
};

const playSound = (id, name) => {
  if (!state.power) return;

  const sound = document.getElementById(id);
  sound.currentTime = 0;
  sound.play();
  activatePad(sound.parentNode);
  setDisplay(name);
};

const renderPadBank = () => {
  const padBank = document.querySelector('#pad-bank');
  padBank.innerHTML = '';

  state.currentPadBank.forEach(drum => {
    const pad = document.createElement('div');
    pad.classList.add('drum-pad');
    pad.id = drum.id;
    pad.innerText = drum.keyTrigger;

    const audio = document.createElement('audio');
    audio.classList.add('clip');
    audio.id = drum.keyTrigger;
    audio.src = drum.url;
    pad.appendChild(audio);

    pad.addEventListener('click', () =>
      playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '))
    );
    document.addEventListener('keydown', e => {
      if (e.key === drum.keyTrigger.toLowerCase())
        playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '));
    });

    padBank.appendChild(pad);
  });
};

const powerControl = () => {
  state.power = !state.power;
  document.getElementById('power-status').textContent = state.power
    ? '开'
    : '关';

  setDisplay('');

  const powerSlider = document.querySelector('.power-control .select .inner');
  powerSlider.style.left = state.power ? 'calc(100% - 18px)' : '0';
};

const selectBank = () => {
  if (!state.power) return;
  setDisplay(' ');
  state.currentPadBank =
    state.currentPadBankId === 'Heater Kit' ? bankTwo : bankOne;
  state.currentPadBankId =
    state.currentPadBankId === 'Heater Kit' ? 'Smooth Piano Kit' : 'Heater Kit';

  renderPadBank();
  document.getElementById('current-bank').textContent =
    state.currentPadBankId === 'Heater Kit' ? 'Heater Kit' : 'Smooth Piano Kit';
  const bankControl = document.querySelector('.bank-control .select .inner');
  bankControl.style.left =
    state.currentPadBankId === 'Heater Kit' ? 'calc(100% - 18px)' : '0';
};

document
  .querySelector('.power-control')
  .addEventListener('click', powerControl);
document.querySelector('.bank-control').addEventListener('click', selectBank);
document
  .querySelector('.volume-slider input')
  .addEventListener('input', adjustVolume);

renderPadBank();
```

