---
id: 66e7ee20b79186306fc12da5
title: 构建一组足球队卡片
challengeType: 25
dashedName: lab-football-team-cards
demoType: onClick
---

# --description--

在本实验中，你将构建一组足球队卡片。用户可以通过下拉菜单，根据球员的位置筛选不同的球员。

请完成以下用户故事，并通过所有测试以完成本实验。

**用户需求：**

1. 你需要创建一个名为 `footballTeam` 的对象，包含以下属性：`team`、`year`、`headCoach`、`players`。
1. `team` 属性应为球队名称的字符串。
1. `year` 属性应为年份的数字。
1. `headCoach` 属性应为主教练姓名的字符串。
1. `players` 属性应为一个包含至少四个元素的数组。
1. `players` 数组中的每个元素都应为一个对象，包含 `name`、`position`、`isCaptain` 属性。
1. `name` 属性应为球员姓名的字符串。
1. `position` 属性的值应为以下之一：`"forward"`（前锋）、`"midfielder"`（中场）、`"defender"`（后卫）、`"goalkeeper"`（守门员）。
1. `isCaptain` 属性应为布尔值。其中一名球员的 `isCaptain` 属性应为 `true`。
1. 你需要在页面上展示 `coach`、`team` 和 `year` 的值。这些值应显示在 `id` 分别为 `head-coach`、`team` 和 `year` 的 HTML 元素中。
1. 你需要在 `#player-cards` 元素中展示球员数据，每位球员应显示在一个带有 `player-card` 类的 `div` 中，内部包含一个 `h2`，显示球员姓名，如果是队长则加上 `(Captain)`，以及一个 `p`，内容为 `Position:` 和球员的位置。

  ```html
  <div class="player-card">
    <h2>Sergio Batista</h2>
    <p>Position: midfielder</p>
  </div>
  <div class="player-card">
    <h2>(Captain) Diego Maradona</h2>
    <p>Position: midfielder</p>
  </div>
  ```

1. 当下拉菜单选择某个位置时，只显示该位置的球员。如果选择 `"All Players"`，则显示所有球员。

# --hints--

你应该有一个 `footballTeam` 变量。

```js
assert.exists(footballTeam);
```

`footballTeam` 变量应为一个包含 `team`、`year`、`headCoach` 和 `players` 四个属性的对象。

```js
assert.isObject(footballTeam);
assert.containsAllKeys(footballTeam, ['team', 'year', 'headCoach', 'players']);
```

`team` 属性应为字符串。

```js
assert.isString(footballTeam.team);
```

`year` 属性应为数字。

```js
assert.isNumber(footballTeam.year);
```

`headCoach` 属性应为字符串。

```js
assert.isString(footballTeam.headCoach);
```

`players` 属性应为一个包含至少四个对象的数组，每个对象应包含 `name`、`position`、`isCaptain` 键。

```js
assert.isArray(footballTeam.players);
assert.isAtLeast(footballTeam.players.length, 4);
footballTeam.players.forEach(player => {
  assert.isObject(player);
  assert.containsAllKeys(player, ['name', 'position', 'isCaptain']);
})
```

`name` 属性应为字符串。

```js
footballTeam.players.forEach(({name}) => assert.isString(name));
```

`position` 属性的值应为 `"forward"`、`"midfielder"`、`"defender"` 或 `"goalkeeper"` 之一。

```js
footballTeam.players.forEach(({position}) => {
  assert.isString(position);
  assert.oneOf(position, ["forward", "midfielder", "defender", "goalkeeper"]);
});
```

`isCaptain` 属性应为布尔值，并且只有一名队长。

```js
footballTeam.players.forEach(({isCaptain}) => assert.isBoolean(isCaptain));
const listOfCaptains = footballTeam.players.filter(({isCaptain}) => isCaptain);
assert.lengthOf(listOfCaptains, 1);
```

你应将 `footballTeam` 对象中的 `coach`、`team` 和 `year` 的值显示在 `id` 分别为 `head-coach`、`team` 和 `year` 的 HTML 元素中。

```js
const teamElement = document.querySelector('.team-stats #team');
const yearElement = document.querySelector('.team-stats #year');
const headCoachElement = document.querySelector('.team-stats #head-coach');
assert.equal(teamElement?.innerText.trim(), footballTeam.team);
assert.equal(yearElement?.innerText.trim(), footballTeam.year);
assert.equal(headCoachElement?.innerText.trim(), footballTeam.headCoach);
```

当选择 `All Players` 选项时，所有球员应显示在 `#player-cards` 中。

```js
const select = document.querySelector('#players')
select.value = 'all';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
  name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
  position: el.querySelector('p').innerText.replace('Position:', '').trim(),
  isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, footballTeam.players);
```

当选择 `Position Forward` 选项时，只显示前锋球员。

```js
const forwards = footballTeam.players.filter(({position}) => position === 'forward')
const select = document.querySelector('#players')
select.value = 'forward';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
  name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
  position: el.querySelector('p').innerText.replace('Position:', '').trim(),
  isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, forwards);
```

当选择 `Position Midfielder` 选项时，只显示中场球员。

```js
const midfielders = footballTeam.players.filter(({position}) => position === 'midfielder')
const select = document.querySelector('#players')
select.value = 'midfielder';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
  name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
  position: el.querySelector('p').innerText.replace('Position:', '').trim(),
  isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, midfielders);
```

当选择 `Position Defender` 选项时，只显示后卫球员。

```js
const defenders = footballTeam.players.filter(({position}) => position === 'defender')
const select = document.querySelector('#players')
select.value = 'defender';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
  name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
  position: el.querySelector('p').innerText.replace('Position:', '').trim(),
  isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, defenders);
```

当选择 `Position Goalkeeper` 选项时，只显示守门员球员。

```js
const goalkeepers = footballTeam.players.filter(({position}) => position === 'goalkeeper')
const select = document.querySelector('#players')
select.value = 'goalkeeper';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
  name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
  position: el.querySelector('p').innerText.replace('Position:', '').trim(),
  isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, goalkeepers);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  <meta charset="UTF-8" />
  <meta http-equiv="X-UA-Compatible" content="IE=edge" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>
    Build a Set of Football Team Cards
  </title>
  <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
  <h1 class="title">Team stats</h1>
  <main>
    <div class="team-stats">
    <p>Team: <span id="team"></span></p>
    <p>Year: <span id="year"></span></p>
    <p>Head coach: <span id="head-coach"></span></p>
    </div>
    <label class="options-label" for="players">Filter Teammates:</label>
    <select name="players" id="players">
    <option value="all">All Players</option>
    <option value="forward">Position Forward</option>
    <option value="midfielder">Position Midfielder</option>
    <option value="defender">Position Defender</option>
    <option value="goalkeeper">Position Goalkeeper</option>
    </select>
    <div class="cards" id="player-cards"></div>
  </main>
  <footer>&copy; freeCodeCamp</footer>
  <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
  flex-direction: column;
  }
}

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  <meta charset="UTF-8" />
  <meta http-equiv="X-UA-Compatible" content="IE=edge" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>
    Build a Set of Football Team Cards
  </title>
  <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
  <h1 class="title">Team stats</h1>
  <main>
    <div class="team-stats">
    <p>Team: <span id="team"></span></p>
    <p>Year: <span id="year"></span></p>
    <p>Head coach: <span id="head-coach"></span></p>
    </div>
    <label class="options-label" for="players">Filter Teammates:</label>
    <select name="players" id="players">
    <option value="all">All Players</option>
    <option value="forward">Position Forward</option>
    <option value="midfielder">Position Midfielder</option>
    <option value="defender">Position Defender</option>
    <option value="goalkeeper">Position Goalkeeper</option>
    </select>
    <div class="cards" id="player-cards"></div>
  </main>
  <footer>&copy; freeCodeCamp</footer>
  <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
  flex-direction: column;
  }
}

```

```js
const teamName = document.getElementById("team");
const worldCupYear = document.getElementById("year");
const headCoach = document.getElementById("head-coach");
const playerCards = document.getElementById("player-cards");
const playersDropdownList = document.getElementById("players");
const footballTeam = {
  team: "Argentina",
  year: 1986,
  headCoach: "Carlos Bilardo",
  players: [
  {
    name: "Sergio Almirón",
    position: "forward",
    isCaptain: false,
  },
  {
    name: "Sergio Batista",
    position: "midfielder",
    isCaptain: false,
  },
   {
    name: "Ricardo Bochini",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "Claudio Borghi",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "José Luis Brown",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Daniel Passarella",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Jorge Burruchaga",
    position: "forward",
    isCaptain: false,
  },
  {
    name: "Néstor Clausen",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "José Luis Cuciuffo",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Diego Maradona",
    position: "midfielder",
    isCaptain: true,
  },
  {
    name: "Jorge Valdano",
    position: "forward",
    isCaptain: false,
  },
  {
    name: "Héctor Enrique",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "Oscar Garré",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Ricardo Giusti",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "Luis Islas",
    position: "goalkeeper",
    isCaptain: false,
  },
  {
    name: "Julio Olarticoechea",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Pedro Pasculli",
    position: "forward",
    isCaptain: false,
  },
  {
    name: "Nery Pumpido",
    position: "goalkeeper",
    isCaptain: false,
  },
  {
    name: "Oscar Ruggeri",
    position: "defender",
    isCaptain: false,
  },
  {
    name: "Carlos Tapia",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "Marcelo Trobbiani",
    position: "midfielder",
    isCaptain: false,
  },
  {
    name: "Héctor Zelada",
    position: "goalkeeper",
    isCaptain: false,
  },
  ],
};

const { sport, team, year, players } = footballTeam;
const coachName = footballTeam.headCoach;

teamName.textContent = team;
worldCupYear.textContent = year;
headCoach.textContent = coachName;

const setPlayerCards = (arr = players) => {
  playerCards.innerHTML += arr
  .map(
    ({ name, position, isCaptain }) =>
    `
    <div class="player-card">
      <h2>${isCaptain ? "(Captain)" : ""} ${name}</h2>
      <p>Position: ${position}</p>
    </div>
    `
  )
  .join("");
};

playersDropdownList.addEventListener("change", (e) => {
  playerCards.innerHTML = "";

  switch (e.target.value) {
  case "forward":
    setPlayerCards(players.filter((player) => player.position === "forward"));
    break;
  case "midfielder":
    setPlayerCards(
    players.filter((player) => player.position === "midfielder")
    );
    break;
  case "defender":
    setPlayerCards(
    players.filter((player) => player.position === "defender")
    );
    break;
  case "goalkeeper":
    setPlayerCards(
    players.filter((player) => player.position === "goalkeeper")
    );
    break;
  default:
    setPlayerCards();
  }
});

setPlayerCards();

```

